% Refutation found. Thanks to Tanya!
% SZS status Theorem for test4
% SZS output start Proof for test4
1. ! [X0] : (main_query_0(X0) <=> main_query_1(X0)) [input]
2. ~! [X0] : (main_query_0(X0) <=> main_query_1(X0)) [negated conjecture 1]
3. ! [X0] : (main_query_0(X0) <=> ? [X1,X2] : (X0 = X1 & $true & tab(X0,X1,X2))) [input]
5. ! [X0] : (main_query_1(X0) <=> ? [X1,X2,X6,X7,X8] : ((! [X3,X4,X5] : (b(X3,X4,X5) <=> tab(X3,X4,X5)) => b(X6,X7,X8)) & tab(X0,X1,X2) & X0 = X1)) [input]
6. ! [X0] : (main_query_0(X0) <=> ? [X1,X2] : (X0 = X1 & tab(X0,X1,X2))) [true and false elimination 3]
8. ! [X0] : (main_query_1(X0) <=> ? [X1,X2,X3,X4,X5] : ((! [X6,X7,X8] : (b(X6,X7,X8) <=> tab(X6,X7,X8)) => b(X3,X4,X5)) & tab(X0,X1,X2) & X0 = X1)) [rectify 5]
9. ? [X0] : (main_query_0(X0) <~> main_query_1(X0)) [ennf transformation 2]
10. ! [X0] : (main_query_1(X0) <=> ? [X1,X2,X3,X4,X5] : ((b(X3,X4,X5) | ? [X6,X7,X8] : (b(X6,X7,X8) <~> tab(X6,X7,X8))) & tab(X0,X1,X2) & X0 = X1)) [ennf transformation 8]
11. ? [X0] : ((~main_query_1(X0) | ~main_query_0(X0)) & (main_query_1(X0) | main_query_0(X0))) [nnf transformation 9]
12. ? [X0] : ((~main_query_1(X0) | ~main_query_0(X0)) & (main_query_1(X0) | main_query_0(X0))) => ((~main_query_1(sK0) | ~main_query_0(sK0)) & (main_query_1(sK0) | main_query_0(sK0))) [choice axiom]
13. (~main_query_1(sK0) | ~main_query_0(sK0)) & (main_query_1(sK0) | main_query_0(sK0)) [skolemisation 11,12]
14. ! [X0] : ((main_query_0(X0) | ! [X1,X2] : (X0 != X1 | ~tab(X0,X1,X2))) & (? [X1,X2] : (X0 = X1 & tab(X0,X1,X2)) | ~main_query_0(X0))) [nnf transformation 6]
15. ! [X0] : ((main_query_0(X0) | ! [X1,X2] : (X0 != X1 | ~tab(X0,X1,X2))) & (? [X3,X4] : (X0 = X3 & tab(X0,X3,X4)) | ~main_query_0(X0))) [rectify 14]
16. ! [X0] : (? [X3,X4] : (X0 = X3 & tab(X0,X3,X4)) => (sK1(X0) = X0 & tab(X0,sK1(X0),sK2(X0)))) [choice axiom]
17. ! [X0] : ((main_query_0(X0) | ! [X1,X2] : (X0 != X1 | ~tab(X0,X1,X2))) & ((sK1(X0) = X0 & tab(X0,sK1(X0),sK2(X0))) | ~main_query_0(X0))) [skolemisation 15,16]
18. ! [X0] : ((main_query_1(X0) | ! [X1,X2,X3,X4,X5] : ((~b(X3,X4,X5) & ! [X6,X7,X8] : ((b(X6,X7,X8) | ~tab(X6,X7,X8)) & (tab(X6,X7,X8) | ~b(X6,X7,X8)))) | ~tab(X0,X1,X2) | X0 != X1)) & (? [X1,X2,X3,X4,X5] : ((b(X3,X4,X5) | ? [X6,X7,X8] : ((~tab(X6,X7,X8) | ~b(X6,X7,X8)) & (tab(X6,X7,X8) | b(X6,X7,X8)))) & tab(X0,X1,X2) & X0 = X1) | ~main_query_1(X0))) [nnf transformation 10]
19. ! [X0] : ((main_query_1(X0) | ! [X1,X2,X3,X4,X5] : ((~b(X3,X4,X5) & ! [X6,X7,X8] : ((b(X6,X7,X8) | ~tab(X6,X7,X8)) & (tab(X6,X7,X8) | ~b(X6,X7,X8)))) | ~tab(X0,X1,X2) | X0 != X1)) & (? [X9,X10,X11,X12,X13] : ((b(X11,X12,X13) | ? [X14,X15,X16] : ((~tab(X14,X15,X16) | ~b(X14,X15,X16)) & (tab(X14,X15,X16) | b(X14,X15,X16)))) & tab(X0,X9,X10) & X0 = X9) | ~main_query_1(X0))) [rectify 18]
20. ! [X0] : (? [X9,X10,X11,X12,X13] : ((b(X11,X12,X13) | ? [X14,X15,X16] : ((~tab(X14,X15,X16) | ~b(X14,X15,X16)) & (tab(X14,X15,X16) | b(X14,X15,X16)))) & tab(X0,X9,X10) & X0 = X9) => ((b(sK5(X0),sK6(X0),sK7(X0)) | ? [X14,X15,X16] : ((~tab(X14,X15,X16) | ~b(X14,X15,X16)) & (tab(X14,X15,X16) | b(X14,X15,X16)))) & tab(X0,sK3(X0),sK4(X0)) & sK3(X0) = X0)) [choice axiom]
21. ? [X14,X15,X16] : ((~tab(X14,X15,X16) | ~b(X14,X15,X16)) & (tab(X14,X15,X16) | b(X14,X15,X16))) => ((~tab(sK8,sK9,sK10) | ~b(sK8,sK9,sK10)) & (tab(sK8,sK9,sK10) | b(sK8,sK9,sK10))) [choice axiom]
22. ! [X0] : ((main_query_1(X0) | ! [X1,X2,X3,X4,X5] : ((~b(X3,X4,X5) & ! [X6,X7,X8] : ((b(X6,X7,X8) | ~tab(X6,X7,X8)) & (tab(X6,X7,X8) | ~b(X6,X7,X8)))) | ~tab(X0,X1,X2) | X0 != X1)) & (((b(sK5(X0),sK6(X0),sK7(X0)) | ((~tab(sK8,sK9,sK10) | ~b(sK8,sK9,sK10)) & (tab(sK8,sK9,sK10) | b(sK8,sK9,sK10)))) & tab(X0,sK3(X0),sK4(X0)) & sK3(X0) = X0) | ~main_query_1(X0))) [skolemisation 19,21,20]
23. main_query_1(sK0) | main_query_0(sK0) [cnf transformation 13]
24. ~main_query_1(sK0) | ~main_query_0(sK0) [cnf transformation 13]
25. tab(X0,sK1(X0),sK2(X0)) | ~main_query_0(X0) [cnf transformation 17]
26. ~main_query_0(X0) | sK1(X0) = X0 [cnf transformation 17]
27. main_query_0(X0) | X0 != X1 | ~tab(X0,X1,X2) [cnf transformation 17]
28. ~main_query_1(X0) | sK3(X0) = X0 [cnf transformation 22]
29. tab(X0,sK3(X0),sK4(X0)) | ~main_query_1(X0) [cnf transformation 22]
33. main_query_1(X0) | b(X6,X7,X8) | ~tab(X6,X7,X8) | ~tab(X0,X1,X2) | X0 != X1 [cnf transformation 22]
34. main_query_1(X0) | ~b(X3,X4,X5) | ~tab(X0,X1,X2) | X0 != X1 [cnf transformation 22]
35. ~tab(X1,X1,X2) | main_query_0(X1) [equality resolution 27]
36. main_query_1(X1) | ~b(X3,X4,X5) | ~tab(X1,X1,X2) [equality resolution 34]
37. main_query_1(X1) | b(X6,X7,X8) | ~tab(X6,X7,X8) | ~tab(X1,X1,X2) [equality resolution 33]
40. 0 <=> main_query_0(sK0) [avatar definition]
41. main_query_0(sK0) <- {0} [avatar component clause 40]
43. 1 <=> ~main_query_0(sK0) [avatar definition]
44. ~main_query_0(sK0) <- {1} [avatar component clause 43]
46. 2 <=> main_query_1(sK0) [avatar definition]
47. main_query_1(sK0) <- {2} [avatar component clause 46]
49. 3 <=> ~main_query_1(sK0) [avatar definition]
50. ~main_query_1(sK0) <- {3} [avatar component clause 49]
51. ~1 | ~3 [avatar split clause 24,49,43]
52. 0 | 2 [avatar split clause 23,46,40]
54. 4 <=> ! [X3,X5,X4] : ~b(X3,X4,X5) [avatar definition]
55. ~b(X3,X4,X5) <- {4} [avatar component clause 54]
57. 6 <=> ! [X1,X2] : (main_query_1(X1) | ~tab(X1,X1,X2)) [avatar definition]
58. ~tab(X1,X1,X2) | main_query_1(X1) <- {6} [avatar component clause 57]
59. 4 | 6 [avatar split clause 36,57,54]
61. 8 <=> ! [X7,X8,X6] : (b(X6,X7,X8) | ~tab(X6,X7,X8)) [avatar definition]
62. ~tab(X6,X7,X8) | b(X6,X7,X8) <- {8} [avatar component clause 61]
63. 8 | 6 [avatar split clause 37,57,61]
85. sK0 = sK3(sK0) <- {2} [resolution 28,47]
86. tab(sK0,sK0,sK4(sK0)) | ~main_query_1(sK0) <- {2} [superposition 29,85]
87. tab(sK0,sK0,sK4(sK0)) <- {2} [subsumption resolution 86,47]
89. main_query_0(sK0) <- {2} [resolution 87,35]
90. $false <- {1, 2} [subsumption resolution 89,44]
91. 1 | ~2 [avatar contradiction clause 90]
92. sK0 = sK1(sK0) <- {0} [resolution 41,26]
93. tab(sK0,sK0,sK2(sK0)) | ~main_query_0(sK0) <- {0} [superposition 25,92]
94. tab(sK0,sK0,sK2(sK0)) <- {0} [subsumption resolution 93,41]
95. main_query_1(sK0) <- {0, 6} [resolution 94,58]
97. $false <- {0, 3, 6} [subsumption resolution 95,50]
98. ~0 | 3 | ~6 [avatar contradiction clause 97]
104. b(sK0,sK0,sK2(sK0)) <- {0, 8} [resolution 62,94]
106. $false <- {0, 4, 8} [subsumption resolution 104,55]
107. ~0 | ~4 | ~8 [avatar contradiction clause 106]
108. $false [avatar sat refutation 51,52,59,63,91,98,107]
% SZS output end Proof for test4
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.056 s
% ------------------------------
% ------------------------------
