% Refutation found. Thanks to Tanya!
% SZS status Theorem for test3
% SZS output start Proof for test3
1. ! [X0] : (main_query_0(X0) <=> main_query_1(X0)) [input]
2. ~! [X0] : (main_query_0(X0) <=> main_query_1(X0)) [negated conjecture 1]
3. ! [X0] : (main_query_0(X0) <=> ? [X1,X2] : (X0 = X1 & $true & tab(X0,X1,X2))) [input]
4. k <=> ! [X3,X4,X5] : (b(X3,X4,X5) <=> tab(X3,X4,X5)) [input]
5. ! [X0] : (main_query_1(X0) <=> ? [X1,X2,X6,X7,X8] : ((k => b(X6,X7,X8)) & tab(X0,X1,X2) & X0 = X1)) [input]
6. ! [X0] : (main_query_0(X0) <=> ? [X1,X2] : (X0 = X1 & tab(X0,X1,X2))) [true and false elimination 3]
7. k <=> ! [X0,X1,X2] : (b(X0,X1,X2) <=> tab(X0,X1,X2)) [rectify 4]
8. ! [X0] : (main_query_1(X0) <=> ? [X1,X2,X3,X4,X5] : ((k => b(X3,X4,X5)) & tab(X0,X1,X2) & X0 = X1)) [rectify 5]
9. ? [X0] : (main_query_0(X0) <~> main_query_1(X0)) [ennf transformation 2]
10. ! [X0] : (main_query_1(X0) <=> ? [X1,X2,X3,X4,X5] : ((b(X3,X4,X5) | ~k) & tab(X0,X1,X2) & X0 = X1)) [ennf transformation 8]
11. ? [X0] : ((~main_query_1(X0) | ~main_query_0(X0)) & (main_query_1(X0) | main_query_0(X0))) [nnf transformation 9]
12. ? [X0] : ((~main_query_1(X0) | ~main_query_0(X0)) & (main_query_1(X0) | main_query_0(X0))) => ((~main_query_1(sK0) | ~main_query_0(sK0)) & (main_query_1(sK0) | main_query_0(sK0))) [choice axiom]
13. (~main_query_1(sK0) | ~main_query_0(sK0)) & (main_query_1(sK0) | main_query_0(sK0)) [skolemisation 11,12]
14. ! [X0] : ((main_query_0(X0) | ! [X1,X2] : (X0 != X1 | ~tab(X0,X1,X2))) & (? [X1,X2] : (X0 = X1 & tab(X0,X1,X2)) | ~main_query_0(X0))) [nnf transformation 6]
15. ! [X0] : ((main_query_0(X0) | ! [X1,X2] : (X0 != X1 | ~tab(X0,X1,X2))) & (? [X3,X4] : (X0 = X3 & tab(X0,X3,X4)) | ~main_query_0(X0))) [rectify 14]
16. ! [X0] : (? [X3,X4] : (X0 = X3 & tab(X0,X3,X4)) => (sK1(X0) = X0 & tab(X0,sK1(X0),sK2(X0)))) [choice axiom]
17. ! [X0] : ((main_query_0(X0) | ! [X1,X2] : (X0 != X1 | ~tab(X0,X1,X2))) & ((sK1(X0) = X0 & tab(X0,sK1(X0),sK2(X0))) | ~main_query_0(X0))) [skolemisation 15,16]
18. (k | ? [X0,X1,X2] : ((~tab(X0,X1,X2) | ~b(X0,X1,X2)) & (tab(X0,X1,X2) | b(X0,X1,X2)))) & (! [X0,X1,X2] : ((b(X0,X1,X2) | ~tab(X0,X1,X2)) & (tab(X0,X1,X2) | ~b(X0,X1,X2))) | ~k) [nnf transformation 7]
19. (k | ? [X0,X1,X2] : ((~tab(X0,X1,X2) | ~b(X0,X1,X2)) & (tab(X0,X1,X2) | b(X0,X1,X2)))) & (! [X3,X4,X5] : ((b(X3,X4,X5) | ~tab(X3,X4,X5)) & (tab(X3,X4,X5) | ~b(X3,X4,X5))) | ~k) [rectify 18]
20. ? [X0,X1,X2] : ((~tab(X0,X1,X2) | ~b(X0,X1,X2)) & (tab(X0,X1,X2) | b(X0,X1,X2))) => ((~tab(sK3,sK4,sK5) | ~b(sK3,sK4,sK5)) & (tab(sK3,sK4,sK5) | b(sK3,sK4,sK5))) [choice axiom]
21. (k | ((~tab(sK3,sK4,sK5) | ~b(sK3,sK4,sK5)) & (tab(sK3,sK4,sK5) | b(sK3,sK4,sK5)))) & (! [X3,X4,X5] : ((b(X3,X4,X5) | ~tab(X3,X4,X5)) & (tab(X3,X4,X5) | ~b(X3,X4,X5))) | ~k) [skolemisation 19,20]
22. ! [X0] : ((main_query_1(X0) | ! [X1,X2,X3,X4,X5] : ((~b(X3,X4,X5) & k) | ~tab(X0,X1,X2) | X0 != X1)) & (? [X1,X2,X3,X4,X5] : ((b(X3,X4,X5) | ~k) & tab(X0,X1,X2) & X0 = X1) | ~main_query_1(X0))) [nnf transformation 10]
23. ! [X0] : ((main_query_1(X0) | ! [X1,X2,X3,X4,X5] : ((~b(X3,X4,X5) & k) | ~tab(X0,X1,X2) | X0 != X1)) & (? [X6,X7,X8,X9,X10] : ((b(X8,X9,X10) | ~k) & tab(X0,X6,X7) & X0 = X6) | ~main_query_1(X0))) [rectify 22]
24. ! [X0] : (? [X6,X7,X8,X9,X10] : ((b(X8,X9,X10) | ~k) & tab(X0,X6,X7) & X0 = X6) => ((b(sK8(X0),sK9(X0),sK10(X0)) | ~k) & tab(X0,sK6(X0),sK7(X0)) & sK6(X0) = X0)) [choice axiom]
25. ! [X0] : ((main_query_1(X0) | ! [X1,X2,X3,X4,X5] : ((~b(X3,X4,X5) & k) | ~tab(X0,X1,X2) | X0 != X1)) & (((b(sK8(X0),sK9(X0),sK10(X0)) | ~k) & tab(X0,sK6(X0),sK7(X0)) & sK6(X0) = X0) | ~main_query_1(X0))) [skolemisation 23,24]
26. main_query_1(sK0) | main_query_0(sK0) [cnf transformation 13]
27. ~main_query_1(sK0) | ~main_query_0(sK0) [cnf transformation 13]
28. tab(X0,sK1(X0),sK2(X0)) | ~main_query_0(X0) [cnf transformation 17]
29. ~main_query_0(X0) | sK1(X0) = X0 [cnf transformation 17]
30. main_query_0(X0) | X0 != X1 | ~tab(X0,X1,X2) [cnf transformation 17]
32. b(X3,X4,X5) | ~tab(X3,X4,X5) | ~k [cnf transformation 21]
35. ~main_query_1(X0) | sK6(X0) = X0 [cnf transformation 25]
36. tab(X0,sK6(X0),sK7(X0)) | ~main_query_1(X0) [cnf transformation 25]
38. main_query_1(X0) | k | ~tab(X0,X1,X2) | X0 != X1 [cnf transformation 25]
39. main_query_1(X0) | ~b(X3,X4,X5) | ~tab(X0,X1,X2) | X0 != X1 [cnf transformation 25]
40. ~tab(X1,X1,X2) | main_query_0(X1) [equality resolution 30]
41. main_query_1(X1) | ~b(X3,X4,X5) | ~tab(X1,X1,X2) [equality resolution 39]
42. main_query_1(X1) | k | ~tab(X1,X1,X2) [equality resolution 38]
44. 0 <=> main_query_0(sK0) [avatar definition]
45. main_query_0(sK0) <- {0} [avatar component clause 44]
47. 1 <=> ~main_query_0(sK0) [avatar definition]
48. ~main_query_0(sK0) <- {1} [avatar component clause 47]
50. 2 <=> main_query_1(sK0) [avatar definition]
51. main_query_1(sK0) <- {2} [avatar component clause 50]
53. 3 <=> ~main_query_1(sK0) [avatar definition]
54. ~main_query_1(sK0) <- {3} [avatar component clause 53]
55. ~1 | ~3 [avatar split clause 27,53,47]
56. 0 | 2 [avatar split clause 26,50,44]
70. 9 <=> ~k [avatar definition]
73. 8 <=> k [avatar definition]
78. 10 <=> ! [X3,X5,X4] : (b(X3,X4,X5) | ~tab(X3,X4,X5)) [avatar definition]
79. b(X3,X4,X5) | ~tab(X3,X4,X5) <- {10} [avatar component clause 78]
80. ~9 | 10 [avatar split clause 32,78,70]
86. 14 <=> ! [X3,X5,X4] : ~b(X3,X4,X5) [avatar definition]
87. ~b(X3,X4,X5) <- {14} [avatar component clause 86]
89. 16 <=> ! [X1,X2] : (main_query_1(X1) | ~tab(X1,X1,X2)) [avatar definition]
90. ~tab(X1,X1,X2) | main_query_1(X1) <- {16} [avatar component clause 89]
91. 14 | 16 [avatar split clause 41,89,86]
92. 8 | 16 [avatar split clause 42,89,73]
97. sK0 = sK6(sK0) <- {2} [resolution 35,51]
98. tab(sK0,sK0,sK7(sK0)) | ~main_query_1(sK0) <- {2} [superposition 36,97]
99. tab(sK0,sK0,sK7(sK0)) <- {2} [subsumption resolution 98,51]
101. main_query_0(sK0) <- {2} [resolution 99,40]
102. $false <- {1, 2} [subsumption resolution 101,48]
103. 1 | ~2 [avatar contradiction clause 102]
104. sK0 = sK1(sK0) <- {0} [resolution 45,29]
105. tab(sK0,sK0,sK2(sK0)) | ~main_query_0(sK0) <- {0} [superposition 28,104]
106. tab(sK0,sK0,sK2(sK0)) <- {0} [subsumption resolution 105,45]
107. main_query_1(sK0) <- {0, 16} [resolution 106,90]
109. $false <- {0, 3, 16} [subsumption resolution 107,54]
110. ~0 | 3 | ~16 [avatar contradiction clause 109]
111. ~tab(X3,X4,X5) <- {10, 14} [subsumption resolution 79,87]
116. $false <- {0, 10, 14} [resolution 111,106]
117. ~0 | ~10 | ~14 [avatar contradiction clause 116]
118. $false [avatar sat refutation 55,56,80,91,92,103,110,117]
% SZS output end Proof for test3
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.067 s
% ------------------------------
% ------------------------------
