% Refutation found. Thanks to Tanya!
% SZS status Theorem for query05_query08_tptp4X
% SZS output start Proof for query05_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
7. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [input]
9. ! [X5] : eq(X5,X5) [input]
12. ! [X5] : lte(X5,X5) [input]
18. ! [X5,X6] : (gte(X5,X6) <=> lte(X6,X5)) [input]
42. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
43. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 42]
44. ! [X0] : eq(X0,X0) [rectify 9]
47. ! [X0] : lte(X0,X0) [rectify 12]
53. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 18]
128. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 43]
129. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
130. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
131. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
132. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 131]
133. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 132,133]
135. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
136. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
137. ! [X0] : ((where_6(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [nnf transformation 6]
138. ! [X0] : ((where_6(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [flattening 137]
139. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 7]
140. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 139]
141. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
142. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 140,141]
143. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 8]
146. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 53]
147. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 128]
148. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
149. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 147,148]
150. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 129]
151. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 129]
153. ~eq(X0,X0) | where_3(X0) [cnf transformation 130]
154. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 134]
156. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 134]
157. ~main_query_1(X0) | select_1(X0) [cnf transformation 135]
158. ~select_1(X0) | main_query_1(X0) [cnf transformation 135]
159. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 136]
160. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 136]
163. ~gte(X0,X0) | where_6(X0) | ~lte(X0,X0) [cnf transformation 138]
164. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 142]
166. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 142]
167. ~main_query_2(X0) | select_4(X0) [cnf transformation 143]
168. ~select_4(X0) | main_query_2(X0) [cnf transformation 143]
169. eq(X0,X0) [cnf transformation 44]
172. lte(X0,X0) [cnf transformation 47]
178. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 146]
202. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 149]
203. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 149]
228. 0 <=> main_query_1(sK8) [avatar definition]
229. main_query_1(sK8) <- {0} [avatar component clause 228]
231. 1 <=> ~main_query_1(sK8) [avatar definition]
232. ~main_query_1(sK8) <- {1} [avatar component clause 231]
234. 2 <=> main_query_2(sK8) [avatar definition]
235. main_query_2(sK8) <- {2} [avatar component clause 234]
237. 3 <=> ~main_query_2(sK8) [avatar definition]
238. ~main_query_2(sK8) <- {3} [avatar component clause 237]
239. ~1 | ~3 [avatar split clause 203,237,231]
240. 0 | 2 [avatar split clause 202,234,228]
241. select_4(sK8) <- {2} [resolution 167,235]
245. where_3(X1) [resolution 153,169]
254. gte(X1,X1) [resolution 178,172]
256. where_6(X0) | ~lte(X0,X0) [resolution 163,254]
258. where_6(X0) [subsumption resolution 256,172]
265. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 154,150]
267. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 164,159]
269. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 265,160]
272. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 267,151]
274. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 269,166]
275. ~select_1(X1) | select_4(X1) [subsumption resolution 274,258]
277. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 272,156]
278. ~select_4(X1) | select_1(X1) [subsumption resolution 277,245]
279. select_1(sK8) <- {2} [resolution 278,241]
281. main_query_1(sK8) <- {2} [resolution 279,158]
283. $false <- {1, 2} [subsumption resolution 281,232]
284. 1 | ~2 [avatar contradiction clause 283]
285. select_1(sK8) <- {0} [resolution 229,157]
286. select_4(sK8) <- {0} [resolution 285,275]
290. main_query_2(sK8) <- {0} [resolution 286,168]
292. $false <- {0, 3} [subsumption resolution 290,238]
293. ~0 | 3 [avatar contradiction clause 292]
294. $false [avatar sat refutation 239,240,284,293]
% SZS output end Proof for query05_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.015 s
% ------------------------------
% ------------------------------
