% Refutation found. Thanks to Tanya!
% SZS status Theorem for query12_query14_tptp4X
% SZS output start Proof for query12_query14_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [input]
3. ! [X1,X0] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (query_7(X5) <=> select_8(X5)) [input]
9. ! [X0,X1] : (where_6(X0,X1) <=> (query_7(X0) & eq(X0,X1))) [input]
10. ! [X1,X0] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
16. ! [X7,X8] : ((lte(X8,X7) & lte(X7,X8)) <=> eq(X7,X8)) [input]
62. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
63. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 62]
64. ! [X0,X1] : (select_1(X1,X0) <=> ? [X2,X3,X4] : (where_3(X1,X0) & from_2(X1,X0,X2,X3,X4))) [rectify 3]
65. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
66. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 7]
67. ! [X0] : (query_7(X0) <=> select_8(X0)) [rectify 8]
68. ! [X0,X1] : (select_4(X1,X0) <=> ? [X2,X3,X4] : (where_6(X1,X0) & from_5(X1,X0,X2,X3,X4))) [rectify 10]
69. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
74. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 16]
120. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 63]
206. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 120]
207. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
208. ! [X0,X1] : ((where_3(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [nnf transformation 2]
209. ! [X0,X1] : ((where_3(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [flattening 208]
210. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X1,X0) | ~from_2(X1,X0,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X1,X0) & from_2(X1,X0,X2,X3,X4)) | ~select_1(X1,X0))) [nnf transformation 64]
211. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X1,X0) | ~from_2(X1,X0,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X1,X0) & from_2(X1,X0,X5,X6,X7)) | ~select_1(X1,X0))) [rectify 210]
212. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X1,X0) & from_2(X1,X0,X5,X6,X7)) => (where_3(X1,X0) & from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
213. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X1,X0) | ~from_2(X1,X0,X2,X3,X4))) & ((where_3(X1,X0) & from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X1,X0))) [skolemisation 211,212]
214. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 65]
215. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
216. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
217. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 66]
218. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 217]
219. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
220. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 218,219]
221. ! [X0] : ((query_7(X0) | ~select_8(X0)) & (select_8(X0) | ~query_7(X0))) [nnf transformation 67]
222. ! [X0,X1] : ((where_6(X0,X1) | (~query_7(X0) | ~eq(X0,X1))) & ((query_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 9]
223. ! [X0,X1] : ((where_6(X0,X1) | ~query_7(X0) | ~eq(X0,X1)) & ((query_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 222]
224. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X1,X0) | ~from_5(X1,X0,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X1,X0) & from_5(X1,X0,X2,X3,X4)) | ~select_4(X1,X0))) [nnf transformation 68]
225. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X1,X0) | ~from_5(X1,X0,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X1,X0) & from_5(X1,X0,X5,X6,X7)) | ~select_4(X1,X0))) [rectify 224]
226. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X1,X0) & from_5(X1,X0,X5,X6,X7)) => (where_6(X1,X0) & from_5(X1,X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
227. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X1,X0) | ~from_5(X1,X0,X2,X3,X4))) & ((where_6(X1,X0) & from_5(X1,X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X1,X0))) [skolemisation 225,226]
228. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 69]
229. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 74]
230. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 229]
231. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 206]
232. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
233. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 231,232]
234. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 207]
235. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 207]
236. ~where_3(X0,X1) | lte(X0,X1) [cnf transformation 209]
237. ~where_3(X0,X1) | lte(X1,X0) [cnf transformation 209]
238. ~lte(X1,X0) | where_3(X0,X1) | ~lte(X0,X1) [cnf transformation 209]
239. from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X1,X0) [cnf transformation 213]
240. ~select_1(X1,X0) | where_3(X1,X0) [cnf transformation 213]
241. ~from_2(X1,X0,X2,X3,X4) | ~where_3(X1,X0) | select_1(X1,X0) [cnf transformation 213]
242. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 214]
243. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 214]
244. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 215]
245. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 215]
247. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 216]
249. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 220]
251. ~select_8(X0) | query_7(X0) [cnf transformation 221]
252. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 223]
254. ~eq(X0,X1) | ~query_7(X0) | where_6(X0,X1) [cnf transformation 223]
255. from_5(X1,X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X1,X0) [cnf transformation 227]
256. ~select_4(X1,X0) | where_6(X1,X0) [cnf transformation 227]
257. ~from_5(X1,X0,X2,X3,X4) | ~where_6(X1,X0) | select_4(X1,X0) [cnf transformation 227]
258. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 228]
259. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 228]
264. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 230]
265. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 230]
266. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 230]
308. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 233]
309. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 233]
351. 0 <=> main_query_1(sK10,sK11) [avatar definition]
352. main_query_1(sK10,sK11) <- {0} [avatar component clause 351]
354. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
355. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 354]
357. 2 <=> main_query_2(sK10,sK11) [avatar definition]
358. main_query_2(sK10,sK11) <- {2} [avatar component clause 357]
360. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
361. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 360]
362. ~1 | ~3 [avatar split clause 309,360,354]
363. 0 | 2 [avatar split clause 308,357,351]
364. select_4(sK10,sK11) <- {2} [resolution 258,358]
366. where_6(sK10,sK11) <- {2} [resolution 364,256]
368. eq(sK10,sK11) <- {2} [resolution 366,252]
374. lte(sK11,sK10) <- {2} [resolution 368,266]
375. lte(sK10,sK11) <- {2} [resolution 368,265]
380. where_3(sK10,sK11) | ~lte(sK10,sK11) <- {2} [resolution 238,374]
382. where_3(sK10,sK11) <- {2} [subsumption resolution 380,375]
430. tab1(X2,X3,sK0(X3,X2),sK1(X3,X2),sK2(X3,X2)) | ~select_1(X2,X3) [resolution 239,234]
435. tab1(X2,X3,sK7(X3,X2),sK8(X3,X2),sK9(X3,X2)) | ~select_4(X2,X3) [resolution 255,244]
438. from_9(X0,X1,sK0(X1,X0),sK1(X1,X0),sK2(X1,X0)) | ~select_1(X0,X1) [resolution 430,247]
439. from_5(X2,X3,sK0(X3,X2),sK1(X3,X2),sK2(X3,X2)) | ~select_1(X2,X3) [resolution 430,245]
443. from_2(X4,X5,sK7(X5,X4),sK8(X5,X4),sK9(X5,X4)) | ~select_4(X4,X5) [resolution 435,235]
447. ~select_1(X2,X3) | select_8(X2) [resolution 438,249]
448. ~where_6(X0,X1) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 439,257]
455. ~select_4(X0,X1) | ~where_3(X0,X1) | select_1(X0,X1) [resolution 443,241]
457. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 455,364]
458. select_1(sK10,sK11) <- {2} [subsumption resolution 457,382]
460. main_query_1(sK10,sK11) <- {2} [resolution 458,243]
462. $false <- {1, 2} [subsumption resolution 460,355]
463. 1 | ~2 [avatar contradiction clause 462]
464. select_1(sK10,sK11) <- {0} [resolution 352,242]
465. select_8(sK10) <- {0} [resolution 464,447]
467. where_3(sK10,sK11) <- {0} [resolution 464,240]
468. query_7(sK10) <- {0} [resolution 465,251]
469. lte(sK11,sK10) <- {0} [resolution 467,237]
470. lte(sK10,sK11) <- {0} [resolution 467,236]
473. eq(sK10,sK11) | ~lte(sK10,sK11) <- {0} [resolution 469,264]
475. eq(sK10,sK11) <- {0} [subsumption resolution 473,470]
482. ~query_7(sK10) | where_6(sK10,sK11) <- {0} [resolution 475,254]
486. where_6(sK10,sK11) <- {0} [subsumption resolution 482,468]
496. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {0} [resolution 486,448]
499. select_4(sK10,sK11) <- {0} [subsumption resolution 496,464]
504. main_query_2(sK10,sK11) <- {0} [resolution 499,259]
506. $false <- {0, 3} [subsumption resolution 504,361]
507. ~0 | 3 [avatar contradiction clause 506]
508. $false [avatar sat refutation 362,363,463,507]
% SZS output end Proof for query12_query14_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.014 s
% ------------------------------
% ------------------------------
