% Refutation found. Thanks to Tanya!
% SZS status Theorem for query30_query33_tptp4X
% SZS output start Proof for query30_query33_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> lt(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1] : (where_6(X0,X1) <=> lt(X0,X1)) [input]
7. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
55. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
56. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 55]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
58. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
105. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 56]
181. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 105]
182. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
183. ! [X0,X1] : ((where_3(X0,X1) | ~lt(X0,X1)) & (lt(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
184. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
185. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 184]
186. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
187. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 185,186]
188. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
189. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
190. ! [X0,X1] : ((where_6(X0,X1) | ~lt(X0,X1)) & (lt(X0,X1) | ~where_6(X0,X1))) [nnf transformation 6]
191. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 7]
192. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 191]
193. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
194. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 192,193]
195. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 58]
200. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 181]
201. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
202. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 200,201]
203. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 182]
204. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 182]
205. ~where_3(X0,X1) | lt(X0,X1) [cnf transformation 183]
206. ~lt(X0,X1) | where_3(X0,X1) [cnf transformation 183]
207. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 187]
208. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 187]
209. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 187]
210. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 188]
211. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 188]
212. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 189]
213. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 189]
214. ~where_6(X0,X1) | lt(X0,X1) [cnf transformation 190]
215. ~lt(X0,X1) | where_6(X0,X1) [cnf transformation 190]
216. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 194]
217. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 194]
218. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 194]
219. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 195]
220. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 195]
268. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 202]
269. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 202]
271. 0 <=> main_query_1(sK6,sK7) [avatar definition]
272. main_query_1(sK6,sK7) <- {0} [avatar component clause 271]
274. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
275. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 274]
277. 2 <=> main_query_2(sK6,sK7) [avatar definition]
278. main_query_2(sK6,sK7) <- {2} [avatar component clause 277]
280. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
281. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 280]
282. ~1 | ~3 [avatar split clause 269,280,274]
283. 0 | 2 [avatar split clause 268,277,271]
284. select_4(sK6,sK7) <- {2} [resolution 219,278]
286. where_6(sK6,sK7) <- {2} [resolution 284,217]
288. lt(sK6,sK7) <- {2} [resolution 286,214]
291. where_3(sK6,sK7) <- {2} [resolution 288,206]
468. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 207,203]
474. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 216,212]
641. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 468,213]
657. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 474,204]
910. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 641,218]
934. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 657,209]
940. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 934,284]
941. select_1(sK6,sK7) <- {2} [subsumption resolution 940,291]
949. main_query_1(sK6,sK7) <- {2} [resolution 941,211]
951. $false <- {1, 2} [subsumption resolution 949,275]
952. 1 | ~2 [avatar contradiction clause 951]
953. select_1(sK6,sK7) <- {0} [resolution 272,210]
955. where_3(sK6,sK7) <- {0} [resolution 953,208]
1001. lt(sK6,sK7) <- {0} [resolution 955,205]
1003. where_6(sK6,sK7) <- {0} [resolution 1001,215]
1026. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 1003,910]
1028. select_4(sK6,sK7) <- {0} [subsumption resolution 1026,953]
1074. main_query_2(sK6,sK7) <- {0} [resolution 1028,220]
1076. $false <- {0, 3} [subsumption resolution 1074,281]
1077. ~0 | 3 [avatar contradiction clause 1076]
1078. $false [avatar sat refutation 282,283,952,1077]
% SZS output end Proof for query30_query33_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.030 s
% ------------------------------
% ------------------------------
