% Refutation found. Thanks to Tanya!
% SZS status Theorem for query05_query08_tptp4X
% SZS output start Proof for query05_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
7. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
9. ! [X6] : eq(X6,X6) [input]
12. ! [X6] : lte(X6,X6) [input]
19. ! [X6,X7] : (gte(X6,X7) <=> lte(X7,X6)) [input]
43. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
44. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 43]
45. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
46. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
47. ! [X0] : eq(X0,X0) [rectify 9]
50. ! [X0] : lte(X0,X0) [rectify 12]
57. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 19]
81. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 44]
133. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 81]
134. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
135. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 136]
138. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
139. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 137,138]
140. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 45]
141. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
142. ! [X0] : ((where_6(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [nnf transformation 6]
143. ! [X0] : ((where_6(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [flattening 142]
144. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 7]
145. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 144]
146. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
147. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 145,146]
148. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 46]
151. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 57]
152. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 133]
153. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
154. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 152,153]
155. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 134]
156. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 134]
158. ~eq(X0,X0) | where_3(X0) [cnf transformation 135]
159. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 139]
161. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 139]
162. ~main_query_1(X0) | select_1(X0) [cnf transformation 140]
163. ~select_1(X0) | main_query_1(X0) [cnf transformation 140]
164. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 141]
165. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 141]
168. ~gte(X0,X0) | where_6(X0) | ~lte(X0,X0) [cnf transformation 143]
169. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 147]
171. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 147]
172. ~main_query_2(X0) | select_4(X0) [cnf transformation 148]
173. ~select_4(X0) | main_query_2(X0) [cnf transformation 148]
174. eq(X0,X0) [cnf transformation 47]
177. lte(X0,X0) [cnf transformation 50]
184. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 151]
208. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 154]
209. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 154]
211. 0 <=> main_query_1(sK8) [avatar definition]
212. main_query_1(sK8) <- {0} [avatar component clause 211]
214. 1 <=> ~main_query_1(sK8) [avatar definition]
215. ~main_query_1(sK8) <- {1} [avatar component clause 214]
217. 2 <=> main_query_2(sK8) [avatar definition]
218. main_query_2(sK8) <- {2} [avatar component clause 217]
220. 3 <=> ~main_query_2(sK8) [avatar definition]
221. ~main_query_2(sK8) <- {3} [avatar component clause 220]
222. ~1 | ~3 [avatar split clause 209,220,214]
223. 0 | 2 [avatar split clause 208,217,211]
224. select_4(sK8) <- {2} [resolution 172,218]
228. where_3(X1) [resolution 158,174]
238. gte(X1,X1) [resolution 184,177]
243. where_6(X1) | ~lte(X1,X1) [resolution 238,168]
244. where_6(X1) [subsumption resolution 243,177]
311. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 159,155]
313. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 169,164]
315. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 311,165]
319. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 313,156]
324. ~select_1(X5) | ~where_6(X5) | select_4(X5) [resolution 315,171]
325. ~select_1(X5) | select_4(X5) [subsumption resolution 324,244]
331. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 319,161]
332. ~select_4(X1) | select_1(X1) [subsumption resolution 331,228]
335. select_1(sK8) <- {2} [resolution 332,224]
337. main_query_1(sK8) <- {2} [resolution 335,163]
339. $false <- {1, 2} [subsumption resolution 337,215]
340. 1 | ~2 [avatar contradiction clause 339]
341. select_1(sK8) <- {0} [resolution 212,162]
342. select_4(sK8) <- {0} [resolution 341,325]
346. main_query_2(sK8) <- {0} [resolution 342,173]
348. $false <- {0, 3} [subsumption resolution 346,221]
349. ~0 | 3 [avatar contradiction clause 348]
350. $false [avatar sat refutation 222,223,340,349]
% SZS output end Proof for query05_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
