% Refutation found. Thanks to Tanya!
% SZS status Theorem for query20_query21_tptp4X
% SZS output start Proof for query20_query21_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
3. ! [X0,X1] : (where_10(X0,X1) <=> eq(X0,X1)) [input]
4. ! [X0,X1] : (select_8(X0,X1) <=> ? [X2,X3,X4] : (where_10(X0,X1) & from_9(X0,X1,X2,X3,X4))) [input]
5. ! [X5,X6] : (subquery_from_7(X5,X6) <=> select_8(X5,X6)) [input]
6. ! [X7,X8] : (from_6(X7,X8) <=> subquery_from_7(X7,X8)) [input]
7. ! [X7] : (select_5(X7) <=> ? [X8] : ($true & from_6(X7,X8))) [input]
8. ! [X5] : (where_3_subquery_4(X5) <=> select_5(X5)) [input]
9. ! [X0] : (where_3(X0) <=> where_3_subquery_4(X0)) [input]
10. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
11. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
12. ! [X0,X1,X2,X3,X4] : (from_15(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
13. ! [X0,X1] : (where_16(X0,X1) <=> eq(X0,X1)) [input]
14. ! [X0,X1] : (select_14(X0,X1) <=> ? [X2,X3,X4] : (where_16(X0,X1) & from_15(X0,X1,X2,X3,X4))) [input]
15. ! [X5,X6] : (subquery_from_13(X5,X6) <=> select_14(X5,X6)) [input]
16. ! [X7,X8] : (from_12(X7,X8) <=> subquery_from_13(X7,X8)) [input]
17. ! [X7] : (select_11(X7) <=> ? [X8] : ($true & from_12(X7,X8))) [input]
18. ! [X5] : (main_query_2(X5) <=> select_11(X5)) [input]
19. ! [X9] : eq(X9,X9) [input]
20. ! [X9,X10] : (eq(X9,X10) => eq(X10,X9)) [input]
52. ! [X9,X7,X8] : ((from_6(X7,X8) & eq(X8,X9)) => from_6(X7,X9)) [input]
68. ! [X9,X7,X8] : ((from_12(X7,X8) & eq(X8,X9)) => from_12(X7,X9)) [input]
74. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
75. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 74]
76. ! [X0,X1] : (subquery_from_7(X0,X1) <=> select_8(X0,X1)) [rectify 5]
77. ! [X0,X1] : (from_6(X0,X1) <=> subquery_from_7(X0,X1)) [rectify 6]
78. ! [X0] : (select_5(X0) <=> ? [X1] : ($true & from_6(X0,X1))) [rectify 7]
79. ! [X0] : (select_5(X0) <=> ? [X1] : from_6(X0,X1)) [true and false elimination 78]
80. ! [X0] : (where_3_subquery_4(X0) <=> select_5(X0)) [rectify 8]
81. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 11]
82. ! [X0,X1] : (subquery_from_13(X0,X1) <=> select_14(X0,X1)) [rectify 15]
83. ! [X0,X1] : (from_12(X0,X1) <=> subquery_from_13(X0,X1)) [rectify 16]
84. ! [X0] : (select_11(X0) <=> ? [X1] : ($true & from_12(X0,X1))) [rectify 17]
85. ! [X0] : (select_11(X0) <=> ? [X1] : from_12(X0,X1)) [true and false elimination 84]
86. ! [X0] : (main_query_2(X0) <=> select_11(X0)) [rectify 18]
87. ! [X0] : eq(X0,X0) [rectify 19]
88. ! [X0,X1] : (eq(X0,X1) => eq(X1,X0)) [rectify 20]
120. ! [X0,X1,X2] : ((from_6(X1,X2) & eq(X2,X0)) => from_6(X1,X0)) [rectify 52]
136. ! [X0,X1,X2] : ((from_12(X1,X2) & eq(X2,X0)) => from_12(X1,X0)) [rectify 68]
142. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 75]
143. ! [X0,X1] : (eq(X1,X0) | ~eq(X0,X1)) [ennf transformation 88]
192. ! [X0,X1,X2] : (from_6(X1,X0) | (~from_6(X1,X2) | ~eq(X2,X0))) [ennf transformation 120]
193. ! [X0,X1,X2] : (from_6(X1,X0) | ~from_6(X1,X2) | ~eq(X2,X0)) [flattening 192]
224. ! [X0,X1,X2] : (from_12(X1,X0) | (~from_12(X1,X2) | ~eq(X2,X0))) [ennf transformation 136]
225. ! [X0,X1,X2] : (from_12(X1,X0) | ~from_12(X1,X2) | ~eq(X2,X0)) [flattening 224]
236. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 142]
237. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
238. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 2]
239. ! [X0,X1] : ((where_10(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_10(X0,X1))) [nnf transformation 3]
240. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4] : (~where_10(X0,X1) | ~from_9(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_10(X0,X1) & from_9(X0,X1,X2,X3,X4)) | ~select_8(X0,X1))) [nnf transformation 4]
241. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4] : (~where_10(X0,X1) | ~from_9(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_10(X0,X1) & from_9(X0,X1,X5,X6,X7)) | ~select_8(X0,X1))) [rectify 240]
242. ! [X1,X0] : (? [X5,X6,X7] : (where_10(X0,X1) & from_9(X0,X1,X5,X6,X7)) => (where_10(X0,X1) & from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
243. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4] : (~where_10(X0,X1) | ~from_9(X0,X1,X2,X3,X4))) & ((where_10(X0,X1) & from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_8(X0,X1))) [skolemisation 241,242]
244. ! [X0,X1] : ((subquery_from_7(X0,X1) | ~select_8(X0,X1)) & (select_8(X0,X1) | ~subquery_from_7(X0,X1))) [nnf transformation 76]
245. ! [X0,X1] : ((from_6(X0,X1) | ~subquery_from_7(X0,X1)) & (subquery_from_7(X0,X1) | ~from_6(X0,X1))) [nnf transformation 77]
246. ! [X0] : ((select_5(X0) | ! [X1] : ~from_6(X0,X1)) & (? [X1] : from_6(X0,X1) | ~select_5(X0))) [nnf transformation 79]
247. ! [X0] : ((select_5(X0) | ! [X1] : ~from_6(X0,X1)) & (? [X2] : from_6(X0,X2) | ~select_5(X0))) [rectify 246]
248. ! [X0] : (? [X2] : from_6(X0,X2) => from_6(X0,sK3(X0))) [choice axiom]
249. ! [X0] : ((select_5(X0) | ! [X1] : ~from_6(X0,X1)) & (from_6(X0,sK3(X0)) | ~select_5(X0))) [skolemisation 247,248]
250. ! [X0] : ((where_3_subquery_4(X0) | ~select_5(X0)) & (select_5(X0) | ~where_3_subquery_4(X0))) [nnf transformation 80]
251. ! [X0] : ((where_3(X0) | ~where_3_subquery_4(X0)) & (where_3_subquery_4(X0) | ~where_3(X0))) [nnf transformation 9]
252. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 10]
253. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 252]
254. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
255. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_1(X0))) [skolemisation 253,254]
256. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 81]
257. ! [X0,X1,X2,X3,X4] : ((from_15(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_15(X0,X1,X2,X3,X4))) [nnf transformation 12]
258. ! [X0,X1] : ((where_16(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_16(X0,X1))) [nnf transformation 13]
259. ! [X0,X1] : ((select_14(X0,X1) | ! [X2,X3,X4] : (~where_16(X0,X1) | ~from_15(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_16(X0,X1) & from_15(X0,X1,X2,X3,X4)) | ~select_14(X0,X1))) [nnf transformation 14]
260. ! [X0,X1] : ((select_14(X0,X1) | ! [X2,X3,X4] : (~where_16(X0,X1) | ~from_15(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_16(X0,X1) & from_15(X0,X1,X5,X6,X7)) | ~select_14(X0,X1))) [rectify 259]
261. ! [X1,X0] : (? [X5,X6,X7] : (where_16(X0,X1) & from_15(X0,X1,X5,X6,X7)) => (where_16(X0,X1) & from_15(X0,X1,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)))) [choice axiom]
262. ! [X0,X1] : ((select_14(X0,X1) | ! [X2,X3,X4] : (~where_16(X0,X1) | ~from_15(X0,X1,X2,X3,X4))) & ((where_16(X0,X1) & from_15(X0,X1,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1))) | ~select_14(X0,X1))) [skolemisation 260,261]
263. ! [X0,X1] : ((subquery_from_13(X0,X1) | ~select_14(X0,X1)) & (select_14(X0,X1) | ~subquery_from_13(X0,X1))) [nnf transformation 82]
264. ! [X0,X1] : ((from_12(X0,X1) | ~subquery_from_13(X0,X1)) & (subquery_from_13(X0,X1) | ~from_12(X0,X1))) [nnf transformation 83]
265. ! [X0] : ((select_11(X0) | ! [X1] : ~from_12(X0,X1)) & (? [X1] : from_12(X0,X1) | ~select_11(X0))) [nnf transformation 85]
266. ! [X0] : ((select_11(X0) | ! [X1] : ~from_12(X0,X1)) & (? [X2] : from_12(X0,X2) | ~select_11(X0))) [rectify 265]
267. ! [X0] : (? [X2] : from_12(X0,X2) => from_12(X0,sK11(X0))) [choice axiom]
268. ! [X0] : ((select_11(X0) | ! [X1] : ~from_12(X0,X1)) & (from_12(X0,sK11(X0)) | ~select_11(X0))) [skolemisation 266,267]
269. ! [X0] : ((main_query_2(X0) | ~select_11(X0)) & (select_11(X0) | ~main_query_2(X0))) [nnf transformation 86]
272. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 236]
273. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK12) | ~main_query_1(sK12)) & (main_query_2(sK12) | main_query_1(sK12))) [choice axiom]
274. (~main_query_2(sK12) | ~main_query_1(sK12)) & (main_query_2(sK12) | main_query_1(sK12)) [skolemisation 272,273]
276. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 237]
277. ~from_9(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 238]
278. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 238]
279. ~where_10(X0,X1) | eq(X0,X1) [cnf transformation 239]
280. ~eq(X0,X1) | where_10(X0,X1) [cnf transformation 239]
281. from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_8(X0,X1) [cnf transformation 243]
282. ~select_8(X0,X1) | where_10(X0,X1) [cnf transformation 243]
283. ~from_9(X0,X1,X2,X3,X4) | ~where_10(X0,X1) | select_8(X0,X1) [cnf transformation 243]
284. ~subquery_from_7(X0,X1) | select_8(X0,X1) [cnf transformation 244]
285. ~select_8(X0,X1) | subquery_from_7(X0,X1) [cnf transformation 244]
286. ~from_6(X0,X1) | subquery_from_7(X0,X1) [cnf transformation 245]
287. ~subquery_from_7(X0,X1) | from_6(X0,X1) [cnf transformation 245]
288. from_6(X0,sK3(X0)) | ~select_5(X0) [cnf transformation 249]
289. ~from_6(X0,X1) | select_5(X0) [cnf transformation 249]
290. ~where_3_subquery_4(X0) | select_5(X0) [cnf transformation 250]
291. ~select_5(X0) | where_3_subquery_4(X0) [cnf transformation 250]
292. ~where_3(X0) | where_3_subquery_4(X0) [cnf transformation 251]
293. ~where_3_subquery_4(X0) | where_3(X0) [cnf transformation 251]
295. ~select_1(X0) | where_3(X0) [cnf transformation 255]
296. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 255]
297. ~main_query_1(X0) | select_1(X0) [cnf transformation 256]
298. ~select_1(X0) | main_query_1(X0) [cnf transformation 256]
299. ~from_15(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 257]
300. ~tab1(X0,X1,X2,X3,X4) | from_15(X0,X1,X2,X3,X4) [cnf transformation 257]
301. ~where_16(X0,X1) | eq(X0,X1) [cnf transformation 258]
302. ~eq(X0,X1) | where_16(X0,X1) [cnf transformation 258]
303. from_15(X0,X1,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) | ~select_14(X0,X1) [cnf transformation 262]
304. ~select_14(X0,X1) | where_16(X0,X1) [cnf transformation 262]
305. ~from_15(X0,X1,X2,X3,X4) | ~where_16(X0,X1) | select_14(X0,X1) [cnf transformation 262]
306. ~subquery_from_13(X0,X1) | select_14(X0,X1) [cnf transformation 263]
307. ~select_14(X0,X1) | subquery_from_13(X0,X1) [cnf transformation 263]
308. ~from_12(X0,X1) | subquery_from_13(X0,X1) [cnf transformation 264]
309. ~subquery_from_13(X0,X1) | from_12(X0,X1) [cnf transformation 264]
310. from_12(X0,sK11(X0)) | ~select_11(X0) [cnf transformation 268]
311. ~from_12(X0,X1) | select_11(X0) [cnf transformation 268]
312. ~main_query_2(X0) | select_11(X0) [cnf transformation 269]
313. ~select_11(X0) | main_query_2(X0) [cnf transformation 269]
314. eq(X0,X0) [cnf transformation 87]
315. ~eq(X0,X1) | eq(X1,X0) [cnf transformation 143]
345. ~from_6(X1,X2) | from_6(X1,X0) | ~eq(X2,X0) [cnf transformation 193]
361. ~from_12(X1,X2) | from_12(X1,X0) | ~eq(X2,X0) [cnf transformation 225]
367. main_query_2(sK12) | main_query_1(sK12) [cnf transformation 274]
368. ~main_query_2(sK12) | ~main_query_1(sK12) [cnf transformation 274]
370. 0 <=> main_query_1(sK12) [avatar definition]
371. main_query_1(sK12) <- {0} [avatar component clause 370]
373. 1 <=> ~main_query_1(sK12) [avatar definition]
376. 2 <=> main_query_2(sK12) [avatar definition]
377. main_query_2(sK12) <- {2} [avatar component clause 376]
379. 3 <=> ~main_query_2(sK12) [avatar definition]
380. ~main_query_2(sK12) <- {3} [avatar component clause 379]
381. ~1 | ~3 [avatar split clause 368,379,373]
382. 0 | 2 [avatar split clause 367,376,370]
383. select_11(sK12) <- {2} [resolution 312,377]
385. where_10(X0,X0) [resolution 280,314]
387. subquery_from_7(X0,sK3(X0)) | ~select_5(X0) [resolution 288,286]
389. where_16(X0,X0) [resolution 302,314]
391. subquery_from_13(X0,sK11(X0)) | ~select_11(X0) [resolution 310,308]
398. select_8(X1,sK3(X1)) | ~select_5(X1) [resolution 387,284]
401. select_14(X1,sK11(X1)) | ~select_11(X1) [resolution 391,306]
406. where_10(X1,sK3(X1)) | ~select_5(X1) [resolution 398,282]
411. where_16(X1,sK11(X1)) | ~select_11(X1) [resolution 401,304]
415. eq(X0,sK3(X0)) | ~select_5(X0) [resolution 406,279]
419. eq(X4,sK11(X4)) | ~select_11(X4) [resolution 411,301]
423. eq(sK3(X2),X2) | ~select_5(X2) [resolution 415,315]
431. eq(sK11(X2),X2) | ~select_11(X2) [resolution 419,315]
448. ~eq(sK3(X0),X1) | from_6(X0,X1) | ~select_5(X0) [resolution 345,288]
479. ~eq(sK11(X0),X1) | from_12(X0,X1) | ~select_11(X0) [resolution 361,310]
518. from_6(X0,X0) | ~select_5(X0) | ~select_5(X0) [resolution 448,423]
522. from_6(X0,X0) | ~select_5(X0) [duplicate literal removal 518]
526. subquery_from_7(X4,X4) | ~select_5(X4) [resolution 522,286]
529. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_8(X2,X3) [resolution 281,277]
533. select_8(X5,X5) | ~select_5(X5) [resolution 526,284]
535. tab1(X2,X3,sK8(X2,X3),sK9(X2,X3),sK10(X2,X3)) | ~select_14(X2,X3) [resolution 303,299]
687. from_12(X0,X0) | ~select_11(X0) | ~select_11(X0) [resolution 479,431]
692. from_12(X0,X0) | ~select_11(X0) [duplicate literal removal 687]
705. subquery_from_13(X4,X4) | ~select_11(X4) [resolution 692,308]
709. from_15(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_8(X4,X5) [resolution 529,300]
716. from_2(X0,X1,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) | ~select_14(X0,X1) [resolution 535,276]
717. from_9(X2,X3,sK8(X2,X3),sK9(X2,X3),sK10(X2,X3)) | ~select_14(X2,X3) [resolution 535,278]
727. select_14(X5,X5) | ~select_11(X5) [resolution 705,306]
1148. ~where_16(X2,X3) | ~select_8(X2,X3) | select_14(X2,X3) [resolution 709,305]
1157. ~select_14(X2,X3) | ~where_3(X2) | select_1(X2) [resolution 716,296]
1165. ~select_14(X2,X3) | ~where_10(X2,X3) | select_8(X2,X3) [resolution 717,283]
1202. ~select_11(X0) | select_1(X0) | ~where_3(X0) [resolution 1157,401]
1215. select_1(sK12) | ~where_3(sK12) <- {2} [resolution 1202,383]
1220. 4 <=> where_3(sK12) [avatar definition]
1221. where_3(sK12) <- {4} [avatar component clause 1220]
1223. 5 <=> ~where_3(sK12) [avatar definition]
1224. ~where_3(sK12) <- {5} [avatar component clause 1223]
1229. 6 <=> select_1(sK12) [avatar definition]
1230. select_1(sK12) <- {6} [avatar component clause 1229]
1231. ~5 | 6 | ~2 [avatar split clause 1215,376,1229,1223]
1971. ~select_8(X0,X0) | select_14(X0,X0) [resolution 1148,389]
1989. select_14(X1,X1) | ~select_5(X1) [resolution 1971,533]
1995. ~where_10(X5,X5) | select_8(X5,X5) | ~select_11(X5) [resolution 1165,727]
2004. select_8(X5,X5) | ~select_11(X5) [subsumption resolution 1995,385]
2009. subquery_from_13(X1,X1) | ~select_5(X1) [resolution 1989,307]
2017. subquery_from_7(X1,X1) | ~select_11(X1) [resolution 2004,285]
2026. from_12(X4,X4) | ~select_5(X4) [resolution 2009,309]
2042. from_6(X4,X4) | ~select_11(X4) [resolution 2017,287]
2060. ~select_5(X5) | select_11(X5) [resolution 2026,311]
2097. ~select_11(X5) | select_5(X5) [resolution 2042,289]
2118. select_5(sK12) <- {2} [resolution 2097,383]
2146. where_3_subquery_4(sK12) <- {2} [resolution 2118,291]
2154. where_3(sK12) <- {2} [resolution 2146,293]
2156. $false <- {2, 5} [subsumption resolution 2154,1224]
2157. ~2 | 5 [avatar contradiction clause 2156]
2158. select_1(sK12) <- {0} [resolution 371,297]
2159. 6 | ~0 [avatar split clause 2158,370,1229]
2172. where_3_subquery_4(sK12) <- {4} [resolution 1221,292]
2180. main_query_1(sK12) <- {6} [resolution 1230,298]
2181. where_3(sK12) <- {6} [resolution 1230,295]
2184. 0 | ~6 [avatar split clause 2180,1229,370]
2194. select_5(sK12) <- {4} [resolution 2172,290]
2202. select_11(sK12) <- {4} [resolution 2194,2060]
2213. main_query_2(sK12) <- {4} [resolution 2202,313]
2214. $false <- {3, 4} [subsumption resolution 2213,380]
2215. 3 | ~4 [avatar contradiction clause 2214]
2216. 4 | ~6 [avatar split clause 2181,1229,1220]
2217. $false [avatar sat refutation 381,382,1231,2157,2159,2184,2215,2216]
% SZS output end Proof for query20_query21_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5756
% Time elapsed: 0.026 s
% ------------------------------
% ------------------------------
