% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_01_query_09_tptp4X
% SZS output start Proof for query_01_query_09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [input]
7. ! [X5] : (where_5_subquery_6(X5) <=> select_7(X5)) [input]
8. ! [X0] : (where_5(X0) <=> where_5_subquery_6(X0)) [input]
9. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
10. ! [X5,X6] : (main_query_2(X5,X6) <=> select_3(X5,X6)) [input]
54. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
55. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 54]
56. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 3]
58. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 6]
59. ! [X0] : (where_5_subquery_6(X0) <=> select_7(X0)) [rectify 7]
60. ! [X0,X1] : (main_query_2(X0,X1) <=> select_3(X0,X1)) [rectify 10]
104. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 55]
176. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 104]
177. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
178. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0,X1))) [nnf transformation 56]
179. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) | ~select_1(X0,X1))) [rectify 178]
180. ! [X1,X0] : (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) => from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) [choice axiom]
181. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1))) [skolemisation 179,180]
182. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
183. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
184. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 5]
185. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 58]
186. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 185]
187. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
188. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_7(X0))) [skolemisation 186,187]
189. ! [X0] : ((where_5_subquery_6(X0) | ~select_7(X0)) & (select_7(X0) | ~where_5_subquery_6(X0))) [nnf transformation 59]
190. ! [X0] : ((where_5(X0) | ~where_5_subquery_6(X0)) & (where_5_subquery_6(X0) | ~where_5(X0))) [nnf transformation 8]
191. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0,X1))) [nnf transformation 9]
192. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) | ~select_3(X0,X1))) [rectify 191]
193. ! [X1,X0] : (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) => (where_5(X0) & from_4(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
194. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_3(X0,X1))) [skolemisation 192,193]
195. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 60]
198. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 176]
199. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
200. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 198,199]
201. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 177]
202. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 177]
203. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 181]
204. ~from_2(X0,X1,X2,X3,X4) | select_1(X0,X1) [cnf transformation 181]
205. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 182]
206. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 182]
207. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 183]
208. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 183]
210. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 184]
212. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 188]
214. ~select_7(X0) | where_5_subquery_6(X0) [cnf transformation 189]
216. ~where_5_subquery_6(X0) | where_5(X0) [cnf transformation 190]
217. from_4(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_3(X0,X1) [cnf transformation 194]
219. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0,X1) [cnf transformation 194]
220. ~main_query_2(X0,X1) | select_3(X0,X1) [cnf transformation 195]
221. ~select_3(X0,X1) | main_query_2(X0,X1) [cnf transformation 195]
263. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 200]
264. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 200]
266. 0 <=> main_query_1(sK10,sK11) [avatar definition]
267. main_query_1(sK10,sK11) <- {0} [avatar component clause 266]
269. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
270. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 269]
272. 2 <=> main_query_2(sK10,sK11) [avatar definition]
273. main_query_2(sK10,sK11) <- {2} [avatar component clause 272]
275. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
276. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 275]
277. ~1 | ~3 [avatar split clause 264,275,269]
278. 0 | 2 [avatar split clause 263,272,266]
279. select_3(sK10,sK11) <- {2} [resolution 220,273]
302. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 203,201]
307. tab1(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_3(X0,X1) [resolution 217,207]
349. from_4(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 302,208]
350. from_8(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 302,210]
356. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_3(X0,X1) [resolution 307,202]
483. ~select_1(X2,X3) | ~where_5(X2) | select_3(X2,X3) [resolution 349,219]
505. ~select_1(X2,X3) | select_7(X2) [resolution 350,212]
527. ~select_3(X2,X3) | select_1(X2,X3) [resolution 356,204]
550. select_1(sK10,sK11) <- {2} [resolution 527,279]
562. main_query_1(sK10,sK11) <- {2} [resolution 550,206]
563. $false <- {1, 2} [subsumption resolution 562,270]
564. 1 | ~2 [avatar contradiction clause 563]
565. select_1(sK10,sK11) <- {0} [resolution 267,205]
566. select_7(sK10) <- {0} [resolution 565,505]
567. ~where_5(sK10) | select_3(sK10,sK11) <- {0} [resolution 565,483]
573. 4 <=> select_3(sK10,sK11) [avatar definition]
574. select_3(sK10,sK11) <- {4} [avatar component clause 573]
576. 6 <=> where_5(sK10) [avatar definition]
579. 7 <=> ~where_5(sK10) [avatar definition]
581. 4 | ~7 | ~0 [avatar split clause 567,266,579,573]
582. where_5_subquery_6(sK10) <- {0} [resolution 566,214]
599. where_5(sK10) <- {0} [resolution 582,216]
603. 6 | ~0 [avatar split clause 599,266,576]
623. main_query_2(sK10,sK11) <- {4} [resolution 574,221]
625. $false <- {3, 4} [subsumption resolution 623,276]
626. 3 | ~4 [avatar contradiction clause 625]
627. $false [avatar sat refutation 277,278,564,581,603,626]
% SZS output end Proof for query_01_query_09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.016 s
% ------------------------------
% ------------------------------
