% Refutation found. Thanks to Tanya!
% SZS status Theorem for query21_query22_tptp4X
% SZS output start Proof for query21_query22_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X5,X6] : (where_6(X5,X6) <=> eq(X5,X6)) [input]
3. ! [X5,X6] : (select_4(X5,X6) <=> ? [X0,X1,X2,X3,X4] : (where_6(X5,X6) & from_5(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (subquery_from_3(X5,X6) <=> select_4(X5,X6)) [input]
5. ! [X5,X7] : (from_2(X5,X7) <=> subquery_from_3(X5,X7)) [input]
6. ! [X5] : (select_1(X5) <=> ? [X7] : ($true & from_2(X5,X7))) [input]
7. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
8. ! [X0,X1,X2,X3,X4] : (from_11(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
9. ! [X8,X6] : (where_12(X8,X6) <=> eq(X8,X6)) [input]
10. ! [X8,X6] : (select_10(X8,X6) <=> ? [X0,X1,X2,X3,X4] : (where_12(X8,X6) & from_11(X0,X1,X2,X3,X4))) [input]
11. ! [X8,X6] : (subquery_from_9(X8,X6) <=> select_10(X8,X6)) [input]
12. ! [X8,X7] : (from_8(X8,X7) <=> subquery_from_9(X8,X7)) [input]
13. ! [X8] : (select_7(X8) <=> ? [X7] : ($true & from_8(X8,X7))) [input]
14. ! [X8] : (main_query_2(X8) <=> select_7(X8)) [input]
15. ! [X9] : eq(X9,X9) [input]
64. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
65. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 64]
66. ! [X0,X1] : (where_6(X0,X1) <=> eq(X0,X1)) [rectify 2]
67. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4,X5,X6] : (where_6(X0,X1) & from_5(X2,X3,X4,X5,X6))) [rectify 3]
68. ! [X0,X1] : (subquery_from_3(X0,X1) <=> select_4(X0,X1)) [rectify 4]
69. ! [X0,X1] : (from_2(X0,X1) <=> subquery_from_3(X0,X1)) [rectify 5]
70. ! [X0] : (select_1(X0) <=> ? [X1] : ($true & from_2(X0,X1))) [rectify 6]
71. ! [X0] : (select_1(X0) <=> ? [X1] : from_2(X0,X1)) [true and false elimination 70]
72. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 7]
73. ! [X0,X1] : (where_12(X0,X1) <=> eq(X0,X1)) [rectify 9]
74. ! [X0,X1] : (select_10(X0,X1) <=> ? [X2,X3,X4,X5,X6] : (where_12(X0,X1) & from_11(X2,X3,X4,X5,X6))) [rectify 10]
75. ! [X0,X1] : (subquery_from_9(X0,X1) <=> select_10(X0,X1)) [rectify 11]
76. ! [X0,X1] : (from_8(X0,X1) <=> subquery_from_9(X0,X1)) [rectify 12]
77. ! [X0] : (select_7(X0) <=> ? [X1] : ($true & from_8(X0,X1))) [rectify 13]
78. ! [X0] : (select_7(X0) <=> ? [X1] : from_8(X0,X1)) [true and false elimination 77]
79. ! [X0] : (main_query_2(X0) <=> select_7(X0)) [rectify 14]
80. ! [X0] : eq(X0,X0) [rectify 15]
129. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 65]
213. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 129]
214. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 1]
215. ! [X0,X1] : ((where_6(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_6(X0,X1))) [nnf transformation 66]
216. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_6(X0,X1) | ~from_5(X2,X3,X4,X5,X6))) & (? [X2,X3,X4,X5,X6] : (where_6(X0,X1) & from_5(X2,X3,X4,X5,X6)) | ~select_4(X0,X1))) [nnf transformation 67]
217. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_6(X0,X1) | ~from_5(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_6(X0,X1) & from_5(X7,X8,X9,X10,X11)) | ~select_4(X0,X1))) [rectify 216]
218. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_6(X0,X1) & from_5(X7,X8,X9,X10,X11)) => (where_6(X0,X1) & from_5(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)))) [choice axiom]
219. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_6(X0,X1) | ~from_5(X2,X3,X4,X5,X6))) & ((where_6(X0,X1) & from_5(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1))) | ~select_4(X0,X1))) [skolemisation 217,218]
220. ! [X0,X1] : ((subquery_from_3(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~subquery_from_3(X0,X1))) [nnf transformation 68]
221. ! [X0,X1] : ((from_2(X0,X1) | ~subquery_from_3(X0,X1)) & (subquery_from_3(X0,X1) | ~from_2(X0,X1))) [nnf transformation 69]
222. ! [X0] : ((select_1(X0) | ! [X1] : ~from_2(X0,X1)) & (? [X1] : from_2(X0,X1) | ~select_1(X0))) [nnf transformation 71]
223. ! [X0] : ((select_1(X0) | ! [X1] : ~from_2(X0,X1)) & (? [X2] : from_2(X0,X2) | ~select_1(X0))) [rectify 222]
224. ! [X0] : (? [X2] : from_2(X0,X2) => from_2(X0,sK5(X0))) [choice axiom]
225. ! [X0] : ((select_1(X0) | ! [X1] : ~from_2(X0,X1)) & (from_2(X0,sK5(X0)) | ~select_1(X0))) [skolemisation 223,224]
226. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 72]
227. ! [X0,X1,X2,X3,X4] : ((from_11(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_11(X0,X1,X2,X3,X4))) [nnf transformation 8]
228. ! [X0,X1] : ((where_12(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_12(X0,X1))) [nnf transformation 73]
229. ! [X0,X1] : ((select_10(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_12(X0,X1) | ~from_11(X2,X3,X4,X5,X6))) & (? [X2,X3,X4,X5,X6] : (where_12(X0,X1) & from_11(X2,X3,X4,X5,X6)) | ~select_10(X0,X1))) [nnf transformation 74]
230. ! [X0,X1] : ((select_10(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_12(X0,X1) | ~from_11(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_12(X0,X1) & from_11(X7,X8,X9,X10,X11)) | ~select_10(X0,X1))) [rectify 229]
231. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_12(X0,X1) & from_11(X7,X8,X9,X10,X11)) => (where_12(X0,X1) & from_11(sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)))) [choice axiom]
232. ! [X0,X1] : ((select_10(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_12(X0,X1) | ~from_11(X2,X3,X4,X5,X6))) & ((where_12(X0,X1) & from_11(sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1))) | ~select_10(X0,X1))) [skolemisation 230,231]
233. ! [X0,X1] : ((subquery_from_9(X0,X1) | ~select_10(X0,X1)) & (select_10(X0,X1) | ~subquery_from_9(X0,X1))) [nnf transformation 75]
234. ! [X0,X1] : ((from_8(X0,X1) | ~subquery_from_9(X0,X1)) & (subquery_from_9(X0,X1) | ~from_8(X0,X1))) [nnf transformation 76]
235. ! [X0] : ((select_7(X0) | ! [X1] : ~from_8(X0,X1)) & (? [X1] : from_8(X0,X1) | ~select_7(X0))) [nnf transformation 78]
236. ! [X0] : ((select_7(X0) | ! [X1] : ~from_8(X0,X1)) & (? [X2] : from_8(X0,X2) | ~select_7(X0))) [rectify 235]
237. ! [X0] : (? [X2] : from_8(X0,X2) => from_8(X0,sK11(X0))) [choice axiom]
238. ! [X0] : ((select_7(X0) | ! [X1] : ~from_8(X0,X1)) & (from_8(X0,sK11(X0)) | ~select_7(X0))) [skolemisation 236,237]
239. ! [X0] : ((main_query_2(X0) | ~select_7(X0)) & (select_7(X0) | ~main_query_2(X0))) [nnf transformation 79]
242. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 213]
243. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK12) | ~main_query_1(sK12)) & (main_query_2(sK12) | main_query_1(sK12))) [choice axiom]
244. (~main_query_2(sK12) | ~main_query_1(sK12)) & (main_query_2(sK12) | main_query_1(sK12)) [skolemisation 242,243]
245. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 214]
246. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 214]
248. ~eq(X0,X1) | where_6(X0,X1) [cnf transformation 215]
249. from_5(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_4(X0,X1) [cnf transformation 219]
251. select_4(X0,X1) | ~where_6(X0,X1) | ~from_5(X2,X3,X4,X5,X6) [cnf transformation 219]
252. ~subquery_from_3(X0,X1) | select_4(X0,X1) [cnf transformation 220]
253. ~select_4(X0,X1) | subquery_from_3(X0,X1) [cnf transformation 220]
254. ~from_2(X0,X1) | subquery_from_3(X0,X1) [cnf transformation 221]
255. ~subquery_from_3(X0,X1) | from_2(X0,X1) [cnf transformation 221]
256. from_2(X0,sK5(X0)) | ~select_1(X0) [cnf transformation 225]
257. ~from_2(X0,X1) | select_1(X0) [cnf transformation 225]
258. ~main_query_1(X0) | select_1(X0) [cnf transformation 226]
259. ~select_1(X0) | main_query_1(X0) [cnf transformation 226]
260. ~from_11(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 227]
261. ~tab1(X0,X1,X2,X3,X4) | from_11(X0,X1,X2,X3,X4) [cnf transformation 227]
263. ~eq(X0,X1) | where_12(X0,X1) [cnf transformation 228]
264. from_11(sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) | ~select_10(X0,X1) [cnf transformation 232]
266. select_10(X0,X1) | ~where_12(X0,X1) | ~from_11(X2,X3,X4,X5,X6) [cnf transformation 232]
267. ~subquery_from_9(X0,X1) | select_10(X0,X1) [cnf transformation 233]
268. ~select_10(X0,X1) | subquery_from_9(X0,X1) [cnf transformation 233]
269. ~from_8(X0,X1) | subquery_from_9(X0,X1) [cnf transformation 234]
270. ~subquery_from_9(X0,X1) | from_8(X0,X1) [cnf transformation 234]
271. from_8(X0,sK11(X0)) | ~select_7(X0) [cnf transformation 238]
272. ~from_8(X0,X1) | select_7(X0) [cnf transformation 238]
273. ~main_query_2(X0) | select_7(X0) [cnf transformation 239]
274. ~select_7(X0) | main_query_2(X0) [cnf transformation 239]
275. eq(X0,X0) [cnf transformation 80]
322. main_query_2(sK12) | main_query_1(sK12) [cnf transformation 244]
323. ~main_query_2(sK12) | ~main_query_1(sK12) [cnf transformation 244]
364. 0 <=> ! [X3,X5,X2,X4,X6] : ~from_5(X2,X3,X4,X5,X6) [avatar definition]
365. ~from_5(X2,X3,X4,X5,X6) <- {0} [avatar component clause 364]
367. 2 <=> ! [X1,X0] : (select_4(X0,X1) | ~where_6(X0,X1)) [avatar definition]
368. ~where_6(X0,X1) | select_4(X0,X1) <- {2} [avatar component clause 367]
369. 0 | 2 [avatar split clause 251,367,364]
371. 4 <=> ! [X3,X5,X2,X4,X6] : ~from_11(X2,X3,X4,X5,X6) [avatar definition]
372. ~from_11(X2,X3,X4,X5,X6) <- {4} [avatar component clause 371]
374. 6 <=> ! [X1,X0] : (select_10(X0,X1) | ~where_12(X0,X1)) [avatar definition]
375. ~where_12(X0,X1) | select_10(X0,X1) <- {6} [avatar component clause 374]
376. 4 | 6 [avatar split clause 266,374,371]
378. 8 <=> main_query_1(sK12) [avatar definition]
379. main_query_1(sK12) <- {8} [avatar component clause 378]
381. 9 <=> ~main_query_1(sK12) [avatar definition]
382. ~main_query_1(sK12) <- {9} [avatar component clause 381]
384. 10 <=> main_query_2(sK12) [avatar definition]
385. main_query_2(sK12) <- {10} [avatar component clause 384]
387. 11 <=> ~main_query_2(sK12) [avatar definition]
388. ~main_query_2(sK12) <- {11} [avatar component clause 387]
389. ~9 | ~11 [avatar split clause 323,387,381]
390. 8 | 10 [avatar split clause 322,384,378]
391. select_1(sK12) <- {8} [resolution 258,379]
393. where_6(X0,X0) [resolution 248,275]
395. subquery_from_3(X0,sK5(X0)) | ~select_1(X0) [resolution 256,254]
397. where_12(X0,X0) [resolution 263,275]
399. subquery_from_9(X0,sK11(X0)) | ~select_7(X0) [resolution 271,269]
417. select_4(X1,sK5(X1)) | ~select_1(X1) [resolution 395,252]
423. tab1(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_4(X0,X1) [resolution 249,245]
426. tab1(sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) | ~select_10(X0,X1) [resolution 264,260]
427. ~select_10(X2,X3) <- {4} [resolution 264,372]
429. select_10(X1,sK11(X1)) | ~select_7(X1) [resolution 399,267]
430. ~select_7(X1) <- {4} [subsumption resolution 429,427]
431. from_11(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_4(X0,X1) [resolution 423,261]
433. ~select_4(X0,X1) <- {4} [subsumption resolution 431,372]
438. ~select_1(X1) <- {4} [subsumption resolution 417,433]
440. $false <- {4, 8} [resolution 438,391]
441. ~4 | ~8 [avatar contradiction clause 440]
442. select_7(sK12) <- {10} [resolution 385,273]
443. $false <- {4, 10} [subsumption resolution 442,430]
444. ~4 | ~10 [avatar contradiction clause 443]
446. select_10(X0,X0) <- {6} [resolution 375,397]
449. from_5(sK6(X2,X3),sK7(X2,X3),sK8(X2,X3),sK9(X2,X3),sK10(X2,X3)) | ~select_10(X2,X3) [resolution 426,246]
450. ~select_10(X2,X3) <- {0} [subsumption resolution 449,365]
451. subquery_from_9(X0,X0) <- {6} [resolution 446,268]
453. $false <- {0, 6} [resolution 450,446]
454. ~0 | ~6 [avatar contradiction clause 453]
455. from_8(X0,X0) <- {6} [resolution 451,270]
457. select_4(X0,X0) <- {2} [resolution 368,393]
458. subquery_from_3(X0,X0) <- {2} [resolution 457,253]
462. from_2(X0,X0) <- {2} [resolution 458,255]
467. select_7(X1) <- {6} [resolution 455,272]
469. main_query_2(X0) <- {6} [resolution 467,274]
475. select_1(X1) <- {2} [resolution 462,257]
479. main_query_1(X0) <- {2} [resolution 475,259]
487. $false <- {2, 9} [resolution 479,382]
488. ~2 | 9 [avatar contradiction clause 487]
489. $false <- {6, 11} [subsumption resolution 388,469]
490. ~6 | 11 [avatar contradiction clause 489]
491. $false [avatar sat refutation 369,376,389,390,441,444,454,488,490]
% SZS output end Proof for query21_query22_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.097 s
% ------------------------------
% ------------------------------
