% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_08_query_10_tptp4X
% SZS output start Proof for query_08_query_10_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X0,X5,X6,X7,X8] : (from_8(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
6. ! [X0] : (select_7(X0) <=> ? [X5,X6,X7,X8] : ($true & from_8(X0,X5,X6,X7,X8))) [input]
7. ! [X0] : (subquery_from_6(X0) <=> select_7(X0)) [input]
8. ! [X9] : (from_5(X9) <=> subquery_from_6(X9)) [input]
9. ! [X9] : (select_4(X9) <=> ($true & from_5(X9))) [input]
10. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [input]
11. ! [X10] : eq(X10,X10) [input]
18. ! [X10,X11] : (lt(X10,X11) <=> (~eq(X10,X11) & lte(X10,X11))) [input]
19. ! [X10,X11] : (gt(X10,X11) <=> lt(X11,X10)) [input]
42. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
43. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 42]
44. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
45. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
46. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 45]
47. ! [X0] : (from_5(X0) <=> subquery_from_6(X0)) [rectify 8]
48. ! [X0] : (select_4(X0) <=> ($true & from_5(X0))) [rectify 9]
49. ! [X0] : (select_4(X0) <=> from_5(X0)) [true and false elimination 48]
50. ! [X0] : eq(X0,X0) [rectify 11]
57. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 18]
58. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 19]
126. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 43]
127. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
128. ! [X0] : ((where_3(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
129. ! [X0] : ((where_3(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [flattening 128]
130. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
131. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 130]
132. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
133. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 131,132]
134. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
135. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 44]
136. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 46]
137. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 136]
138. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
139. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0))) [skolemisation 137,138]
140. ! [X0] : ((subquery_from_6(X0) | ~select_7(X0)) & (select_7(X0) | ~subquery_from_6(X0))) [nnf transformation 7]
141. ! [X0] : ((from_5(X0) | ~subquery_from_6(X0)) & (subquery_from_6(X0) | ~from_5(X0))) [nnf transformation 47]
142. ! [X0] : ((select_4(X0) | ~from_5(X0)) & (from_5(X0) | ~select_4(X0))) [nnf transformation 49]
143. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 10]
146. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 57]
147. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 146]
148. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 58]
149. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 126]
150. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
151. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 149,150]
152. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 127]
153. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 127]
156. gt(X0,X0) | where_3(X0) | lt(X0,X0) [cnf transformation 129]
157. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 133]
159. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 133]
160. ~main_query_1(X0) | select_1(X0) [cnf transformation 134]
161. ~select_1(X0) | main_query_1(X0) [cnf transformation 134]
162. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 135]
163. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 135]
164. from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [cnf transformation 139]
165. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 139]
166. ~subquery_from_6(X0) | select_7(X0) [cnf transformation 140]
167. ~select_7(X0) | subquery_from_6(X0) [cnf transformation 140]
168. ~from_5(X0) | subquery_from_6(X0) [cnf transformation 141]
169. ~subquery_from_6(X0) | from_5(X0) [cnf transformation 141]
170. ~select_4(X0) | from_5(X0) [cnf transformation 142]
171. ~from_5(X0) | select_4(X0) [cnf transformation 142]
172. ~main_query_2(X0) | select_4(X0) [cnf transformation 143]
173. ~select_4(X0) | main_query_2(X0) [cnf transformation 143]
174. eq(X0,X0) [cnf transformation 50]
184. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 147]
186. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 148]
208. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 151]
209. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 151]
211. 0 <=> main_query_1(sK8) [avatar definition]
212. main_query_1(sK8) <- {0} [avatar component clause 211]
214. 1 <=> ~main_query_1(sK8) [avatar definition]
215. ~main_query_1(sK8) <- {1} [avatar component clause 214]
217. 2 <=> main_query_2(sK8) [avatar definition]
218. main_query_2(sK8) <- {2} [avatar component clause 217]
220. 3 <=> ~main_query_2(sK8) [avatar definition]
221. ~main_query_2(sK8) <- {3} [avatar component clause 220]
222. ~1 | ~3 [avatar split clause 209,220,214]
223. 0 | 2 [avatar split clause 208,217,211]
224. select_4(sK8) <- {2} [resolution 172,218]
225. from_5(sK8) <- {2} [resolution 224,170]
228. subquery_from_6(sK8) <- {2} [resolution 225,168]
230. select_7(sK8) <- {2} [resolution 228,166]
236. ~lt(X0,X0) [resolution 184,174]
238. where_3(X1) | lt(X1,X1) | lt(X1,X1) [resolution 156,186]
239. where_3(X1) | lt(X1,X1) [duplicate literal removal 238]
240. where_3(X1) [subsumption resolution 239,236]
264. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 157,152]
306. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [resolution 164,162]
425. from_8(X1,sK0(X1),sK1(X1),sK2(X1),sK3(X1)) | ~select_1(X1) [resolution 264,163]
431. from_2(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [resolution 306,153]
439. ~select_1(X1) | select_7(X1) [resolution 425,165]
441. ~select_7(X1) | ~where_3(X1) | select_1(X1) [resolution 431,159]
447. ~select_7(X1) | select_1(X1) [subsumption resolution 441,240]
450. select_1(sK8) <- {2} [resolution 447,230]
454. main_query_1(sK8) <- {2} [resolution 450,161]
456. $false <- {1, 2} [subsumption resolution 454,215]
457. 1 | ~2 [avatar contradiction clause 456]
458. select_1(sK8) <- {0} [resolution 212,160]
461. select_7(sK8) <- {0} [resolution 458,439]
467. subquery_from_6(sK8) <- {0} [resolution 461,167]
470. from_5(sK8) <- {0} [resolution 467,169]
480. select_4(sK8) <- {0} [resolution 470,171]
489. main_query_2(sK8) <- {0} [resolution 480,173]
491. $false <- {0, 3} [subsumption resolution 489,221]
492. ~0 | 3 [avatar contradiction clause 491]
493. $false [avatar sat refutation 222,223,457,492]
% SZS output end Proof for query_08_query_10_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.032 s
% ------------------------------
% ------------------------------
