% Refutation found. Thanks to Tanya!
% SZS status Theorem for query05_query08_tptp4X
% SZS output start Proof for query05_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
7. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
9. ! [X6] : eq(X6,X6) [input]
12. ! [X6] : lte(X6,X6) [input]
18. ! [X6,X7] : (gte(X6,X7) <=> lte(X7,X6)) [input]
48. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
49. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 48]
50. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
51. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
52. ! [X0] : eq(X0,X0) [rectify 9]
55. ! [X0] : lte(X0,X0) [rectify 12]
61. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 18]
91. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 49]
155. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 91]
156. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
157. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
158. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
159. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 158]
160. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
161. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 159,160]
162. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 50]
163. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
164. ! [X0] : ((where_6(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [nnf transformation 6]
165. ! [X0] : ((where_6(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [flattening 164]
166. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 7]
167. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 166]
168. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
169. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 167,168]
170. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 51]
173. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 61]
174. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 155]
175. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
176. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 174,175]
177. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 156]
178. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 156]
180. ~eq(X0,X0) | where_3(X0) [cnf transformation 157]
181. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 161]
183. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 161]
184. ~main_query_1(X0) | select_1(X0) [cnf transformation 162]
185. ~select_1(X0) | main_query_1(X0) [cnf transformation 162]
186. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 163]
187. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 163]
190. ~gte(X0,X0) | where_6(X0) | ~lte(X0,X0) [cnf transformation 165]
191. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 169]
193. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 169]
194. ~main_query_2(X0) | select_4(X0) [cnf transformation 170]
195. ~select_4(X0) | main_query_2(X0) [cnf transformation 170]
196. eq(X0,X0) [cnf transformation 52]
199. lte(X0,X0) [cnf transformation 55]
205. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 173]
235. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 176]
236. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 176]
267. 0 <=> main_query_1(sK8) [avatar definition]
268. main_query_1(sK8) <- {0} [avatar component clause 267]
270. 1 <=> ~main_query_1(sK8) [avatar definition]
271. ~main_query_1(sK8) <- {1} [avatar component clause 270]
273. 2 <=> main_query_2(sK8) [avatar definition]
274. main_query_2(sK8) <- {2} [avatar component clause 273]
276. 3 <=> ~main_query_2(sK8) [avatar definition]
277. ~main_query_2(sK8) <- {3} [avatar component clause 276]
278. ~1 | ~3 [avatar split clause 236,276,270]
279. 0 | 2 [avatar split clause 235,273,267]
280. select_4(sK8) <- {2} [resolution 194,274]
284. where_3(X1) [resolution 180,196]
293. gte(X1,X1) [resolution 205,199]
295. where_6(X0) | ~lte(X0,X0) [resolution 190,293]
297. where_6(X0) [subsumption resolution 295,199]
304. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 181,177]
306. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 191,186]
308. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 304,187]
311. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 306,178]
313. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 308,193]
314. ~select_1(X1) | select_4(X1) [subsumption resolution 313,297]
316. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 311,183]
317. ~select_4(X1) | select_1(X1) [subsumption resolution 316,284]
318. select_1(sK8) <- {2} [resolution 317,280]
320. main_query_1(sK8) <- {2} [resolution 318,185]
322. $false <- {1, 2} [subsumption resolution 320,271]
323. 1 | ~2 [avatar contradiction clause 322]
324. select_1(sK8) <- {0} [resolution 268,184]
325. select_4(sK8) <- {0} [resolution 324,314]
329. main_query_2(sK8) <- {0} [resolution 325,195]
331. $false <- {0, 3} [subsumption resolution 329,277]
332. ~0 | 3 [avatar contradiction clause 331]
333. $false [avatar sat refutation 278,279,323,332]
% SZS output end Proof for query05_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
