% Refutation found. Thanks to Tanya!
% SZS status Theorem for query02_query09_tptp4X
% SZS output start Proof for query02_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X5,X1,X2,X3,X4] : (from_7(X5,X1,X2,X3,X4) <=> tab1(X5,X1,X2,X3,X4)) [input]
5. ! [X5] : (select_6(X5) <=> ? [X1,X2,X3,X4] : ($true & from_7(X5,X1,X2,X3,X4))) [input]
6. ! [X5] : (subquery_from_5(X5) <=> select_6(X5)) [input]
7. ! [X5] : (from_4(X5) <=> subquery_from_5(X5)) [input]
8. ! [X5] : (select_3(X5) <=> ($true & from_4(X5))) [input]
9. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
43. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
44. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 43]
45. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
46. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
47. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [rectify 5]
48. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 47]
49. ! [X0] : (subquery_from_5(X0) <=> select_6(X0)) [rectify 6]
50. ! [X0] : (from_4(X0) <=> subquery_from_5(X0)) [rectify 7]
51. ! [X0] : (select_3(X0) <=> ($true & from_4(X0))) [rectify 8]
52. ! [X0] : (select_3(X0) <=> from_4(X0)) [true and false elimination 51]
53. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 9]
138. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 44]
139. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
140. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 45]
141. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 140]
142. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
143. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 141,142]
144. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
145. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 46]
146. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 48]
147. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 146]
148. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
149. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 147,148]
150. ! [X0] : ((subquery_from_5(X0) | ~select_6(X0)) & (select_6(X0) | ~subquery_from_5(X0))) [nnf transformation 49]
151. ! [X0] : ((from_4(X0) | ~subquery_from_5(X0)) & (subquery_from_5(X0) | ~from_4(X0))) [nnf transformation 50]
152. ! [X0] : ((select_3(X0) | ~from_4(X0)) & (from_4(X0) | ~select_3(X0))) [nnf transformation 52]
153. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 53]
156. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 138]
157. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
158. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 156,157]
159. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 139]
160. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 139]
161. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 143]
162. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 143]
163. ~main_query_1(X0) | select_1(X0) [cnf transformation 144]
164. ~select_1(X0) | main_query_1(X0) [cnf transformation 144]
165. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 145]
166. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 145]
167. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 149]
168. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 149]
169. ~subquery_from_5(X0) | select_6(X0) [cnf transformation 150]
170. ~select_6(X0) | subquery_from_5(X0) [cnf transformation 150]
171. ~from_4(X0) | subquery_from_5(X0) [cnf transformation 151]
172. ~subquery_from_5(X0) | from_4(X0) [cnf transformation 151]
173. ~select_3(X0) | from_4(X0) [cnf transformation 152]
174. ~from_4(X0) | select_3(X0) [cnf transformation 152]
175. ~main_query_2(X0) | select_3(X0) [cnf transformation 153]
176. ~select_3(X0) | main_query_2(X0) [cnf transformation 153]
208. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 158]
209. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 158]
234. 0 <=> main_query_1(sK8) [avatar definition]
235. main_query_1(sK8) <- {0} [avatar component clause 234]
237. 1 <=> ~main_query_1(sK8) [avatar definition]
238. ~main_query_1(sK8) <- {1} [avatar component clause 237]
240. 2 <=> main_query_2(sK8) [avatar definition]
241. main_query_2(sK8) <- {2} [avatar component clause 240]
243. 3 <=> ~main_query_2(sK8) [avatar definition]
244. ~main_query_2(sK8) <- {3} [avatar component clause 243]
245. ~1 | ~3 [avatar split clause 209,243,237]
246. 0 | 2 [avatar split clause 208,240,234]
247. select_3(sK8) <- {2} [resolution 175,241]
249. from_4(sK8) <- {2} [resolution 247,173]
252. subquery_from_5(sK8) <- {2} [resolution 249,171]
255. select_6(sK8) <- {2} [resolution 252,169]
261. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 161,159]
263. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 167,165]
265. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 261,166]
268. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 263,160]
270. ~select_1(X1) | select_6(X1) [resolution 265,168]
272. ~select_6(X1) | select_1(X1) [resolution 268,162]
273. select_1(sK8) <- {2} [resolution 272,255]
275. main_query_1(sK8) <- {2} [resolution 273,164]
276. $false <- {1, 2} [subsumption resolution 275,238]
277. 1 | ~2 [avatar contradiction clause 276]
278. select_1(sK8) <- {0} [resolution 235,163]
279. select_6(sK8) <- {0} [resolution 278,270]
282. subquery_from_5(sK8) <- {0} [resolution 279,170]
283. from_4(sK8) <- {0} [resolution 282,172]
285. select_3(sK8) <- {0} [resolution 283,174]
287. main_query_2(sK8) <- {0} [resolution 285,176]
289. $false <- {0, 3} [subsumption resolution 287,244]
290. ~0 | 3 [avatar contradiction clause 289]
291. $false [avatar sat refutation 245,246,277,290]
% SZS output end Proof for query02_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
