% Refutation found. Thanks to Tanya!
% SZS status Theorem for query10_query11_tptp4X
% SZS output start Proof for query10_query11_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X1,X0] : (where_3(X1,X0) <=> eq(X0,X1)) [input]
3. ! [X1,X0] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X1,X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8] : (where_6(X7,X8) <=> eq(X7,X8)) [input]
7. ! [X7,X8] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X7,X8) & from_5(X7,X8,X9,X10,X11))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
54. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
55. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 54]
56. ! [X0,X1] : (where_3(X0,X1) <=> eq(X1,X0)) [rectify 2]
57. ! [X0,X1] : (select_1(X1,X0) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X1,X0,X2,X3,X4))) [rectify 3]
58. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
59. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
60. ! [X0,X1] : (where_6(X0,X1) <=> eq(X0,X1)) [rectify 6]
61. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
62. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
108. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 55]
184. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 108]
185. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
186. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X1,X0)) & (eq(X1,X0) | ~where_3(X0,X1))) [nnf transformation 56]
187. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X1,X0,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X1,X0,X2,X3,X4)) | ~select_1(X1,X0))) [nnf transformation 57]
188. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X1,X0,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X1,X0,X5,X6,X7)) | ~select_1(X1,X0))) [rectify 187]
189. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X1,X0,X5,X6,X7)) => (where_3(X0,X1) & from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
190. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X1,X0,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X1,X0))) [skolemisation 188,189]
191. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 58]
192. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 59]
193. ! [X0,X1] : ((where_6(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_6(X0,X1))) [nnf transformation 60]
194. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 61]
195. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 194]
196. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
197. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 195,196]
198. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 62]
201. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 184]
202. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
203. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 201,202]
204. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 185]
205. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 185]
206. ~where_3(X0,X1) | eq(X1,X0) [cnf transformation 186]
207. ~eq(X1,X0) | where_3(X0,X1) [cnf transformation 186]
208. from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X1,X0) [cnf transformation 190]
209. ~select_1(X1,X0) | where_3(X0,X1) [cnf transformation 190]
210. ~from_2(X1,X0,X2,X3,X4) | ~where_3(X0,X1) | select_1(X1,X0) [cnf transformation 190]
211. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 191]
212. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 191]
213. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 192]
214. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 192]
215. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 193]
216. ~eq(X0,X1) | where_6(X0,X1) [cnf transformation 193]
217. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 197]
218. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 197]
219. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 197]
220. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 198]
221. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 198]
265. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 203]
266. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 203]
303. 0 <=> main_query_1(sK6,sK7) [avatar definition]
304. main_query_1(sK6,sK7) <- {0} [avatar component clause 303]
306. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
307. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 306]
309. 2 <=> main_query_2(sK6,sK7) [avatar definition]
310. main_query_2(sK6,sK7) <- {2} [avatar component clause 309]
312. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
313. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 312]
314. ~1 | ~3 [avatar split clause 266,312,306]
315. 0 | 2 [avatar split clause 265,309,303]
320. select_4(sK6,sK7) <- {2} [resolution 220,310]
322. where_6(sK6,sK7) <- {2} [resolution 320,218]
324. eq(sK6,sK7) <- {2} [resolution 322,215]
329. where_3(sK7,sK6) <- {2} [resolution 324,207]
358. tab1(X2,X3,sK0(X3,X2),sK1(X3,X2),sK2(X3,X2)) | ~select_1(X2,X3) [resolution 208,204]
360. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 217,213]
361. from_5(X0,X1,sK0(X1,X0),sK1(X1,X0),sK2(X1,X0)) | ~select_1(X0,X1) [resolution 358,214]
364. from_2(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 360,205]
365. ~where_6(X0,X1) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 361,219]
383. ~select_4(X0,X1) | ~where_3(X1,X0) | select_1(X0,X1) [resolution 364,210]
385. ~where_3(sK7,sK6) | select_1(sK6,sK7) <- {2} [resolution 383,320]
386. select_1(sK6,sK7) <- {2} [subsumption resolution 385,329]
387. main_query_1(sK6,sK7) <- {2} [resolution 386,212]
389. $false <- {1, 2} [subsumption resolution 387,307]
390. 1 | ~2 [avatar contradiction clause 389]
391. select_1(sK6,sK7) <- {0} [resolution 304,211]
393. where_3(sK7,sK6) <- {0} [resolution 391,209]
394. eq(sK6,sK7) <- {0} [resolution 393,206]
399. where_6(sK6,sK7) <- {0} [resolution 394,216]
411. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 399,365]
413. select_4(sK6,sK7) <- {0} [subsumption resolution 411,391]
420. main_query_2(sK6,sK7) <- {0} [resolution 413,221]
422. $false <- {0, 3} [subsumption resolution 420,313]
423. ~0 | 3 [avatar contradiction clause 422]
424. $false [avatar sat refutation 314,315,390,423]
% SZS output end Proof for query10_query11_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.034 s
% ------------------------------
% ------------------------------
