% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_11_query_15_tptp4X
% SZS output start Proof for query_11_query_15_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (where_6_subquery_7(X5) <=> select_8(X5)) [input]
9. ! [X0,X1] : (where_6(X0,X1) <=> (where_6_subquery_7(X0) & eq(X0,X1))) [input]
10. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
59. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
60. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 59]
61. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
62. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 7]
63. ! [X0] : (where_6_subquery_7(X0) <=> select_8(X0)) [rectify 8]
64. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
112. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 60]
192. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 112]
193. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
194. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
195. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
196. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 195]
197. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
198. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 196,197]
199. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 61]
200. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
201. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
202. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 62]
203. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 202]
204. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
205. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 203,204]
206. ! [X0] : ((where_6_subquery_7(X0) | ~select_8(X0)) & (select_8(X0) | ~where_6_subquery_7(X0))) [nnf transformation 63]
207. ! [X0,X1] : ((where_6(X0,X1) | (~where_6_subquery_7(X0) | ~eq(X0,X1))) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 9]
208. ! [X0,X1] : ((where_6(X0,X1) | ~where_6_subquery_7(X0) | ~eq(X0,X1)) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 207]
209. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 10]
210. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 209]
211. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
212. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 210,211]
213. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 64]
216. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 192]
217. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
218. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 216,217]
219. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 193]
220. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 193]
221. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 194]
222. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 194]
223. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 198]
224. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 198]
225. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 198]
226. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 199]
227. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 199]
228. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 200]
229. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 200]
231. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 201]
233. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 205]
235. ~select_8(X0) | where_6_subquery_7(X0) [cnf transformation 206]
236. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 208]
238. ~eq(X0,X1) | ~where_6_subquery_7(X0) | where_6(X0,X1) [cnf transformation 208]
239. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 212]
240. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 212]
241. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 212]
242. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 213]
243. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 213]
289. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 218]
290. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 218]
292. 0 <=> main_query_1(sK10,sK11) [avatar definition]
293. main_query_1(sK10,sK11) <- {0} [avatar component clause 292]
295. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
296. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 295]
298. 2 <=> main_query_2(sK10,sK11) [avatar definition]
299. main_query_2(sK10,sK11) <- {2} [avatar component clause 298]
301. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
302. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 301]
303. ~1 | ~3 [avatar split clause 290,301,295]
304. 0 | 2 [avatar split clause 289,298,292]
371. select_1(sK10,sK11) <- {0} [resolution 293,226]
376. where_3(sK10,sK11) <- {0} [resolution 371,224]
378. eq(sK10,sK11) <- {0} [resolution 376,221]
381. ~where_6_subquery_7(sK10) | where_6(sK10,sK11) <- {0} [resolution 378,238]
390. 8 <=> where_6(sK10,sK11) [avatar definition]
391. where_6(sK10,sK11) <- {8} [avatar component clause 390]
393. 10 <=> where_6_subquery_7(sK10) [avatar definition]
396. 11 <=> ~where_6_subquery_7(sK10) [avatar definition]
398. 8 | ~11 | ~0 [avatar split clause 381,292,396,390]
439. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 223,219]
443. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 239,228]
527. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 439,229]
528. from_9(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 439,231]
534. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 443,220]
566. 14 <=> select_8(sK10) [avatar definition]
567. select_8(sK10) <- {14} [avatar component clause 566]
677. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 527,241]
708. ~select_1(X2,X3) | select_8(X2) [resolution 528,233]
721. select_8(sK10) <- {0} [resolution 708,371]
736. 14 | ~0 [avatar split clause 721,292,566]
757. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 534,225]
763. where_6_subquery_7(sK10) <- {14} [resolution 567,235]
766. 10 | ~14 [avatar split clause 763,566,393]
799. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {8} [resolution 391,677]
804. select_4(sK10,sK11) <- {0, 8} [subsumption resolution 799,371]
829. main_query_2(sK10,sK11) <- {0, 8} [resolution 804,243]
831. $false <- {0, 3, 8} [subsumption resolution 829,302]
832. ~0 | 3 | ~8 [avatar contradiction clause 831]
848. select_4(sK10,sK11) <- {2} [resolution 299,242]
852. where_6(sK10,sK11) <- {2} [resolution 848,240]
853. 8 | ~2 [avatar split clause 852,298,390]
857. eq(sK10,sK11) <- {8} [resolution 391,236]
860. where_3(sK10,sK11) <- {8} [resolution 857,222]
1215. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 757,848]
1219. select_1(sK10,sK11) <- {2, 8} [subsumption resolution 1215,860]
1231. main_query_1(sK10,sK11) <- {2, 8} [resolution 1219,227]
1237. $false <- {1, 2, 8} [subsumption resolution 1231,296]
1238. 1 | ~2 | ~8 [avatar contradiction clause 1237]
1239. $false [avatar sat refutation 303,304,398,736,766,832,853,1238]
% SZS output end Proof for query_11_query_15_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.016 s
% ------------------------------
% ------------------------------
