% Refutation found. Thanks to Tanya!
% SZS status Theorem for query01_query03_tptp4X
% SZS output start Proof for query01_query03_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X5,X6,X7,X8] : (from_4(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
5. ! [X0] : (select_3(X0) <=> ? [X5,X6,X7,X8] : ($true & from_4(X0,X5,X6,X7,X8))) [input]
6. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
40. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
41. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 40]
42. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
43. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
44. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [rectify 5]
45. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 44]
130. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 41]
131. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
132. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 42]
133. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 132]
134. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 133,134]
136. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
137. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 43]
138. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 45]
139. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 138]
140. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
141. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0))) [skolemisation 139,140]
142. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 6]
145. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 130]
146. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
147. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 145,146]
148. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 131]
149. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 131]
150. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 135]
151. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 135]
152. ~main_query_1(X0) | select_1(X0) [cnf transformation 136]
153. ~select_1(X0) | main_query_1(X0) [cnf transformation 136]
154. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 137]
155. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 137]
156. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 141]
157. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 141]
158. ~main_query_2(X0) | select_3(X0) [cnf transformation 142]
159. ~select_3(X0) | main_query_2(X0) [cnf transformation 142]
191. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 147]
192. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 147]
217. 0 <=> main_query_1(sK8) [avatar definition]
218. main_query_1(sK8) <- {0} [avatar component clause 217]
220. 1 <=> ~main_query_1(sK8) [avatar definition]
221. ~main_query_1(sK8) <- {1} [avatar component clause 220]
223. 2 <=> main_query_2(sK8) [avatar definition]
224. main_query_2(sK8) <- {2} [avatar component clause 223]
226. 3 <=> ~main_query_2(sK8) [avatar definition]
227. ~main_query_2(sK8) <- {3} [avatar component clause 226]
228. ~1 | ~3 [avatar split clause 192,226,220]
229. 0 | 2 [avatar split clause 191,223,217]
230. select_3(sK8) <- {2} [resolution 158,224]
238. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 150,148]
240. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 156,154]
242. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 238,155]
245. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 240,149]
247. ~select_1(X1) | select_3(X1) [resolution 242,157]
249. ~select_3(X1) | select_1(X1) [resolution 245,151]
250. select_1(sK8) <- {2} [resolution 249,230]
252. main_query_1(sK8) <- {2} [resolution 250,153]
253. $false <- {1, 2} [subsumption resolution 252,221]
254. 1 | ~2 [avatar contradiction clause 253]
255. select_1(sK8) <- {0} [resolution 218,152]
256. select_3(sK8) <- {0} [resolution 255,247]
259. main_query_2(sK8) <- {0} [resolution 256,159]
260. $false <- {0, 3} [subsumption resolution 259,227]
261. ~0 | 3 [avatar contradiction clause 260]
262. $false [avatar sat refutation 228,229,254,261]
% SZS output end Proof for query01_query03_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
