% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_04_query_05_tptp4X
% SZS output start Proof for query_04_query_05_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [input]
4. ! [X5,X6,X7,X8,X9] : (from_4(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
5. ! [X5] : (where_5(X5) <=> eq(X5,X5)) [input]
6. ! [X5,X6] : (select_3(X5,X6) <=> ? [X7,X8,X9] : (where_5(X5) & from_4(X5,X6,X7,X8,X9))) [input]
7. ! [X10,X11] : (main_query_2(X10,X11) <=> select_3(X10,X11)) [input]
8. ! [X12] : eq(X12,X12) [input]
46. ! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [input]
47. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [negated conjecture 46]
48. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
50. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [rectify 5]
51. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
52. ! [X0,X1] : (main_query_2(X0,X1) <=> select_3(X0,X1)) [rectify 7]
53. ! [X0] : eq(X0,X0) [rectify 8]
152. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 47]
153. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
154. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0,X1))) [nnf transformation 48]
155. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) | ~select_1(X0,X1))) [rectify 154]
156. ! [X1,X0] : (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) => from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) [choice axiom]
157. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1))) [skolemisation 155,156]
158. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 3]
159. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 49]
160. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 50]
161. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0,X1))) [nnf transformation 51]
162. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) | ~select_3(X0,X1))) [rectify 161]
163. ! [X1,X0] : (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) => (where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
164. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_3(X0,X1))) [skolemisation 162,163]
165. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 52]
168. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 152]
169. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
170. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 168,169]
171. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 153]
172. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 153]
173. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 157]
174. ~from_2(X0,X1,X2,X3,X4) | select_1(X0,X1) [cnf transformation 157]
175. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 158]
176. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 158]
177. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 159]
178. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 159]
180. ~eq(X0,X0) | where_5(X0) [cnf transformation 160]
181. from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [cnf transformation 164]
183. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0,X1) [cnf transformation 164]
184. ~main_query_2(X0,X1) | select_3(X0,X1) [cnf transformation 165]
185. ~select_3(X0,X1) | main_query_2(X0,X1) [cnf transformation 165]
186. eq(X0,X0) [cnf transformation 53]
222. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 170]
223. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 170]
225. 0 <=> main_query_1(sK6,sK7) [avatar definition]
226. main_query_1(sK6,sK7) <- {0} [avatar component clause 225]
228. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
229. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 228]
231. 2 <=> main_query_2(sK6,sK7) [avatar definition]
232. main_query_2(sK6,sK7) <- {2} [avatar component clause 231]
234. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
235. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 234]
236. ~1 | ~3 [avatar split clause 223,234,228]
237. 0 | 2 [avatar split clause 222,231,225]
239. where_5(X1) [resolution 180,186]
240. select_3(sK6,sK7) <- {2} [resolution 184,232]
273. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 173,171]
275. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 181,177]
288. from_4(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 273,178]
294. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 275,172]
309. ~select_1(X2,X3) | ~where_5(X2) | select_3(X2,X3) [resolution 288,183]
315. ~select_1(X2,X3) | select_3(X2,X3) [subsumption resolution 309,239]
331. ~select_3(X2,X3) | select_1(X2,X3) [resolution 294,174]
344. select_1(sK6,sK7) <- {2} [resolution 331,240]
353. main_query_1(sK6,sK7) <- {2} [resolution 344,176]
354. $false <- {1, 2} [subsumption resolution 353,229]
355. 1 | ~2 [avatar contradiction clause 354]
356. select_1(sK6,sK7) <- {0} [resolution 226,175]
357. select_3(sK6,sK7) <- {0} [resolution 356,315]
362. main_query_2(sK6,sK7) <- {0} [resolution 357,185]
364. $false <- {0, 3} [subsumption resolution 362,235]
365. ~0 | 3 [avatar contradiction clause 364]
366. $false [avatar sat refutation 236,237,355,365]
% SZS output end Proof for query_04_query_05_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
