% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_02_query_03_tptp4X
% SZS output start Proof for query_02_query_03_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X5,X6,X7,X8,X9] : (from_4(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
5. ! [X5] : (select_3(X5) <=> ? [X6,X7,X8,X9] : ($true & from_4(X5,X6,X7,X8,X9))) [input]
6. ! [X10] : (main_query_2(X10) <=> select_3(X10)) [input]
38. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
39. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 38]
40. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
41. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
42. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [rectify 5]
43. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 42]
44. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 6]
121. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 39]
122. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
123. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 40]
124. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 123]
125. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
126. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 124,125]
127. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
128. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 41]
129. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 43]
130. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 129]
131. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
132. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0))) [skolemisation 130,131]
133. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 44]
136. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 121]
137. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
138. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 136,137]
139. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 122]
140. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 122]
141. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 126]
142. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 126]
143. ~main_query_1(X0) | select_1(X0) [cnf transformation 127]
144. ~select_1(X0) | main_query_1(X0) [cnf transformation 127]
145. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 128]
146. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 128]
147. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 132]
148. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 132]
149. ~main_query_2(X0) | select_3(X0) [cnf transformation 133]
150. ~select_3(X0) | main_query_2(X0) [cnf transformation 133]
180. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 138]
181. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 138]
183. 0 <=> main_query_1(sK8) [avatar definition]
184. main_query_1(sK8) <- {0} [avatar component clause 183]
186. 1 <=> ~main_query_1(sK8) [avatar definition]
187. ~main_query_1(sK8) <- {1} [avatar component clause 186]
189. 2 <=> main_query_2(sK8) [avatar definition]
190. main_query_2(sK8) <- {2} [avatar component clause 189]
192. 3 <=> ~main_query_2(sK8) [avatar definition]
193. ~main_query_2(sK8) <- {3} [avatar component clause 192]
194. ~1 | ~3 [avatar split clause 181,192,186]
195. 0 | 2 [avatar split clause 180,189,183]
196. select_3(sK8) <- {2} [resolution 149,190]
227. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 141,139]
229. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 147,145]
231. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 227,146]
235. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 229,140]
240. ~select_1(X5) | select_3(X5) [resolution 231,148]
246. ~select_3(X1) | select_1(X1) [resolution 235,142]
249. select_1(sK8) <- {2} [resolution 246,196]
251. main_query_1(sK8) <- {2} [resolution 249,144]
252. $false <- {1, 2} [subsumption resolution 251,187]
253. 1 | ~2 [avatar contradiction clause 252]
254. select_1(sK8) <- {0} [resolution 184,143]
255. select_3(sK8) <- {0} [resolution 254,240]
258. main_query_2(sK8) <- {0} [resolution 255,150]
259. $false <- {0, 3} [subsumption resolution 258,193]
260. ~0 | 3 [avatar contradiction clause 259]
261. $false [avatar sat refutation 194,195,253,260]
% SZS output end Proof for query_02_query_03_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.022 s
% ------------------------------
% ------------------------------
