% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_13_query_15_tptp4X
% SZS output start Proof for query_13_query_15_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (where_6_subquery_7(X5) <=> select_8(X5)) [input]
9. ! [X0,X1] : (where_6(X0,X1) <=> (where_6_subquery_7(X0) & eq(X0,X1))) [input]
10. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
13. ! [X7,X8] : (eq(X7,X8) => eq(X8,X7)) [input]
16. ! [X7,X8] : ((lte(X8,X7) & lte(X7,X8)) <=> eq(X7,X8)) [input]
22. ! [X7,X5,X6] : ((select_4(X5,X6) & eq(X5,X7)) => select_4(X7,X6)) [input]
59. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
60. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 59]
61. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
62. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 7]
63. ! [X0] : (where_6_subquery_7(X0) <=> select_8(X0)) [rectify 8]
64. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
66. ! [X0,X1] : (eq(X0,X1) => eq(X1,X0)) [rectify 13]
69. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 16]
75. ! [X0,X1,X2] : ((select_4(X1,X2) & eq(X1,X0)) => select_4(X0,X2)) [rectify 22]
112. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 60]
113. ! [X0,X1] : (eq(X1,X0) | ~eq(X0,X1)) [ennf transformation 66]
118. ! [X0,X1,X2] : (select_4(X0,X2) | (~select_4(X1,X2) | ~eq(X1,X0))) [ennf transformation 75]
119. ! [X0,X1,X2] : (select_4(X0,X2) | ~select_4(X1,X2) | ~eq(X1,X0)) [flattening 118]
192. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 112]
193. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
194. ! [X0,X1] : ((where_3(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [nnf transformation 2]
195. ! [X0,X1] : ((where_3(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [flattening 194]
196. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
197. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 196]
198. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
199. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 197,198]
200. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 61]
201. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
202. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
203. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 62]
204. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 203]
205. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
206. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 204,205]
207. ! [X0] : ((where_6_subquery_7(X0) | ~select_8(X0)) & (select_8(X0) | ~where_6_subquery_7(X0))) [nnf transformation 63]
208. ! [X0,X1] : ((where_6(X0,X1) | (~where_6_subquery_7(X0) | ~eq(X0,X1))) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 9]
209. ! [X0,X1] : ((where_6(X0,X1) | ~where_6_subquery_7(X0) | ~eq(X0,X1)) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 208]
210. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 10]
211. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 210]
212. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
213. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 211,212]
214. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 64]
215. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 69]
216. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 215]
217. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 192]
218. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
219. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 217,218]
220. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 193]
221. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 193]
222. ~where_3(X0,X1) | lte(X0,X1) [cnf transformation 195]
223. ~where_3(X0,X1) | lte(X1,X0) [cnf transformation 195]
224. ~lte(X1,X0) | where_3(X0,X1) | ~lte(X0,X1) [cnf transformation 195]
225. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 199]
226. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 199]
227. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 199]
228. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 200]
229. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 200]
230. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 201]
231. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 201]
233. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 202]
235. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 206]
237. ~select_8(X0) | where_6_subquery_7(X0) [cnf transformation 207]
238. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 209]
239. ~where_6(X0,X1) | where_6_subquery_7(X0) [cnf transformation 209]
240. ~eq(X0,X1) | ~where_6_subquery_7(X0) | where_6(X0,X1) [cnf transformation 209]
241. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 213]
242. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 213]
243. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 213]
244. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 214]
245. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 214]
247. ~eq(X0,X1) | eq(X1,X0) [cnf transformation 113]
250. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 216]
251. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 216]
252. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 216]
254. ~select_4(X1,X2) | select_4(X0,X2) | ~eq(X1,X0) [cnf transformation 119]
291. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 219]
292. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 219]
294. 0 <=> main_query_1(sK10,sK11) [avatar definition]
295. main_query_1(sK10,sK11) <- {0} [avatar component clause 294]
297. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
298. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 297]
300. 2 <=> main_query_2(sK10,sK11) [avatar definition]
301. main_query_2(sK10,sK11) <- {2} [avatar component clause 300]
303. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
304. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 303]
305. ~1 | ~3 [avatar split clause 292,303,297]
306. 0 | 2 [avatar split clause 291,300,294]
307. select_4(sK10,sK11) <- {2} [resolution 244,301]
309. where_6(sK10,sK11) <- {2} [resolution 307,242]
311. eq(sK10,sK11) <- {2} [resolution 309,238]
319. eq(sK11,sK10) <- {2} [resolution 311,247]
336. ~where_6_subquery_7(sK11) | where_6(sK11,sK10) <- {2} [resolution 319,240]
344. 4 <=> where_6(sK11,sK10) [avatar definition]
345. where_6(sK11,sK10) <- {4} [avatar component clause 344]
350. 7 <=> ~where_6_subquery_7(sK11) [avatar definition]
351. ~where_6_subquery_7(sK11) <- {7} [avatar component clause 350]
352. 4 | ~7 | ~2 [avatar split clause 336,300,350,344]
356. ~eq(sK10,X0) | select_4(X0,sK11) <- {2} [resolution 254,307]
370. select_4(sK11,sK11) <- {2} [resolution 356,311]
378. where_6(sK11,sK11) <- {2} [resolution 370,242]
386. where_6_subquery_7(sK11) <- {2} [resolution 378,239]
387. $false <- {2, 7} [subsumption resolution 386,351]
388. ~2 | 7 [avatar contradiction clause 387]
389. select_1(sK10,sK11) <- {0} [resolution 295,228]
391. where_3(sK10,sK11) <- {0} [resolution 389,226]
393. lte(sK11,sK10) <- {0} [resolution 391,223]
394. lte(sK10,sK11) <- {0} [resolution 391,222]
397. eq(sK10,sK11) | ~lte(sK10,sK11) <- {0} [resolution 393,250]
399. eq(sK10,sK11) <- {0} [subsumption resolution 397,394]
407. ~where_6_subquery_7(sK10) | where_6(sK10,sK11) <- {0} [resolution 399,240]
424. 8 <=> where_6(sK10,sK11) [avatar definition]
425. where_6(sK10,sK11) <- {8} [avatar component clause 424]
427. 10 <=> where_6_subquery_7(sK10) [avatar definition]
430. 11 <=> ~where_6_subquery_7(sK10) [avatar definition]
432. 8 | ~11 | ~0 [avatar split clause 407,294,430,424]
461. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 225,220]
464. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 241,230]
538. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 461,231]
539. from_9(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 461,233]
545. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 464,221]
577. 14 <=> select_8(sK10) [avatar definition]
578. select_8(sK10) <- {14} [avatar component clause 577]
681. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 538,243]
712. ~select_1(X2,X3) | select_8(X2) [resolution 539,235]
725. select_8(sK10) <- {0} [resolution 712,389]
740. 14 | ~0 [avatar split clause 725,294,577]
762. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 545,227]
776. where_6_subquery_7(sK10) <- {14} [resolution 578,237]
779. 10 | ~14 [avatar split clause 776,577,427]
807. eq(sK11,sK10) <- {4} [resolution 345,238]
820. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {8} [resolution 425,681]
823. select_4(sK10,sK11) <- {0, 8} [subsumption resolution 820,389]
872. main_query_2(sK10,sK11) <- {0, 8} [resolution 823,245]
874. $false <- {0, 3, 8} [subsumption resolution 872,304]
875. ~0 | 3 | ~8 [avatar contradiction clause 874]
891. select_4(sK10,sK11) <- {2} [resolution 301,244]
905. lte(sK11,sK10) <- {4} [resolution 807,251]
906. lte(sK10,sK11) <- {4} [resolution 807,252]
973. where_3(sK10,sK11) | ~lte(sK10,sK11) <- {4} [resolution 905,224]
974. where_3(sK10,sK11) <- {4} [subsumption resolution 973,906]
1284. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 762,891]
1288. select_1(sK10,sK11) <- {2, 4} [subsumption resolution 1284,974]
1299. main_query_1(sK10,sK11) <- {2, 4} [resolution 1288,229]
1301. $false <- {1, 2, 4} [subsumption resolution 1299,298]
1302. 1 | ~2 | ~4 [avatar contradiction clause 1301]
1303. $false [avatar sat refutation 305,306,352,388,432,740,779,875,1302]
% SZS output end Proof for query_13_query_15_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.054 s
% ------------------------------
% ------------------------------
