% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_04_query_06_tptp4X
% SZS output start Proof for query_04_query_06_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [input]
4. ! [X5,X6,X7,X8,X9] : (from_4(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
5. ! [X5] : (where_5(X5) <=> (gte(X5,X5) & lte(X5,X5))) [input]
6. ! [X5,X6] : (select_3(X5,X6) <=> ? [X7,X8,X9] : (where_5(X5) & from_4(X5,X6,X7,X8,X9))) [input]
7. ! [X10,X11] : (main_query_2(X10,X11) <=> select_3(X10,X11)) [input]
11. ! [X12] : lte(X12,X12) [input]
17. ! [X12,X13] : (gte(X12,X13) <=> lte(X13,X12)) [input]
46. ! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [input]
47. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [negated conjecture 46]
48. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
50. ! [X0] : (where_5(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 5]
51. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
52. ! [X0,X1] : (main_query_2(X0,X1) <=> select_3(X0,X1)) [rectify 7]
56. ! [X0] : lte(X0,X0) [rectify 11]
62. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 17]
152. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 47]
153. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
154. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0,X1))) [nnf transformation 48]
155. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) | ~select_1(X0,X1))) [rectify 154]
156. ! [X1,X0] : (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) => from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) [choice axiom]
157. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1))) [skolemisation 155,156]
158. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 3]
159. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 49]
160. ! [X0] : ((where_5(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [nnf transformation 50]
161. ! [X0] : ((where_5(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [flattening 160]
162. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0,X1))) [nnf transformation 51]
163. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) | ~select_3(X0,X1))) [rectify 162]
164. ! [X1,X0] : (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) => (where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
165. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_3(X0,X1))) [skolemisation 163,164]
166. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 52]
169. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 62]
170. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 152]
171. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
172. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 170,171]
173. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 153]
174. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 153]
175. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 157]
176. ~from_2(X0,X1,X2,X3,X4) | select_1(X0,X1) [cnf transformation 157]
177. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 158]
178. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 158]
179. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 159]
180. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 159]
183. ~gte(X0,X0) | where_5(X0) | ~lte(X0,X0) [cnf transformation 161]
184. from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [cnf transformation 165]
186. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0,X1) [cnf transformation 165]
187. ~main_query_2(X0,X1) | select_3(X0,X1) [cnf transformation 166]
188. ~select_3(X0,X1) | main_query_2(X0,X1) [cnf transformation 166]
192. lte(X0,X0) [cnf transformation 56]
198. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 169]
227. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 172]
228. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 172]
230. 0 <=> main_query_1(sK6,sK7) [avatar definition]
231. main_query_1(sK6,sK7) <- {0} [avatar component clause 230]
233. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
234. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 233]
236. 2 <=> main_query_2(sK6,sK7) [avatar definition]
237. main_query_2(sK6,sK7) <- {2} [avatar component clause 236]
239. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
240. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 239]
241. ~1 | ~3 [avatar split clause 228,239,233]
242. 0 | 2 [avatar split clause 227,236,230]
243. select_3(sK6,sK7) <- {2} [resolution 187,237]
251. gte(X1,X1) [resolution 198,192]
253. where_5(X0) | ~lte(X0,X0) [resolution 183,251]
255. where_5(X0) [subsumption resolution 253,192]
269. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 175,173]
271. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 184,179]
284. from_4(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 269,180]
290. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 271,174]
298. ~select_1(X2,X3) | ~where_5(X2) | select_3(X2,X3) [resolution 284,186]
304. ~select_1(X2,X3) | select_3(X2,X3) [subsumption resolution 298,255]
320. ~select_3(X2,X3) | select_1(X2,X3) [resolution 290,176]
333. select_1(sK6,sK7) <- {2} [resolution 320,243]
342. main_query_1(sK6,sK7) <- {2} [resolution 333,178]
343. $false <- {1, 2} [subsumption resolution 342,234]
344. 1 | ~2 [avatar contradiction clause 343]
345. select_1(sK6,sK7) <- {0} [resolution 231,177]
346. select_3(sK6,sK7) <- {0} [resolution 345,304]
349. main_query_2(sK6,sK7) <- {0} [resolution 346,188]
351. $false <- {0, 3} [subsumption resolution 349,240]
352. ~0 | 3 [avatar contradiction clause 351]
353. $false [avatar sat refutation 241,242,344,352]
% SZS output end Proof for query_04_query_06_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
