% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query02_tptp4X
% SZS output start Proof for query00_query02_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_4(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (select_3(X6) <=> ? [X7,X8,X9,X10] : ($true & from_4(X6,X7,X8,X9,X10))) [input]
6. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
40. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
41. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 40]
42. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
43. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
44. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
45. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [rectify 5]
46. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 45]
47. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 6]
81. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 41]
133. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 81]
134. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 42]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 135]
137. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 136,137]
139. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 43]
140. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 44]
141. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 46]
142. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 141]
143. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
144. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0))) [skolemisation 142,143]
145. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 47]
148. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 133]
149. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
150. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 148,149]
151. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 134]
152. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 134]
153. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 138]
154. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 138]
155. ~main_query_1(X0) | select_1(X0) [cnf transformation 139]
156. ~select_1(X0) | main_query_1(X0) [cnf transformation 139]
157. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 140]
158. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 140]
159. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 144]
160. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 144]
161. ~main_query_2(X0) | select_3(X0) [cnf transformation 145]
162. ~select_3(X0) | main_query_2(X0) [cnf transformation 145]
194. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 150]
195. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 150]
220. 0 <=> main_query_1(sK8) [avatar definition]
221. main_query_1(sK8) <- {0} [avatar component clause 220]
223. 1 <=> ~main_query_1(sK8) [avatar definition]
224. ~main_query_1(sK8) <- {1} [avatar component clause 223]
226. 2 <=> main_query_2(sK8) [avatar definition]
227. main_query_2(sK8) <- {2} [avatar component clause 226]
229. 3 <=> ~main_query_2(sK8) [avatar definition]
230. ~main_query_2(sK8) <- {3} [avatar component clause 229]
231. ~1 | ~3 [avatar split clause 195,229,223]
232. 0 | 2 [avatar split clause 194,226,220]
233. select_3(sK8) <- {2} [resolution 161,227]
241. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 153,151]
243. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 159,157]
245. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 241,158]
248. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 243,152]
250. ~select_1(X1) | select_3(X1) [resolution 245,160]
252. ~select_3(X1) | select_1(X1) [resolution 248,154]
253. select_1(sK8) <- {2} [resolution 252,233]
255. main_query_1(sK8) <- {2} [resolution 253,156]
256. $false <- {1, 2} [subsumption resolution 255,224]
257. 1 | ~2 [avatar contradiction clause 256]
258. select_1(sK8) <- {0} [resolution 221,155]
259. select_3(sK8) <- {0} [resolution 258,250]
262. main_query_2(sK8) <- {0} [resolution 259,162]
263. $false <- {0, 3} [subsumption resolution 262,230]
264. ~0 | 3 [avatar contradiction clause 263]
265. $false [avatar sat refutation 231,232,257,264]
% SZS output end Proof for query00_query02_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.034 s
% ------------------------------
% ------------------------------
