% Refutation found. Thanks to Tanya!
% SZS status Theorem for query02_query08_tptp4X
% SZS output start Proof for query02_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X5,X6,X7,X8] : (from_4(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
5. ! [X0] : (where_5(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
6. ! [X0] : (select_3(X0) <=> ? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8))) [input]
7. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
11. ! [X9] : lte(X9,X9) [input]
17. ! [X9,X10] : (gte(X9,X10) <=> lte(X10,X9)) [input]
41. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
42. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
45. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
49. ! [X0] : lte(X0,X0) [rectify 11]
55. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 17]
130. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 42]
131. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
132. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
133. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 132]
134. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 133,134]
136. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
137. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 44]
138. ! [X0] : ((where_5(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [nnf transformation 5]
139. ! [X0] : ((where_5(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [flattening 138]
140. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 45]
141. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 140]
142. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
143. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 141,142]
144. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 7]
147. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 55]
148. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 130]
149. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
150. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 148,149]
151. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 131]
152. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 131]
153. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 135]
154. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 135]
155. ~main_query_1(X0) | select_1(X0) [cnf transformation 136]
156. ~select_1(X0) | main_query_1(X0) [cnf transformation 136]
157. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 137]
158. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 137]
161. ~gte(X0,X0) | where_5(X0) | ~lte(X0,X0) [cnf transformation 139]
162. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 143]
164. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 143]
165. ~main_query_2(X0) | select_3(X0) [cnf transformation 144]
166. ~select_3(X0) | main_query_2(X0) [cnf transformation 144]
170. lte(X0,X0) [cnf transformation 49]
176. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 147]
200. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 150]
201. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 150]
226. 0 <=> main_query_1(sK8) [avatar definition]
227. main_query_1(sK8) <- {0} [avatar component clause 226]
229. 1 <=> ~main_query_1(sK8) [avatar definition]
230. ~main_query_1(sK8) <- {1} [avatar component clause 229]
232. 2 <=> main_query_2(sK8) [avatar definition]
233. main_query_2(sK8) <- {2} [avatar component clause 232]
235. 3 <=> ~main_query_2(sK8) [avatar definition]
236. ~main_query_2(sK8) <- {3} [avatar component clause 235]
237. ~1 | ~3 [avatar split clause 201,235,229]
238. 0 | 2 [avatar split clause 200,232,226]
239. select_3(sK8) <- {2} [resolution 165,233]
247. gte(X1,X1) [resolution 176,170]
249. where_5(X0) | ~lte(X0,X0) [resolution 161,247]
251. where_5(X0) [subsumption resolution 249,170]
257. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 153,151]
259. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 162,157]
261. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 257,158]
264. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 259,152]
266. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 261,164]
267. ~select_1(X1) | select_3(X1) [subsumption resolution 266,251]
269. ~select_3(X1) | select_1(X1) [resolution 264,154]
270. select_1(sK8) <- {2} [resolution 269,239]
272. main_query_1(sK8) <- {2} [resolution 270,156]
273. $false <- {1, 2} [subsumption resolution 272,230]
274. 1 | ~2 [avatar contradiction clause 273]
275. select_1(sK8) <- {0} [resolution 227,155]
276. select_3(sK8) <- {0} [resolution 275,267]
279. main_query_2(sK8) <- {0} [resolution 276,166]
281. $false <- {0, 3} [subsumption resolution 279,236]
282. ~0 | 3 [avatar contradiction clause 281]
283. $false [avatar sat refutation 237,238,274,282]
% SZS output end Proof for query02_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.010 s
% ------------------------------
% ------------------------------
