% Refutation found. Thanks to Tanya!
% SZS status Theorem for query08_query09_tptp4X
% SZS output start Proof for query08_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X6,X7,X8,X9,X10] : (from_8(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
6. ! [X6] : (select_7(X6) <=> ? [X7,X8,X9,X10] : ($true & from_8(X6,X7,X8,X9,X10))) [input]
7. ! [X6] : (subquery_from_6(X6) <=> select_7(X6)) [input]
8. ! [X11] : (from_5(X11) <=> subquery_from_6(X11)) [input]
9. ! [X11] : (select_4(X11) <=> ($true & from_5(X11))) [input]
10. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
14. ! [X12] : lte(X12,X12) [input]
21. ! [X12,X13] : (gte(X12,X13) <=> lte(X13,X12)) [input]
45. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
46. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 45]
47. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
48. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
49. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
50. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 49]
51. ! [X0] : (subquery_from_6(X0) <=> select_7(X0)) [rectify 7]
52. ! [X0] : (from_5(X0) <=> subquery_from_6(X0)) [rectify 8]
53. ! [X0] : (select_4(X0) <=> ($true & from_5(X0))) [rectify 9]
54. ! [X0] : (select_4(X0) <=> from_5(X0)) [true and false elimination 53]
55. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 10]
59. ! [X0] : lte(X0,X0) [rectify 14]
66. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 21]
90. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 46]
142. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 90]
143. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
144. ! [X0] : ((where_3(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
145. ! [X0] : ((where_3(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_3(X0))) [flattening 144]
146. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
147. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 146]
148. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
149. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 147,148]
150. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 47]
151. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 48]
152. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 50]
153. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 152]
154. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
155. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0))) [skolemisation 153,154]
156. ! [X0] : ((subquery_from_6(X0) | ~select_7(X0)) & (select_7(X0) | ~subquery_from_6(X0))) [nnf transformation 51]
157. ! [X0] : ((from_5(X0) | ~subquery_from_6(X0)) & (subquery_from_6(X0) | ~from_5(X0))) [nnf transformation 52]
158. ! [X0] : ((select_4(X0) | ~from_5(X0)) & (from_5(X0) | ~select_4(X0))) [nnf transformation 54]
159. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 55]
162. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 66]
163. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 142]
164. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
165. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 163,164]
166. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 143]
167. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 143]
170. ~gte(X0,X0) | where_3(X0) | ~lte(X0,X0) [cnf transformation 145]
171. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 149]
173. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 149]
174. ~main_query_1(X0) | select_1(X0) [cnf transformation 150]
175. ~select_1(X0) | main_query_1(X0) [cnf transformation 150]
176. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 151]
177. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 151]
178. from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [cnf transformation 155]
179. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 155]
180. ~subquery_from_6(X0) | select_7(X0) [cnf transformation 156]
181. ~select_7(X0) | subquery_from_6(X0) [cnf transformation 156]
182. ~from_5(X0) | subquery_from_6(X0) [cnf transformation 157]
183. ~subquery_from_6(X0) | from_5(X0) [cnf transformation 157]
184. ~select_4(X0) | from_5(X0) [cnf transformation 158]
185. ~from_5(X0) | select_4(X0) [cnf transformation 158]
186. ~main_query_2(X0) | select_4(X0) [cnf transformation 159]
187. ~select_4(X0) | main_query_2(X0) [cnf transformation 159]
191. lte(X0,X0) [cnf transformation 59]
198. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 162]
222. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 165]
223. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 165]
225. 0 <=> main_query_1(sK8) [avatar definition]
226. main_query_1(sK8) <- {0} [avatar component clause 225]
228. 1 <=> ~main_query_1(sK8) [avatar definition]
229. ~main_query_1(sK8) <- {1} [avatar component clause 228]
231. 2 <=> main_query_2(sK8) [avatar definition]
232. main_query_2(sK8) <- {2} [avatar component clause 231]
234. 3 <=> ~main_query_2(sK8) [avatar definition]
235. ~main_query_2(sK8) <- {3} [avatar component clause 234]
236. ~1 | ~3 [avatar split clause 223,234,228]
237. 0 | 2 [avatar split clause 222,231,225]
238. select_4(sK8) <- {2} [resolution 186,232]
239. from_5(sK8) <- {2} [resolution 238,184]
242. subquery_from_6(sK8) <- {2} [resolution 239,182]
244. select_7(sK8) <- {2} [resolution 242,180]
252. gte(X1,X1) [resolution 198,191]
257. where_3(X1) | ~lte(X1,X1) [resolution 252,170]
258. where_3(X1) [subsumption resolution 257,191]
324. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 171,166]
326. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [resolution 178,176]
328. from_8(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 324,177]
331. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_7(X1) [resolution 326,167]
333. ~select_1(X1) | select_7(X1) [resolution 328,179]
335. ~select_7(X1) | ~where_3(X1) | select_1(X1) [resolution 331,173]
336. ~select_7(X1) | select_1(X1) [subsumption resolution 335,258]
337. select_1(sK8) <- {2} [resolution 336,244]
341. main_query_1(sK8) <- {2} [resolution 337,175]
343. $false <- {1, 2} [subsumption resolution 341,229]
344. 1 | ~2 [avatar contradiction clause 343]
345. select_1(sK8) <- {0} [resolution 226,174]
348. select_7(sK8) <- {0} [resolution 345,333]
354. subquery_from_6(sK8) <- {0} [resolution 348,181]
357. from_5(sK8) <- {0} [resolution 354,183]
361. select_4(sK8) <- {0} [resolution 357,185]
365. main_query_2(sK8) <- {0} [resolution 361,187]
367. $false <- {0, 3} [subsumption resolution 365,235]
368. ~0 | 3 [avatar contradiction clause 367]
369. $false [avatar sat refutation 236,237,344,368]
% SZS output end Proof for query08_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.014 s
% ------------------------------
% ------------------------------
