% Refutation found. Thanks to Tanya!
% SZS status Theorem for query21_query80_tptp4X
% SZS output start Proof for query21_query80_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X5,X6] : (where_5(X5,X6) <=> eq(X5,X6)) [input]
3. ! [X5,X6] : (select_3(X5,X6) <=> ? [X0,X1,X2,X3,X4] : (where_5(X5,X6) & from_4(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (query_2(X5,X6) <=> select_3(X5,X6)) [input]
5. ! [X7,X8] : (from_6(X7,X8) <=> query_2(X7,X8)) [input]
6. ! [X7] : (select_1(X7) <=> ? [X8] : ($true & from_6(X7,X8))) [input]
7. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
8. ! [X9,X10,X11,X12,X13] : (from_8(X9,X10,X11,X12,X13) <=> tab1(X9,X10,X11,X12,X13)) [input]
9. ! [X0] : (select_7(X0) <=> ? [X9,X10,X11,X12,X13] : ($true & from_8(X9,X10,X11,X12,X13))) [input]
10. ! [X14] : (main_query_2(X14) <=> select_7(X14)) [input]
11. ! [X15] : eq(X15,X15) [input]
52. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
53. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 52]
54. ! [X0,X1] : (where_5(X0,X1) <=> eq(X0,X1)) [rectify 2]
55. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4,X5,X6] : (where_5(X0,X1) & from_4(X2,X3,X4,X5,X6))) [rectify 3]
56. ! [X0,X1] : (query_2(X0,X1) <=> select_3(X0,X1)) [rectify 4]
57. ! [X0,X1] : (from_6(X0,X1) <=> query_2(X0,X1)) [rectify 5]
58. ! [X0] : (select_1(X0) <=> ? [X1] : ($true & from_6(X0,X1))) [rectify 6]
59. ! [X0] : (select_1(X0) <=> ? [X1] : from_6(X0,X1)) [true and false elimination 58]
60. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 7]
61. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 8]
62. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4,X5] : ($true & from_8(X1,X2,X3,X4,X5))) [rectify 9]
63. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4,X5] : from_8(X1,X2,X3,X4,X5)) [true and false elimination 62]
64. ! [X0] : (main_query_2(X0) <=> select_7(X0)) [rectify 10]
65. ! [X0] : eq(X0,X0) [rectify 11]
106. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 53]
174. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 106]
175. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 1]
176. ! [X0,X1] : ((where_5(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_5(X0,X1))) [nnf transformation 54]
177. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_5(X0,X1) | ~from_4(X2,X3,X4,X5,X6))) & (? [X2,X3,X4,X5,X6] : (where_5(X0,X1) & from_4(X2,X3,X4,X5,X6)) | ~select_3(X0,X1))) [nnf transformation 55]
178. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_5(X0,X1) | ~from_4(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_5(X0,X1) & from_4(X7,X8,X9,X10,X11)) | ~select_3(X0,X1))) [rectify 177]
179. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_5(X0,X1) & from_4(X7,X8,X9,X10,X11)) => (where_5(X0,X1) & from_4(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)))) [choice axiom]
180. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_5(X0,X1) | ~from_4(X2,X3,X4,X5,X6))) & ((where_5(X0,X1) & from_4(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1))) | ~select_3(X0,X1))) [skolemisation 178,179]
181. ! [X0,X1] : ((query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~query_2(X0,X1))) [nnf transformation 56]
182. ! [X0,X1] : ((from_6(X0,X1) | ~query_2(X0,X1)) & (query_2(X0,X1) | ~from_6(X0,X1))) [nnf transformation 57]
183. ! [X0] : ((select_1(X0) | ! [X1] : ~from_6(X0,X1)) & (? [X1] : from_6(X0,X1) | ~select_1(X0))) [nnf transformation 59]
184. ! [X0] : ((select_1(X0) | ! [X1] : ~from_6(X0,X1)) & (? [X2] : from_6(X0,X2) | ~select_1(X0))) [rectify 183]
185. ! [X0] : (? [X2] : from_6(X0,X2) => from_6(X0,sK5(X0))) [choice axiom]
186. ! [X0] : ((select_1(X0) | ! [X1] : ~from_6(X0,X1)) & (from_6(X0,sK5(X0)) | ~select_1(X0))) [skolemisation 184,185]
187. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 60]
188. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 61]
189. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4,X5] : ~from_8(X1,X2,X3,X4,X5)) & (? [X1,X2,X3,X4,X5] : from_8(X1,X2,X3,X4,X5) | ~select_7(X0))) [nnf transformation 63]
190. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4,X5] : ~from_8(X1,X2,X3,X4,X5)) & (? [X6,X7,X8,X9,X10] : from_8(X6,X7,X8,X9,X10) | ~select_7(X0))) [rectify 189]
191. ? [X6,X7,X8,X9,X10] : from_8(X6,X7,X8,X9,X10) => from_8(sK6,sK7,sK8,sK9,sK10) [choice axiom]
192. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4,X5] : ~from_8(X1,X2,X3,X4,X5)) & (from_8(sK6,sK7,sK8,sK9,sK10) | ~select_7(X0))) [skolemisation 190,191]
193. ! [X0] : ((main_query_2(X0) | ~select_7(X0)) & (select_7(X0) | ~main_query_2(X0))) [nnf transformation 64]
196. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 174]
197. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK11) | ~main_query_1(sK11)) & (main_query_2(sK11) | main_query_1(sK11))) [choice axiom]
198. (~main_query_2(sK11) | ~main_query_1(sK11)) & (main_query_2(sK11) | main_query_1(sK11)) [skolemisation 196,197]
199. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 175]
200. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 175]
202. ~eq(X0,X1) | where_5(X0,X1) [cnf transformation 176]
203. from_4(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_3(X0,X1) [cnf transformation 180]
205. select_3(X0,X1) | ~where_5(X0,X1) | ~from_4(X2,X3,X4,X5,X6) [cnf transformation 180]
206. ~query_2(X0,X1) | select_3(X0,X1) [cnf transformation 181]
207. ~select_3(X0,X1) | query_2(X0,X1) [cnf transformation 181]
208. ~from_6(X0,X1) | query_2(X0,X1) [cnf transformation 182]
209. ~query_2(X0,X1) | from_6(X0,X1) [cnf transformation 182]
210. from_6(X0,sK5(X0)) | ~select_1(X0) [cnf transformation 186]
211. ~from_6(X0,X1) | select_1(X0) [cnf transformation 186]
212. ~main_query_1(X0) | select_1(X0) [cnf transformation 187]
213. ~select_1(X0) | main_query_1(X0) [cnf transformation 187]
214. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 188]
215. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 188]
216. from_8(sK6,sK7,sK8,sK9,sK10) | ~select_7(X0) [cnf transformation 192]
217. select_7(X0) | ~from_8(X1,X2,X3,X4,X5) [cnf transformation 192]
218. ~main_query_2(X0) | select_7(X0) [cnf transformation 193]
219. ~select_7(X0) | main_query_2(X0) [cnf transformation 193]
220. eq(X0,X0) [cnf transformation 65]
259. main_query_2(sK11) | main_query_1(sK11) [cnf transformation 198]
260. ~main_query_2(sK11) | ~main_query_1(sK11) [cnf transformation 198]
293. 0 <=> ! [X3,X5,X2,X4,X6] : ~from_4(X2,X3,X4,X5,X6) [avatar definition]
294. ~from_4(X2,X3,X4,X5,X6) <- {0} [avatar component clause 293]
296. 2 <=> ! [X1,X0] : (select_3(X0,X1) | ~where_5(X0,X1)) [avatar definition]
297. ~where_5(X0,X1) | select_3(X0,X1) <- {2} [avatar component clause 296]
298. 0 | 2 [avatar split clause 205,296,293]
300. 4 <=> ! [X1,X3,X5,X2,X4] : ~from_8(X1,X2,X3,X4,X5) [avatar definition]
301. ~from_8(X1,X2,X3,X4,X5) <- {4} [avatar component clause 300]
303. 6 <=> ! [X0] : select_7(X0) [avatar definition]
304. select_7(X0) <- {6} [avatar component clause 303]
305. 4 | 6 [avatar split clause 217,303,300]
307. 8 <=> ! [X0] : ~select_7(X0) [avatar definition]
308. ~select_7(X0) <- {8} [avatar component clause 307]
313. 10 <=> from_8(sK6,sK7,sK8,sK9,sK10) [avatar definition]
314. from_8(sK6,sK7,sK8,sK9,sK10) <- {10} [avatar component clause 313]
315. 8 | 10 [avatar split clause 216,313,307]
317. 12 <=> main_query_1(sK11) [avatar definition]
318. main_query_1(sK11) <- {12} [avatar component clause 317]
320. 13 <=> ~main_query_1(sK11) [avatar definition]
321. ~main_query_1(sK11) <- {13} [avatar component clause 320]
323. 14 <=> main_query_2(sK11) [avatar definition]
324. main_query_2(sK11) <- {14} [avatar component clause 323]
326. 15 <=> ~main_query_2(sK11) [avatar definition]
327. ~main_query_2(sK11) <- {15} [avatar component clause 326]
328. ~13 | ~15 [avatar split clause 260,326,320]
329. 12 | 14 [avatar split clause 259,323,317]
330. select_1(sK11) <- {12} [resolution 212,318]
332. where_5(X0,X0) [resolution 202,220]
334. query_2(X0,sK5(X0)) | ~select_1(X0) [resolution 210,208]
352. select_3(X1,sK5(X1)) | ~select_1(X1) [resolution 334,206]
360. tab1(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_3(X0,X1) [resolution 203,199]
363. from_8(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_3(X0,X1) [resolution 360,215]
365. ~select_3(X0,X1) <- {4} [subsumption resolution 363,301]
370. ~select_1(X1) <- {4} [subsumption resolution 352,365]
372. $false <- {4, 12} [resolution 370,330]
373. ~4 | ~12 [avatar contradiction clause 372]
374. select_7(sK11) <- {14} [resolution 324,218]
375. $false <- {8, 14} [subsumption resolution 374,308]
376. ~8 | ~14 [avatar contradiction clause 375]
383. tab1(sK6,sK7,sK8,sK9,sK10) <- {10} [resolution 314,214]
385. main_query_2(X0) <- {6} [resolution 304,219]
388. from_4(sK6,sK7,sK8,sK9,sK10) <- {10} [resolution 383,200]
389. $false <- {0, 10} [subsumption resolution 388,294]
390. ~0 | ~10 [avatar contradiction clause 389]
391. select_3(X0,X0) <- {2} [resolution 297,332]
392. query_2(X0,X0) <- {2} [resolution 391,207]
396. from_6(X0,X0) <- {2} [resolution 392,209]
400. select_1(X1) <- {2} [resolution 396,211]
402. main_query_1(X0) <- {2} [resolution 400,213]
407. $false <- {2, 13} [resolution 402,321]
408. ~2 | 13 [avatar contradiction clause 407]
409. $false <- {6, 15} [subsumption resolution 327,385]
410. ~6 | 15 [avatar contradiction clause 409]
411. $false [avatar sat refutation 298,305,315,328,329,373,376,390,408,410]
% SZS output end Proof for query21_query80_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
