% Refutation found. Thanks to Tanya!
% SZS status Theorem for query05_query06_tptp4X
% SZS output start Proof for query05_query06_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X6,X7,X8,X9,X10] : (from_8(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
6. ! [X6] : (select_7(X6) <=> ? [X7,X8,X9,X10] : ($true & from_8(X6,X7,X8,X9,X10))) [input]
7. ! [X5] : (subquery_from_6(X5) <=> select_7(X5)) [input]
8. ! [X11] : (from_5(X11) <=> subquery_from_6(X11)) [input]
9. ! [X11] : (select_4(X11) <=> ($true & from_5(X11))) [input]
10. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
11. ! [X12] : eq(X12,X12) [input]
45. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
46. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 45]
47. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
48. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
49. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
50. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 49]
51. ! [X0] : (subquery_from_6(X0) <=> select_7(X0)) [rectify 7]
52. ! [X0] : (from_5(X0) <=> subquery_from_6(X0)) [rectify 8]
53. ! [X0] : (select_4(X0) <=> ($true & from_5(X0))) [rectify 9]
54. ! [X0] : (select_4(X0) <=> from_5(X0)) [true and false elimination 53]
55. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 10]
56. ! [X0] : eq(X0,X0) [rectify 11]
90. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 46]
142. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 90]
143. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
144. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
145. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
146. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 145]
147. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
148. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 146,147]
149. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 47]
150. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 48]
151. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 50]
152. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 151]
153. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
154. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0))) [skolemisation 152,153]
155. ! [X0] : ((subquery_from_6(X0) | ~select_7(X0)) & (select_7(X0) | ~subquery_from_6(X0))) [nnf transformation 51]
156. ! [X0] : ((from_5(X0) | ~subquery_from_6(X0)) & (subquery_from_6(X0) | ~from_5(X0))) [nnf transformation 52]
157. ! [X0] : ((select_4(X0) | ~from_5(X0)) & (from_5(X0) | ~select_4(X0))) [nnf transformation 54]
158. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 55]
161. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 142]
162. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
163. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 161,162]
164. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 143]
165. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 143]
167. ~eq(X0,X0) | where_3(X0) [cnf transformation 144]
168. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 148]
170. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 148]
171. ~main_query_1(X0) | select_1(X0) [cnf transformation 149]
172. ~select_1(X0) | main_query_1(X0) [cnf transformation 149]
173. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 150]
174. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 150]
175. from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [cnf transformation 154]
176. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 154]
177. ~subquery_from_6(X0) | select_7(X0) [cnf transformation 155]
178. ~select_7(X0) | subquery_from_6(X0) [cnf transformation 155]
179. ~from_5(X0) | subquery_from_6(X0) [cnf transformation 156]
180. ~subquery_from_6(X0) | from_5(X0) [cnf transformation 156]
181. ~select_4(X0) | from_5(X0) [cnf transformation 157]
182. ~from_5(X0) | select_4(X0) [cnf transformation 157]
183. ~main_query_2(X0) | select_4(X0) [cnf transformation 158]
184. ~select_4(X0) | main_query_2(X0) [cnf transformation 158]
185. eq(X0,X0) [cnf transformation 56]
217. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 163]
218. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 163]
220. 0 <=> main_query_1(sK8) [avatar definition]
221. main_query_1(sK8) <- {0} [avatar component clause 220]
223. 1 <=> ~main_query_1(sK8) [avatar definition]
224. ~main_query_1(sK8) <- {1} [avatar component clause 223]
226. 2 <=> main_query_2(sK8) [avatar definition]
227. main_query_2(sK8) <- {2} [avatar component clause 226]
229. 3 <=> ~main_query_2(sK8) [avatar definition]
230. ~main_query_2(sK8) <- {3} [avatar component clause 229]
231. ~1 | ~3 [avatar split clause 218,229,223]
232. 0 | 2 [avatar split clause 217,226,220]
233. select_4(sK8) <- {2} [resolution 183,227]
234. from_5(sK8) <- {2} [resolution 233,181]
237. subquery_from_6(sK8) <- {2} [resolution 234,179]
239. select_7(sK8) <- {2} [resolution 237,177]
241. where_3(X1) [resolution 167,185]
276. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 168,164]
278. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [resolution 175,173]
280. from_8(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 276,174]
284. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_7(X1) [resolution 278,165]
289. ~select_1(X5) | select_7(X5) [resolution 280,176]
295. ~select_7(X1) | ~where_3(X1) | select_1(X1) [resolution 284,170]
296. ~select_7(X1) | select_1(X1) [subsumption resolution 295,241]
299. select_1(sK8) <- {2} [resolution 296,239]
303. main_query_1(sK8) <- {2} [resolution 299,172]
305. $false <- {1, 2} [subsumption resolution 303,224]
306. 1 | ~2 [avatar contradiction clause 305]
307. select_1(sK8) <- {0} [resolution 221,171]
310. select_7(sK8) <- {0} [resolution 307,289]
316. subquery_from_6(sK8) <- {0} [resolution 310,178]
319. from_5(sK8) <- {0} [resolution 316,180]
323. select_4(sK8) <- {0} [resolution 319,182]
327. main_query_2(sK8) <- {0} [resolution 323,184]
329. $false <- {0, 3} [subsumption resolution 327,230]
330. ~0 | 3 [avatar contradiction clause 329]
331. $false [avatar sat refutation 231,232,306,330]
% SZS output end Proof for query05_query06_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
