% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_08_query_09_tptp4X
% SZS output start Proof for query_08_query_09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X5,X1,X2,X3,X4] : (from_5(X5,X1,X2,X3,X4) <=> tab1(X5,X1,X2,X3,X4)) [input]
6. ! [X5] : (where_6(X5) <=> (gte(X5,X5) & lte(X5,X5))) [input]
7. ! [X5] : (select_4(X5) <=> ? [X1,X2,X3,X4] : (where_6(X5) & from_5(X5,X1,X2,X3,X4))) [input]
8. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
12. ! [X7] : lte(X7,X7) [input]
16. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
17. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
19. ! [X7,X8] : (gte(X7,X8) <=> lte(X8,X7)) [input]
40. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
41. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 40]
42. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
43. ! [X0] : (where_6(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 6]
44. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
45. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
46. ! [X0] : eq(X0,X0) [rectify 9]
49. ! [X0] : lte(X0,X0) [rectify 12]
53. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 16]
54. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 17]
56. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 19]
122. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 41]
123. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
124. ! [X0] : ((where_3(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
125. ! [X0] : ((where_3(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [flattening 124]
126. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
127. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 126]
128. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
129. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 127,128]
130. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
131. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 42]
132. ! [X0] : ((where_6(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [nnf transformation 43]
133. ! [X0] : ((where_6(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [flattening 132]
134. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 44]
135. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 134]
136. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
137. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 135,136]
138. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 45]
141. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 53]
142. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 141]
143. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 54]
144. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 56]
145. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 122]
146. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
147. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 145,146]
148. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 123]
149. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 123]
152. gt(X0,X0) | where_3(X0) | lt(X0,X0) [cnf transformation 125]
153. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 129]
155. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 129]
156. ~main_query_1(X0) | select_1(X0) [cnf transformation 130]
157. ~select_1(X0) | main_query_1(X0) [cnf transformation 130]
158. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 131]
159. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 131]
162. ~gte(X0,X0) | where_6(X0) | ~lte(X0,X0) [cnf transformation 133]
163. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 137]
165. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 137]
166. ~main_query_2(X0) | select_4(X0) [cnf transformation 138]
167. ~select_4(X0) | main_query_2(X0) [cnf transformation 138]
168. eq(X0,X0) [cnf transformation 46]
171. lte(X0,X0) [cnf transformation 49]
178. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 142]
180. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 143]
183. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 144]
204. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 147]
205. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 147]
207. 0 <=> main_query_1(sK8) [avatar definition]
208. main_query_1(sK8) <- {0} [avatar component clause 207]
210. 1 <=> ~main_query_1(sK8) [avatar definition]
211. ~main_query_1(sK8) <- {1} [avatar component clause 210]
213. 2 <=> main_query_2(sK8) [avatar definition]
214. main_query_2(sK8) <- {2} [avatar component clause 213]
216. 3 <=> ~main_query_2(sK8) [avatar definition]
217. ~main_query_2(sK8) <- {3} [avatar component clause 216]
218. ~1 | ~3 [avatar split clause 205,216,210]
219. 0 | 2 [avatar split clause 204,213,207]
220. select_4(sK8) <- {2} [resolution 166,214]
227. ~lt(X0,X0) [resolution 178,168]
230. gte(X1,X1) [resolution 183,171]
234. where_3(X1) | lt(X1,X1) | lt(X1,X1) [resolution 152,180]
235. where_3(X1) | lt(X1,X1) [duplicate literal removal 234]
236. where_3(X1) [subsumption resolution 235,227]
239. where_6(X1) | ~lte(X1,X1) [resolution 230,162]
240. where_6(X1) [subsumption resolution 239,171]
273. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 153,148]
321. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 163,158]
542. from_5(X1,sK0(X1),sK1(X1),sK2(X1),sK3(X1)) | ~select_1(X1) [resolution 273,159]
572. from_2(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 321,149]
727. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 542,165]
733. ~select_1(X1) | select_4(X1) [subsumption resolution 727,240]
735. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 572,155]
741. ~select_4(X1) | select_1(X1) [subsumption resolution 735,236]
742. select_1(sK8) <- {2} [resolution 741,220]
751. main_query_1(sK8) <- {2} [resolution 742,157]
753. $false <- {1, 2} [subsumption resolution 751,211]
754. 1 | ~2 [avatar contradiction clause 753]
755. select_1(sK8) <- {0} [resolution 208,156]
756. select_4(sK8) <- {0} [resolution 755,733]
760. main_query_2(sK8) <- {0} [resolution 756,167]
762. $false <- {0, 3} [subsumption resolution 760,217]
763. ~0 | 3 [avatar contradiction clause 762]
764. $false [avatar sat refutation 218,219,754,763]
% SZS output end Proof for query_08_query_09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.035 s
% ------------------------------
% ------------------------------
