% Refutation found. Thanks to Tanya!
% SZS status Theorem for query02_query08_tptp4X
% SZS output start Proof for query02_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_4(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (where_5(X6) <=> (gte(X6,X6) & lte(X6,X6))) [input]
6. ! [X6] : (select_3(X6) <=> ? [X7,X8,X9,X10] : (where_5(X6) & from_4(X6,X7,X8,X9,X10))) [input]
7. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
11. ! [X11] : lte(X11,X11) [input]
18. ! [X11,X12] : (gte(X11,X12) <=> lte(X12,X11)) [input]
39. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
40. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 39]
41. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
42. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
43. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
44. ! [X0] : (where_5(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 5]
45. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
46. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
50. ! [X0] : lte(X0,X0) [rectify 11]
57. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 18]
78. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 40]
124. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 78]
125. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
126. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 41]
127. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 126]
128. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
129. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 127,128]
130. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 42]
131. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 43]
132. ! [X0] : ((where_5(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [nnf transformation 44]
133. ! [X0] : ((where_5(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [flattening 132]
134. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 45]
135. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 134]
136. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
137. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 135,136]
138. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 46]
141. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 57]
142. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 124]
143. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
144. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 142,143]
145. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 125]
146. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 125]
147. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 129]
148. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 129]
149. ~main_query_1(X0) | select_1(X0) [cnf transformation 130]
150. ~select_1(X0) | main_query_1(X0) [cnf transformation 130]
151. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 131]
152. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 131]
155. ~gte(X0,X0) | where_5(X0) | ~lte(X0,X0) [cnf transformation 133]
156. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 137]
158. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 137]
159. ~main_query_2(X0) | select_3(X0) [cnf transformation 138]
160. ~select_3(X0) | main_query_2(X0) [cnf transformation 138]
164. lte(X0,X0) [cnf transformation 50]
171. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 141]
192. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 144]
193. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 144]
195. 0 <=> main_query_1(sK8) [avatar definition]
196. main_query_1(sK8) <- {0} [avatar component clause 195]
198. 1 <=> ~main_query_1(sK8) [avatar definition]
199. ~main_query_1(sK8) <- {1} [avatar component clause 198]
201. 2 <=> main_query_2(sK8) [avatar definition]
202. main_query_2(sK8) <- {2} [avatar component clause 201]
204. 3 <=> ~main_query_2(sK8) [avatar definition]
205. ~main_query_2(sK8) <- {3} [avatar component clause 204]
206. ~1 | ~3 [avatar split clause 193,204,198]
207. 0 | 2 [avatar split clause 192,201,195]
208. select_3(sK8) <- {2} [resolution 159,202]
217. gte(X1,X1) [resolution 171,164]
221. where_5(X0) | ~lte(X0,X0) [resolution 155,217]
223. where_5(X0) [subsumption resolution 221,164]
268. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 147,145]
291. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 156,151]
294. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 268,152]
298. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 291,146]
301. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 294,158]
302. ~select_1(X1) | select_3(X1) [subsumption resolution 301,223]
306. ~select_3(X1) | select_1(X1) [resolution 298,148]
314. select_1(sK8) <- {2} [resolution 306,208]
318. main_query_1(sK8) <- {2} [resolution 314,150]
319. $false <- {1, 2} [subsumption resolution 318,199]
320. 1 | ~2 [avatar contradiction clause 319]
321. select_1(sK8) <- {0} [resolution 196,149]
324. select_3(sK8) <- {0} [resolution 321,302]
329. main_query_2(sK8) <- {0} [resolution 324,160]
331. $false <- {0, 3} [subsumption resolution 329,205]
332. ~0 | 3 [avatar contradiction clause 331]
333. $false [avatar sat refutation 206,207,320,332]
% SZS output end Proof for query02_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
