% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_04_query_10_tptp4X
% SZS output start Proof for query_04_query_10_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [input]
6. ! [X0] : (subquery_from_5(X0) <=> select_6(X0)) [input]
7. ! [X5] : (from_4(X5) <=> subquery_from_5(X5)) [input]
8. ! [X5] : (select_3(X5) <=> ($true & from_4(X5))) [input]
9. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
41. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
42. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 5]
45. ! [X0] : (from_4(X0) <=> subquery_from_5(X0)) [rectify 7]
46. ! [X0] : (select_3(X0) <=> ($true & from_4(X0))) [rectify 8]
47. ! [X0] : (select_3(X0) <=> from_4(X0)) [true and false elimination 46]
124. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 42]
125. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
126. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
127. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 126]
128. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
129. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 127,128]
130. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
131. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 4]
132. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 44]
133. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 132]
134. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
135. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 133,134]
136. ! [X0] : ((subquery_from_5(X0) | ~select_6(X0)) & (select_6(X0) | ~subquery_from_5(X0))) [nnf transformation 6]
137. ! [X0] : ((from_4(X0) | ~subquery_from_5(X0)) & (subquery_from_5(X0) | ~from_4(X0))) [nnf transformation 45]
138. ! [X0] : ((select_3(X0) | ~from_4(X0)) & (from_4(X0) | ~select_3(X0))) [nnf transformation 47]
139. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 9]
142. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 124]
143. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
144. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 142,143]
145. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 125]
146. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 125]
147. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 129]
148. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 129]
149. ~main_query_1(X0) | select_1(X0) [cnf transformation 130]
150. ~select_1(X0) | main_query_1(X0) [cnf transformation 130]
151. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 131]
152. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 131]
153. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 135]
154. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 135]
155. ~subquery_from_5(X0) | select_6(X0) [cnf transformation 136]
156. ~select_6(X0) | subquery_from_5(X0) [cnf transformation 136]
157. ~from_4(X0) | subquery_from_5(X0) [cnf transformation 137]
158. ~subquery_from_5(X0) | from_4(X0) [cnf transformation 137]
159. ~select_3(X0) | from_4(X0) [cnf transformation 138]
160. ~from_4(X0) | select_3(X0) [cnf transformation 138]
161. ~main_query_2(X0) | select_3(X0) [cnf transformation 139]
162. ~select_3(X0) | main_query_2(X0) [cnf transformation 139]
192. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 144]
193. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 144]
195. 0 <=> main_query_1(sK8) [avatar definition]
196. main_query_1(sK8) <- {0} [avatar component clause 195]
198. 1 <=> ~main_query_1(sK8) [avatar definition]
199. ~main_query_1(sK8) <- {1} [avatar component clause 198]
201. 2 <=> main_query_2(sK8) [avatar definition]
202. main_query_2(sK8) <- {2} [avatar component clause 201]
204. 3 <=> ~main_query_2(sK8) [avatar definition]
205. ~main_query_2(sK8) <- {3} [avatar component clause 204]
206. ~1 | ~3 [avatar split clause 193,204,198]
207. 0 | 2 [avatar split clause 192,201,195]
208. select_3(sK8) <- {2} [resolution 161,202]
210. from_4(sK8) <- {2} [resolution 208,159]
213. subquery_from_5(sK8) <- {2} [resolution 210,157]
216. select_6(sK8) <- {2} [resolution 213,155]
245. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 147,145]
247. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 153,151]
250. from_7(X2,sK0(X2),sK1(X2),sK2(X2),sK3(X2)) | ~select_1(X2) [resolution 245,152]
256. from_2(X5,sK4(X5),sK5(X5),sK6(X5),sK7(X5)) | ~select_6(X5) [resolution 247,146]
260. ~select_1(X1) | select_6(X1) [resolution 250,154]
266. ~select_6(X1) | select_1(X1) [resolution 256,148]
267. select_1(sK8) <- {2} [resolution 266,216]
269. main_query_1(sK8) <- {2} [resolution 267,150]
270. $false <- {1, 2} [subsumption resolution 269,199]
271. 1 | ~2 [avatar contradiction clause 270]
272. select_1(sK8) <- {0} [resolution 196,149]
273. select_6(sK8) <- {0} [resolution 272,260]
276. subquery_from_5(sK8) <- {0} [resolution 273,156]
277. from_4(sK8) <- {0} [resolution 276,158]
281. select_3(sK8) <- {0} [resolution 277,160]
285. main_query_2(sK8) <- {0} [resolution 281,162]
287. $false <- {0, 3} [subsumption resolution 285,205]
288. ~0 | 3 [avatar contradiction clause 287]
289. $false [avatar sat refutation 206,207,271,288]
% SZS output end Proof for query_04_query_10_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.025 s
% ------------------------------
% ------------------------------
