% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_02_query_07_tptp4X
% SZS output start Proof for query_02_query_07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X5,X6,X7,X8,X9] : (from_7(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
5. ! [X5] : (select_6(X5) <=> ? [X6,X7,X8,X9] : ($true & from_7(X5,X6,X7,X8,X9))) [input]
6. ! [X10] : (subquery_from_5(X10) <=> select_6(X10)) [input]
7. ! [X11] : (from_4(X11) <=> subquery_from_5(X11)) [input]
8. ! [X11] : (select_3(X11) <=> ($true & from_4(X11))) [input]
9. ! [X10] : (main_query_2(X10) <=> select_3(X10)) [input]
41. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
42. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
45. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [rectify 5]
46. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 45]
47. ! [X0] : (subquery_from_5(X0) <=> select_6(X0)) [rectify 6]
48. ! [X0] : (from_4(X0) <=> subquery_from_5(X0)) [rectify 7]
49. ! [X0] : (select_3(X0) <=> ($true & from_4(X0))) [rectify 8]
50. ! [X0] : (select_3(X0) <=> from_4(X0)) [true and false elimination 49]
51. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 9]
128. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 42]
129. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
130. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
131. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 130]
132. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
133. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 131,132]
134. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
135. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 44]
136. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 46]
137. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 136]
138. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
139. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 137,138]
140. ! [X0] : ((subquery_from_5(X0) | ~select_6(X0)) & (select_6(X0) | ~subquery_from_5(X0))) [nnf transformation 47]
141. ! [X0] : ((from_4(X0) | ~subquery_from_5(X0)) & (subquery_from_5(X0) | ~from_4(X0))) [nnf transformation 48]
142. ! [X0] : ((select_3(X0) | ~from_4(X0)) & (from_4(X0) | ~select_3(X0))) [nnf transformation 50]
143. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 51]
146. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 128]
147. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
148. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 146,147]
149. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 129]
150. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 129]
151. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 133]
152. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 133]
153. ~main_query_1(X0) | select_1(X0) [cnf transformation 134]
154. ~select_1(X0) | main_query_1(X0) [cnf transformation 134]
155. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 135]
156. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 135]
157. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 139]
158. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 139]
159. ~subquery_from_5(X0) | select_6(X0) [cnf transformation 140]
160. ~select_6(X0) | subquery_from_5(X0) [cnf transformation 140]
161. ~from_4(X0) | subquery_from_5(X0) [cnf transformation 141]
162. ~subquery_from_5(X0) | from_4(X0) [cnf transformation 141]
163. ~select_3(X0) | from_4(X0) [cnf transformation 142]
164. ~from_4(X0) | select_3(X0) [cnf transformation 142]
165. ~main_query_2(X0) | select_3(X0) [cnf transformation 143]
166. ~select_3(X0) | main_query_2(X0) [cnf transformation 143]
196. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 148]
197. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 148]
199. 0 <=> main_query_1(sK8) [avatar definition]
200. main_query_1(sK8) <- {0} [avatar component clause 199]
202. 1 <=> ~main_query_1(sK8) [avatar definition]
203. ~main_query_1(sK8) <- {1} [avatar component clause 202]
205. 2 <=> main_query_2(sK8) [avatar definition]
206. main_query_2(sK8) <- {2} [avatar component clause 205]
208. 3 <=> ~main_query_2(sK8) [avatar definition]
209. ~main_query_2(sK8) <- {3} [avatar component clause 208]
210. ~1 | ~3 [avatar split clause 197,208,202]
211. 0 | 2 [avatar split clause 196,205,199]
212. select_3(sK8) <- {2} [resolution 165,206]
214. from_4(sK8) <- {2} [resolution 212,163]
217. subquery_from_5(sK8) <- {2} [resolution 214,161]
220. select_6(sK8) <- {2} [resolution 217,159]
249. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 151,149]
251. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 157,155]
254. from_7(X2,sK0(X2),sK1(X2),sK2(X2),sK3(X2)) | ~select_1(X2) [resolution 249,156]
260. from_2(X5,sK4(X5),sK5(X5),sK6(X5),sK7(X5)) | ~select_6(X5) [resolution 251,150]
264. ~select_1(X1) | select_6(X1) [resolution 254,158]
270. ~select_6(X1) | select_1(X1) [resolution 260,152]
273. select_1(sK8) <- {2} [resolution 270,220]
277. main_query_1(sK8) <- {2} [resolution 273,154]
278. $false <- {1, 2} [subsumption resolution 277,203]
279. 1 | ~2 [avatar contradiction clause 278]
280. select_1(sK8) <- {0} [resolution 200,153]
283. select_6(sK8) <- {0} [resolution 280,264]
288. subquery_from_5(sK8) <- {0} [resolution 283,160]
291. from_4(sK8) <- {0} [resolution 288,162]
293. select_3(sK8) <- {0} [resolution 291,164]
295. main_query_2(sK8) <- {0} [resolution 293,166]
297. $false <- {0, 3} [subsumption resolution 295,209]
298. ~0 | 3 [avatar contradiction clause 297]
299. $false [avatar sat refutation 210,211,279,298]
% SZS output end Proof for query_02_query_07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.018 s
% ------------------------------
% ------------------------------
