% Refutation found. Thanks to Tanya!
% SZS status Theorem for query03_query04_tptp4X
% SZS output start Proof for query03_query04_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X5,X6,X7,X8] : (from_4(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
5. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [input]
6. ! [X0] : (select_3(X0) <=> ? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8))) [input]
7. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
8. ! [X9] : eq(X9,X9) [input]
39. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
40. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 39]
41. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
42. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
43. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
44. ! [X0] : eq(X0,X0) [rectify 8]
120. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 40]
121. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
122. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 41]
123. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 122]
124. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
125. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 123,124]
126. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
127. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 42]
128. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 5]
129. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 43]
130. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 129]
131. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
132. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 130,131]
133. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 7]
136. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 120]
137. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
138. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 136,137]
139. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 121]
140. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 121]
141. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 125]
142. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 125]
143. ~main_query_1(X0) | select_1(X0) [cnf transformation 126]
144. ~select_1(X0) | main_query_1(X0) [cnf transformation 126]
145. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 127]
146. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 127]
148. ~eq(X0,X0) | where_5(X0) [cnf transformation 128]
149. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 132]
151. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 132]
152. ~main_query_2(X0) | select_3(X0) [cnf transformation 133]
153. ~select_3(X0) | main_query_2(X0) [cnf transformation 133]
154. eq(X0,X0) [cnf transformation 44]
183. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 138]
184. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 138]
186. 0 <=> main_query_1(sK8) [avatar definition]
187. main_query_1(sK8) <- {0} [avatar component clause 186]
189. 1 <=> ~main_query_1(sK8) [avatar definition]
190. ~main_query_1(sK8) <- {1} [avatar component clause 189]
192. 2 <=> main_query_2(sK8) [avatar definition]
193. main_query_2(sK8) <- {2} [avatar component clause 192]
195. 3 <=> ~main_query_2(sK8) [avatar definition]
196. ~main_query_2(sK8) <- {3} [avatar component clause 195]
197. ~1 | ~3 [avatar split clause 184,195,189]
198. 0 | 2 [avatar split clause 183,192,186]
199. select_3(sK8) <- {2} [resolution 152,193]
203. where_5(X1) [resolution 148,154]
237. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 141,139]
239. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 149,145]
242. from_4(X2,sK0(X2),sK1(X2),sK2(X2),sK3(X2)) | ~select_1(X2) [resolution 237,146]
248. from_2(X5,sK4(X5),sK5(X5),sK6(X5),sK7(X5)) | ~select_3(X5) [resolution 239,140]
252. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 242,151]
253. ~select_1(X1) | select_3(X1) [subsumption resolution 252,203]
259. ~select_3(X1) | select_1(X1) [resolution 248,142]
260. select_1(sK8) <- {2} [resolution 259,199]
262. main_query_1(sK8) <- {2} [resolution 260,144]
263. $false <- {1, 2} [subsumption resolution 262,190]
264. 1 | ~2 [avatar contradiction clause 263]
265. select_1(sK8) <- {0} [resolution 187,143]
266. select_3(sK8) <- {0} [resolution 265,253]
269. main_query_2(sK8) <- {0} [resolution 266,153]
271. $false <- {0, 3} [subsumption resolution 269,196]
272. ~0 | 3 [avatar contradiction clause 271]
273. $false [avatar sat refutation 197,198,264,272]
% SZS output end Proof for query03_query04_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.032 s
% ------------------------------
% ------------------------------
