% Refutation found. Thanks to Tanya!
% SZS status Theorem for query04_query06_tptp4X
% SZS output start Proof for query04_query06_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X5,X6,X7,X8,X9] : (from_8(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
6. ! [X5] : (select_7(X5) <=> ? [X6,X7,X8,X9] : ($true & from_8(X5,X6,X7,X8,X9))) [input]
7. ! [X10] : (subquery_from_6(X10) <=> select_7(X10)) [input]
8. ! [X11] : (from_5(X11) <=> subquery_from_6(X11)) [input]
9. ! [X11] : (select_4(X11) <=> ($true & from_5(X11))) [input]
10. ! [X10] : (main_query_2(X10) <=> select_4(X10)) [input]
11. ! [X12] : eq(X12,X12) [input]
42. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
43. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 42]
44. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
45. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
46. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 45]
47. ! [X0] : (subquery_from_6(X0) <=> select_7(X0)) [rectify 7]
48. ! [X0] : (from_5(X0) <=> subquery_from_6(X0)) [rectify 8]
49. ! [X0] : (select_4(X0) <=> ($true & from_5(X0))) [rectify 9]
50. ! [X0] : (select_4(X0) <=> from_5(X0)) [true and false elimination 49]
51. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 10]
52. ! [X0] : eq(X0,X0) [rectify 11]
128. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 43]
129. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
130. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
131. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
132. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 131]
133. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 132,133]
135. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
136. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 44]
137. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 46]
138. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 137]
139. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
140. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0))) [skolemisation 138,139]
141. ! [X0] : ((subquery_from_6(X0) | ~select_7(X0)) & (select_7(X0) | ~subquery_from_6(X0))) [nnf transformation 47]
142. ! [X0] : ((from_5(X0) | ~subquery_from_6(X0)) & (subquery_from_6(X0) | ~from_5(X0))) [nnf transformation 48]
143. ! [X0] : ((select_4(X0) | ~from_5(X0)) & (from_5(X0) | ~select_4(X0))) [nnf transformation 50]
144. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 51]
147. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 128]
148. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
149. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 147,148]
150. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 129]
151. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 129]
153. ~eq(X0,X0) | where_3(X0) [cnf transformation 130]
154. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 134]
156. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 134]
157. ~main_query_1(X0) | select_1(X0) [cnf transformation 135]
158. ~select_1(X0) | main_query_1(X0) [cnf transformation 135]
159. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 136]
160. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 136]
161. from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [cnf transformation 140]
162. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 140]
163. ~subquery_from_6(X0) | select_7(X0) [cnf transformation 141]
164. ~select_7(X0) | subquery_from_6(X0) [cnf transformation 141]
165. ~from_5(X0) | subquery_from_6(X0) [cnf transformation 142]
166. ~subquery_from_6(X0) | from_5(X0) [cnf transformation 142]
167. ~select_4(X0) | from_5(X0) [cnf transformation 143]
168. ~from_5(X0) | select_4(X0) [cnf transformation 143]
169. ~main_query_2(X0) | select_4(X0) [cnf transformation 144]
170. ~select_4(X0) | main_query_2(X0) [cnf transformation 144]
171. eq(X0,X0) [cnf transformation 52]
200. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 149]
201. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 149]
203. 0 <=> main_query_1(sK8) [avatar definition]
204. main_query_1(sK8) <- {0} [avatar component clause 203]
206. 1 <=> ~main_query_1(sK8) [avatar definition]
207. ~main_query_1(sK8) <- {1} [avatar component clause 206]
209. 2 <=> main_query_2(sK8) [avatar definition]
210. main_query_2(sK8) <- {2} [avatar component clause 209]
212. 3 <=> ~main_query_2(sK8) [avatar definition]
213. ~main_query_2(sK8) <- {3} [avatar component clause 212]
214. ~1 | ~3 [avatar split clause 201,212,206]
215. 0 | 2 [avatar split clause 200,209,203]
216. select_4(sK8) <- {2} [resolution 169,210]
217. from_5(sK8) <- {2} [resolution 216,167]
220. subquery_from_6(sK8) <- {2} [resolution 217,165]
222. select_7(sK8) <- {2} [resolution 220,163]
224. where_3(X1) [resolution 153,171]
259. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 154,150]
261. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [resolution 161,159]
264. from_8(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 259,160]
268. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_7(X1) [resolution 261,151]
271. ~select_1(X1) | select_7(X1) [resolution 264,162]
275. ~select_7(X1) | ~where_3(X1) | select_1(X1) [resolution 268,156]
281. ~select_7(X1) | select_1(X1) [subsumption resolution 275,224]
284. select_1(sK8) <- {2} [resolution 281,222]
288. main_query_1(sK8) <- {2} [resolution 284,158]
290. $false <- {1, 2} [subsumption resolution 288,207]
291. 1 | ~2 [avatar contradiction clause 290]
292. select_1(sK8) <- {0} [resolution 204,157]
295. select_7(sK8) <- {0} [resolution 292,271]
301. subquery_from_6(sK8) <- {0} [resolution 295,164]
304. from_5(sK8) <- {0} [resolution 301,166]
308. select_4(sK8) <- {0} [resolution 304,168]
312. main_query_2(sK8) <- {0} [resolution 308,170]
314. $false <- {0, 3} [subsumption resolution 312,213]
315. ~0 | 3 [avatar contradiction clause 314]
316. $false [avatar sat refutation 214,215,291,315]
% SZS output end Proof for query04_query06_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
