% Refutation found. Thanks to Tanya!
% SZS status Theorem for query10_query12_tptp4X
% SZS output start Proof for query10_query12_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X1,X0] : (where_3(X1,X0) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X1,X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X1,X0] : (where_6(X1,X0) <=> (lte(X1,X0) & lte(X0,X1))) [input]
7. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X1,X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
13. ! [X7,X8] : ((lte(X8,X7) & lte(X7,X8)) <=> eq(X7,X8)) [input]
54. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
55. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 54]
56. ! [X0,X1] : (where_3(X0,X1) <=> eq(X1,X0)) [rectify 2]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
58. ! [X0,X1] : (where_6(X0,X1) <=> (lte(X0,X1) & lte(X1,X0))) [rectify 6]
59. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
64. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
105. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 55]
181. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 105]
182. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
183. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X1,X0)) & (eq(X1,X0) | ~where_3(X0,X1))) [nnf transformation 56]
184. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X1,X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X1,X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
185. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X1,X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X1,X0) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 184]
186. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X1,X0) & from_2(X0,X1,X5,X6,X7)) => (where_3(X1,X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
187. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X1,X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X1,X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 185,186]
188. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
189. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
190. ! [X0,X1] : ((where_6(X0,X1) | (~lte(X0,X1) | ~lte(X1,X0))) & ((lte(X0,X1) & lte(X1,X0)) | ~where_6(X0,X1))) [nnf transformation 58]
191. ! [X0,X1] : ((where_6(X0,X1) | ~lte(X0,X1) | ~lte(X1,X0)) & ((lte(X0,X1) & lte(X1,X0)) | ~where_6(X0,X1))) [flattening 190]
192. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X1,X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X1,X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 7]
193. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X1,X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X1,X0) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 192]
194. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X1,X0) & from_5(X0,X1,X5,X6,X7)) => (where_6(X1,X0) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
195. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X1,X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X1,X0) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 193,194]
196. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 59]
197. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 64]
198. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 197]
199. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 181]
200. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
201. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 199,200]
202. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 182]
203. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 182]
204. ~where_3(X0,X1) | eq(X1,X0) [cnf transformation 183]
205. ~eq(X1,X0) | where_3(X0,X1) [cnf transformation 183]
206. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 187]
207. ~select_1(X0,X1) | where_3(X1,X0) [cnf transformation 187]
208. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X1,X0) | select_1(X0,X1) [cnf transformation 187]
209. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 188]
210. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 188]
211. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 189]
212. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 189]
213. ~where_6(X0,X1) | lte(X1,X0) [cnf transformation 191]
214. ~where_6(X0,X1) | lte(X0,X1) [cnf transformation 191]
215. ~lte(X1,X0) | ~lte(X0,X1) | where_6(X0,X1) [cnf transformation 191]
216. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 195]
217. ~select_4(X0,X1) | where_6(X1,X0) [cnf transformation 195]
218. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X1,X0) | select_4(X0,X1) [cnf transformation 195]
219. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 196]
220. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 196]
225. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 198]
226. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 198]
227. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 198]
264. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 201]
265. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 201]
302. 0 <=> main_query_1(sK6,sK7) [avatar definition]
303. main_query_1(sK6,sK7) <- {0} [avatar component clause 302]
305. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
306. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 305]
308. 2 <=> main_query_2(sK6,sK7) [avatar definition]
309. main_query_2(sK6,sK7) <- {2} [avatar component clause 308]
311. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
312. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 311]
313. ~1 | ~3 [avatar split clause 265,311,305]
314. 0 | 2 [avatar split clause 264,308,302]
317. select_4(sK6,sK7) <- {2} [resolution 219,309]
318. where_6(sK7,sK6) <- {2} [resolution 317,217]
320. lte(sK7,sK6) <- {2} [resolution 318,214]
321. lte(sK6,sK7) <- {2} [resolution 318,213]
334. eq(sK6,sK7) | ~lte(sK6,sK7) <- {2} [resolution 225,320]
336. eq(sK6,sK7) <- {2} [subsumption resolution 334,321]
347. where_3(sK7,sK6) <- {2} [resolution 336,205]
362. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 206,202]
366. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 216,211]
367. from_5(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 362,212]
370. from_2(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 366,203]
371. ~where_6(X1,X0) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 367,218]
373. ~select_4(X0,X1) | ~where_3(X1,X0) | select_1(X0,X1) [resolution 370,208]
391. ~where_3(sK7,sK6) | select_1(sK6,sK7) <- {2} [resolution 373,317]
392. select_1(sK6,sK7) <- {2} [subsumption resolution 391,347]
393. main_query_1(sK6,sK7) <- {2} [resolution 392,210]
395. $false <- {1, 2} [subsumption resolution 393,306]
396. 1 | ~2 [avatar contradiction clause 395]
397. select_1(sK6,sK7) <- {0} [resolution 303,209]
399. where_3(sK7,sK6) <- {0} [resolution 397,207]
400. eq(sK6,sK7) <- {0} [resolution 399,204]
402. lte(sK7,sK6) <- {0} [resolution 400,227]
403. lte(sK6,sK7) <- {0} [resolution 400,226]
412. ~lte(sK7,sK6) | where_6(sK7,sK6) <- {0} [resolution 403,215]
413. where_6(sK7,sK6) <- {0} [subsumption resolution 412,402]
424. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 413,371]
427. select_4(sK6,sK7) <- {0} [subsumption resolution 424,397]
432. main_query_2(sK6,sK7) <- {0} [resolution 427,220]
434. $false <- {0, 3} [subsumption resolution 432,312]
435. ~0 | 3 [avatar contradiction clause 434]
436. $false [avatar sat refutation 313,314,396,435]
% SZS output end Proof for query10_query12_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.011 s
% ------------------------------
% ------------------------------
