% Refutation found. Thanks to Tanya!
% SZS status Theorem for query05_query08_tptp4X
% SZS output start Proof for query05_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
7. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
9. ! [X6] : eq(X6,X6) [input]
12. ! [X6] : lte(X6,X6) [input]
19. ! [X6,X7] : (gte(X6,X7) <=> lte(X7,X6)) [input]
40. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
41. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 40]
42. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
43. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
44. ! [X0] : eq(X0,X0) [rectify 9]
47. ! [X0] : lte(X0,X0) [rectify 12]
54. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 19]
75. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 41]
121. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 75]
122. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
123. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
124. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
125. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 124]
126. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
127. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 125,126]
128. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 42]
129. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
130. ! [X0] : ((where_6(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [nnf transformation 6]
131. ! [X0] : ((where_6(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [flattening 130]
132. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 7]
133. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 132]
134. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
135. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 133,134]
136. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 43]
139. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 54]
140. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 121]
141. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
142. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 140,141]
143. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 122]
144. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 122]
146. ~eq(X0,X0) | where_3(X0) [cnf transformation 123]
147. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 127]
149. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 127]
150. ~main_query_1(X0) | select_1(X0) [cnf transformation 128]
151. ~select_1(X0) | main_query_1(X0) [cnf transformation 128]
152. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 129]
153. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 129]
156. ~gte(X0,X0) | where_6(X0) | ~lte(X0,X0) [cnf transformation 131]
157. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 135]
159. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 135]
160. ~main_query_2(X0) | select_4(X0) [cnf transformation 136]
161. ~select_4(X0) | main_query_2(X0) [cnf transformation 136]
162. eq(X0,X0) [cnf transformation 44]
165. lte(X0,X0) [cnf transformation 47]
172. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 139]
193. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 142]
194. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 142]
196. 0 <=> main_query_1(sK8) [avatar definition]
197. main_query_1(sK8) <- {0} [avatar component clause 196]
199. 1 <=> ~main_query_1(sK8) [avatar definition]
200. ~main_query_1(sK8) <- {1} [avatar component clause 199]
202. 2 <=> main_query_2(sK8) [avatar definition]
203. main_query_2(sK8) <- {2} [avatar component clause 202]
205. 3 <=> ~main_query_2(sK8) [avatar definition]
206. ~main_query_2(sK8) <- {3} [avatar component clause 205]
207. ~1 | ~3 [avatar split clause 194,205,199]
208. 0 | 2 [avatar split clause 193,202,196]
209. select_4(sK8) <- {2} [resolution 160,203]
213. where_3(X1) [resolution 146,162]
223. gte(X1,X1) [resolution 172,165]
228. where_6(X1) | ~lte(X1,X1) [resolution 223,156]
229. where_6(X1) [subsumption resolution 228,165]
296. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 147,143]
298. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 157,152]
301. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 296,153]
305. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 298,144]
311. ~select_1(X7) | ~where_6(X7) | select_4(X7) [resolution 301,159]
312. ~select_1(X7) | select_4(X7) [subsumption resolution 311,229]
318. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 305,149]
319. ~select_4(X1) | select_1(X1) [subsumption resolution 318,213]
320. select_1(sK8) <- {2} [resolution 319,209]
322. main_query_1(sK8) <- {2} [resolution 320,151]
324. $false <- {1, 2} [subsumption resolution 322,200]
325. 1 | ~2 [avatar contradiction clause 324]
326. select_1(sK8) <- {0} [resolution 197,150]
327. select_4(sK8) <- {0} [resolution 326,312]
331. main_query_2(sK8) <- {0} [resolution 327,161]
333. $false <- {0, 3} [subsumption resolution 331,206]
334. ~0 | 3 [avatar contradiction clause 333]
335. $false [avatar sat refutation 207,208,325,334]
% SZS output end Proof for query05_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.014 s
% ------------------------------
% ------------------------------
