% Refutation found. Thanks to Tanya!
% SZS status Theorem for query22_query80_tptp4X
% SZS output start Proof for query22_query80_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X5,X6] : (where_5(X5,X6) <=> eq(X5,X6)) [input]
3. ! [X5,X6] : (select_3(X5,X6) <=> ? [X0,X1,X2,X3,X4] : (where_5(X5,X6) & from_4(X0,X1,X2,X3,X4))) [input]
4. ! [X7,X6] : (query_2(X7,X6) <=> select_3(X7,X6)) [input]
5. ! [X8,X9] : (from_6(X8,X9) <=> query_2(X8,X9)) [input]
6. ! [X8] : (select_1(X8) <=> ? [X9] : ($true & from_6(X8,X9))) [input]
7. ! [X7] : (main_query_1(X7) <=> select_1(X7)) [input]
8. ! [X10,X11,X12,X13,X14] : (from_8(X10,X11,X12,X13,X14) <=> tab1(X10,X11,X12,X13,X14)) [input]
9. ! [X0] : (select_7(X0) <=> ? [X10,X11,X12,X13,X14] : ($true & from_8(X10,X11,X12,X13,X14))) [input]
10. ! [X15] : (main_query_2(X15) <=> select_7(X15)) [input]
11. ! [X16] : eq(X16,X16) [input]
53. ! [X7] : (main_query_1(X7) <=> main_query_2(X7)) [input]
54. ~! [X7] : (main_query_1(X7) <=> main_query_2(X7)) [negated conjecture 53]
55. ! [X0,X1] : (where_5(X0,X1) <=> eq(X0,X1)) [rectify 2]
56. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4,X5,X6] : (where_5(X0,X1) & from_4(X2,X3,X4,X5,X6))) [rectify 3]
57. ! [X0,X1] : (query_2(X0,X1) <=> select_3(X0,X1)) [rectify 4]
58. ! [X0,X1] : (from_6(X0,X1) <=> query_2(X0,X1)) [rectify 5]
59. ! [X0] : (select_1(X0) <=> ? [X1] : ($true & from_6(X0,X1))) [rectify 6]
60. ! [X0] : (select_1(X0) <=> ? [X1] : from_6(X0,X1)) [true and false elimination 59]
61. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 7]
62. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 8]
63. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4,X5] : ($true & from_8(X1,X2,X3,X4,X5))) [rectify 9]
64. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4,X5] : from_8(X1,X2,X3,X4,X5)) [true and false elimination 63]
65. ! [X0] : (main_query_2(X0) <=> select_7(X0)) [rectify 10]
66. ! [X0] : eq(X0,X0) [rectify 11]
108. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 54]
178. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 108]
179. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 1]
180. ! [X0,X1] : ((where_5(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_5(X0,X1))) [nnf transformation 55]
181. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_5(X0,X1) | ~from_4(X2,X3,X4,X5,X6))) & (? [X2,X3,X4,X5,X6] : (where_5(X0,X1) & from_4(X2,X3,X4,X5,X6)) | ~select_3(X0,X1))) [nnf transformation 56]
182. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_5(X0,X1) | ~from_4(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_5(X0,X1) & from_4(X7,X8,X9,X10,X11)) | ~select_3(X0,X1))) [rectify 181]
183. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_5(X0,X1) & from_4(X7,X8,X9,X10,X11)) => (where_5(X0,X1) & from_4(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)))) [choice axiom]
184. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_5(X0,X1) | ~from_4(X2,X3,X4,X5,X6))) & ((where_5(X0,X1) & from_4(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1))) | ~select_3(X0,X1))) [skolemisation 182,183]
185. ! [X0,X1] : ((query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~query_2(X0,X1))) [nnf transformation 57]
186. ! [X0,X1] : ((from_6(X0,X1) | ~query_2(X0,X1)) & (query_2(X0,X1) | ~from_6(X0,X1))) [nnf transformation 58]
187. ! [X0] : ((select_1(X0) | ! [X1] : ~from_6(X0,X1)) & (? [X1] : from_6(X0,X1) | ~select_1(X0))) [nnf transformation 60]
188. ! [X0] : ((select_1(X0) | ! [X1] : ~from_6(X0,X1)) & (? [X2] : from_6(X0,X2) | ~select_1(X0))) [rectify 187]
189. ! [X0] : (? [X2] : from_6(X0,X2) => from_6(X0,sK5(X0))) [choice axiom]
190. ! [X0] : ((select_1(X0) | ! [X1] : ~from_6(X0,X1)) & (from_6(X0,sK5(X0)) | ~select_1(X0))) [skolemisation 188,189]
191. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 61]
192. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 62]
193. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4,X5] : ~from_8(X1,X2,X3,X4,X5)) & (? [X1,X2,X3,X4,X5] : from_8(X1,X2,X3,X4,X5) | ~select_7(X0))) [nnf transformation 64]
194. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4,X5] : ~from_8(X1,X2,X3,X4,X5)) & (? [X6,X7,X8,X9,X10] : from_8(X6,X7,X8,X9,X10) | ~select_7(X0))) [rectify 193]
195. ? [X6,X7,X8,X9,X10] : from_8(X6,X7,X8,X9,X10) => from_8(sK6,sK7,sK8,sK9,sK10) [choice axiom]
196. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4,X5] : ~from_8(X1,X2,X3,X4,X5)) & (from_8(sK6,sK7,sK8,sK9,sK10) | ~select_7(X0))) [skolemisation 194,195]
197. ! [X0] : ((main_query_2(X0) | ~select_7(X0)) & (select_7(X0) | ~main_query_2(X0))) [nnf transformation 65]
200. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 178]
201. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK11) | ~main_query_1(sK11)) & (main_query_2(sK11) | main_query_1(sK11))) [choice axiom]
202. (~main_query_2(sK11) | ~main_query_1(sK11)) & (main_query_2(sK11) | main_query_1(sK11)) [skolemisation 200,201]
203. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 179]
204. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 179]
206. ~eq(X0,X1) | where_5(X0,X1) [cnf transformation 180]
207. from_4(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_3(X0,X1) [cnf transformation 184]
209. select_3(X0,X1) | ~where_5(X0,X1) | ~from_4(X2,X3,X4,X5,X6) [cnf transformation 184]
210. ~query_2(X0,X1) | select_3(X0,X1) [cnf transformation 185]
211. ~select_3(X0,X1) | query_2(X0,X1) [cnf transformation 185]
212. ~from_6(X0,X1) | query_2(X0,X1) [cnf transformation 186]
213. ~query_2(X0,X1) | from_6(X0,X1) [cnf transformation 186]
214. from_6(X0,sK5(X0)) | ~select_1(X0) [cnf transformation 190]
215. ~from_6(X0,X1) | select_1(X0) [cnf transformation 190]
216. ~main_query_1(X0) | select_1(X0) [cnf transformation 191]
217. ~select_1(X0) | main_query_1(X0) [cnf transformation 191]
218. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 192]
219. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 192]
220. from_8(sK6,sK7,sK8,sK9,sK10) | ~select_7(X0) [cnf transformation 196]
221. select_7(X0) | ~from_8(X1,X2,X3,X4,X5) [cnf transformation 196]
222. ~main_query_2(X0) | select_7(X0) [cnf transformation 197]
223. ~select_7(X0) | main_query_2(X0) [cnf transformation 197]
224. eq(X0,X0) [cnf transformation 66]
264. main_query_2(sK11) | main_query_1(sK11) [cnf transformation 202]
265. ~main_query_2(sK11) | ~main_query_1(sK11) [cnf transformation 202]
299. 0 <=> ! [X3,X5,X2,X4,X6] : ~from_4(X2,X3,X4,X5,X6) [avatar definition]
300. ~from_4(X2,X3,X4,X5,X6) <- {0} [avatar component clause 299]
302. 2 <=> ! [X1,X0] : (select_3(X0,X1) | ~where_5(X0,X1)) [avatar definition]
303. ~where_5(X0,X1) | select_3(X0,X1) <- {2} [avatar component clause 302]
304. 0 | 2 [avatar split clause 209,302,299]
306. 4 <=> ! [X1,X3,X5,X2,X4] : ~from_8(X1,X2,X3,X4,X5) [avatar definition]
307. ~from_8(X1,X2,X3,X4,X5) <- {4} [avatar component clause 306]
309. 6 <=> ! [X0] : select_7(X0) [avatar definition]
310. select_7(X0) <- {6} [avatar component clause 309]
311. 4 | 6 [avatar split clause 221,309,306]
313. 8 <=> ! [X0] : ~select_7(X0) [avatar definition]
314. ~select_7(X0) <- {8} [avatar component clause 313]
319. 10 <=> from_8(sK6,sK7,sK8,sK9,sK10) [avatar definition]
320. from_8(sK6,sK7,sK8,sK9,sK10) <- {10} [avatar component clause 319]
321. 8 | 10 [avatar split clause 220,319,313]
323. 12 <=> main_query_1(sK11) [avatar definition]
324. main_query_1(sK11) <- {12} [avatar component clause 323]
326. 13 <=> ~main_query_1(sK11) [avatar definition]
327. ~main_query_1(sK11) <- {13} [avatar component clause 326]
329. 14 <=> main_query_2(sK11) [avatar definition]
330. main_query_2(sK11) <- {14} [avatar component clause 329]
332. 15 <=> ~main_query_2(sK11) [avatar definition]
333. ~main_query_2(sK11) <- {15} [avatar component clause 332]
334. ~13 | ~15 [avatar split clause 265,332,326]
335. 12 | 14 [avatar split clause 264,329,323]
336. select_1(sK11) <- {12} [resolution 216,324]
338. where_5(X0,X0) [resolution 206,224]
340. query_2(X0,sK5(X0)) | ~select_1(X0) [resolution 214,212]
358. select_3(X1,sK5(X1)) | ~select_1(X1) [resolution 340,210]
366. tab1(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_3(X0,X1) [resolution 207,203]
369. from_8(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_3(X0,X1) [resolution 366,219]
371. ~select_3(X0,X1) <- {4} [subsumption resolution 369,307]
376. ~select_1(X1) <- {4} [subsumption resolution 358,371]
378. $false <- {4, 12} [resolution 376,336]
379. ~4 | ~12 [avatar contradiction clause 378]
380. select_7(sK11) <- {14} [resolution 330,222]
381. $false <- {8, 14} [subsumption resolution 380,314]
382. ~8 | ~14 [avatar contradiction clause 381]
389. tab1(sK6,sK7,sK8,sK9,sK10) <- {10} [resolution 320,218]
391. main_query_2(X0) <- {6} [resolution 310,223]
394. from_4(sK6,sK7,sK8,sK9,sK10) <- {10} [resolution 389,204]
395. $false <- {0, 10} [subsumption resolution 394,300]
396. ~0 | ~10 [avatar contradiction clause 395]
397. select_3(X0,X0) <- {2} [resolution 303,338]
398. query_2(X0,X0) <- {2} [resolution 397,211]
402. from_6(X0,X0) <- {2} [resolution 398,213]
406. select_1(X1) <- {2} [resolution 402,215]
408. main_query_1(X0) <- {2} [resolution 406,217]
413. $false <- {2, 13} [resolution 408,327]
414. ~2 | 13 [avatar contradiction clause 413]
415. $false <- {6, 15} [subsumption resolution 333,391]
416. ~6 | 15 [avatar contradiction clause 415]
417. $false [avatar sat refutation 304,311,321,334,335,379,382,396,414,416]
% SZS output end Proof for query22_query80_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
