% Refutation found. Thanks to Tanya!
% SZS status Theorem for query07_query09_tptp4X
% SZS output start Proof for query07_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X0,X5,X6,X7,X8] : (from_8(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
6. ! [X0] : (select_7(X0) <=> ? [X5,X6,X7,X8] : ($true & from_8(X0,X5,X6,X7,X8))) [input]
7. ! [X0] : (subquery_from_6(X0) <=> select_7(X0)) [input]
8. ! [X9] : (from_5(X9) <=> subquery_from_6(X9)) [input]
9. ! [X9] : (select_4(X9) <=> ($true & from_5(X9))) [input]
10. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [input]
11. ! [X10] : eq(X10,X10) [input]
18. ! [X10,X11] : (lt(X10,X11) <=> (~eq(X10,X11) & lte(X10,X11))) [input]
19. ! [X10,X11] : (gt(X10,X11) <=> lt(X11,X10)) [input]
45. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
46. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 45]
47. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
48. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
49. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 48]
50. ! [X0] : (from_5(X0) <=> subquery_from_6(X0)) [rectify 8]
51. ! [X0] : (select_4(X0) <=> ($true & from_5(X0))) [rectify 9]
52. ! [X0] : (select_4(X0) <=> from_5(X0)) [true and false elimination 51]
53. ! [X0] : eq(X0,X0) [rectify 11]
60. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 18]
61. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 19]
138. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 46]
139. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
140. ! [X0] : ((where_3(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
141. ! [X0] : ((where_3(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [flattening 140]
142. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
143. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 142]
144. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
145. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 143,144]
146. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
147. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 47]
148. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 49]
149. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 148]
150. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
151. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0))) [skolemisation 149,150]
152. ! [X0] : ((subquery_from_6(X0) | ~select_7(X0)) & (select_7(X0) | ~subquery_from_6(X0))) [nnf transformation 7]
153. ! [X0] : ((from_5(X0) | ~subquery_from_6(X0)) & (subquery_from_6(X0) | ~from_5(X0))) [nnf transformation 50]
154. ! [X0] : ((select_4(X0) | ~from_5(X0)) & (from_5(X0) | ~select_4(X0))) [nnf transformation 52]
155. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 10]
158. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 60]
159. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 158]
160. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 61]
161. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 138]
162. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
163. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 161,162]
164. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 139]
165. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 139]
168. gt(X0,X0) | where_3(X0) | lt(X0,X0) [cnf transformation 141]
169. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 145]
171. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 145]
172. ~main_query_1(X0) | select_1(X0) [cnf transformation 146]
173. ~select_1(X0) | main_query_1(X0) [cnf transformation 146]
174. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 147]
175. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 147]
176. from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [cnf transformation 151]
177. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 151]
178. ~subquery_from_6(X0) | select_7(X0) [cnf transformation 152]
179. ~select_7(X0) | subquery_from_6(X0) [cnf transformation 152]
180. ~from_5(X0) | subquery_from_6(X0) [cnf transformation 153]
181. ~subquery_from_6(X0) | from_5(X0) [cnf transformation 153]
182. ~select_4(X0) | from_5(X0) [cnf transformation 154]
183. ~from_5(X0) | select_4(X0) [cnf transformation 154]
184. ~main_query_2(X0) | select_4(X0) [cnf transformation 155]
185. ~select_4(X0) | main_query_2(X0) [cnf transformation 155]
186. eq(X0,X0) [cnf transformation 53]
196. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 159]
198. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 160]
223. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 163]
224. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 163]
226. 0 <=> main_query_1(sK8) [avatar definition]
227. main_query_1(sK8) <- {0} [avatar component clause 226]
229. 1 <=> ~main_query_1(sK8) [avatar definition]
230. ~main_query_1(sK8) <- {1} [avatar component clause 229]
232. 2 <=> main_query_2(sK8) [avatar definition]
233. main_query_2(sK8) <- {2} [avatar component clause 232]
235. 3 <=> ~main_query_2(sK8) [avatar definition]
236. ~main_query_2(sK8) <- {3} [avatar component clause 235]
237. ~1 | ~3 [avatar split clause 224,235,229]
238. 0 | 2 [avatar split clause 223,232,226]
239. select_4(sK8) <- {2} [resolution 184,233]
240. from_5(sK8) <- {2} [resolution 239,182]
243. subquery_from_6(sK8) <- {2} [resolution 240,180]
245. select_7(sK8) <- {2} [resolution 243,178]
251. ~lt(X0,X0) [resolution 196,186]
253. where_3(X1) | lt(X1,X1) | lt(X1,X1) [resolution 168,198]
254. where_3(X1) | lt(X1,X1) [duplicate literal removal 253]
255. where_3(X1) [subsumption resolution 254,251]
319. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 169,164]
359. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [resolution 176,174]
442. from_8(X6,sK0(X6),sK1(X6),sK2(X6),sK3(X6)) | ~select_1(X6) [resolution 319,175]
450. from_2(X9,sK4(X9),sK5(X9),sK6(X9),sK7(X9)) | ~select_7(X9) [resolution 359,165]
454. ~select_1(X1) | select_7(X1) [resolution 442,177]
461. ~select_7(X1) | ~where_3(X1) | select_1(X1) [resolution 450,171]
462. ~select_7(X1) | select_1(X1) [subsumption resolution 461,255]
463. select_1(sK8) <- {2} [resolution 462,245]
465. main_query_1(sK8) <- {2} [resolution 463,173]
467. $false <- {1, 2} [subsumption resolution 465,230]
468. 1 | ~2 [avatar contradiction clause 467]
469. select_1(sK8) <- {0} [resolution 227,172]
470. select_7(sK8) <- {0} [resolution 469,454]
476. subquery_from_6(sK8) <- {0} [resolution 470,179]
479. from_5(sK8) <- {0} [resolution 476,181]
483. select_4(sK8) <- {0} [resolution 479,183]
487. main_query_2(sK8) <- {0} [resolution 483,185]
489. $false <- {0, 3} [subsumption resolution 487,236]
490. ~0 | 3 [avatar contradiction clause 489]
491. $false [avatar sat refutation 237,238,468,490]
% SZS output end Proof for query07_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.031 s
% ------------------------------
% ------------------------------
