% Refutation found. Thanks to Tanya!
% SZS status Theorem for query06_query08_tptp4X
% SZS output start Proof for query06_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : ($true & from_5(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (subquery_from_3(X5) <=> select_4(X5)) [input]
4. ! [X6] : (from_2(X6) <=> subquery_from_3(X6)) [input]
5. ! [X6] : (select_1(X6) <=> ($true & from_2(X6))) [input]
6. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
7. ! [X7,X8,X9,X10,X11] : (from_7(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
8. ! [X7] : (where_8(X7) <=> (gte(X7,X7) & lte(X7,X7))) [input]
9. ! [X7] : (select_6(X7) <=> ? [X8,X9,X10,X11] : (where_8(X7) & from_7(X7,X8,X9,X10,X11))) [input]
10. ! [X5] : (main_query_2(X5) <=> select_6(X5)) [input]
14. ! [X12] : lte(X12,X12) [input]
21. ! [X12,X13] : (gte(X12,X13) <=> lte(X13,X12)) [input]
45. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
46. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 45]
47. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4)) [true and false elimination 2]
48. ! [X0] : (subquery_from_3(X0) <=> select_4(X0)) [rectify 3]
49. ! [X0] : (from_2(X0) <=> subquery_from_3(X0)) [rectify 4]
50. ! [X0] : (select_1(X0) <=> ($true & from_2(X0))) [rectify 5]
51. ! [X0] : (select_1(X0) <=> from_2(X0)) [true and false elimination 50]
52. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 6]
53. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 7]
54. ! [X0] : (where_8(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 8]
55. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4))) [rectify 9]
56. ! [X0] : (main_query_2(X0) <=> select_6(X0)) [rectify 10]
60. ! [X0] : lte(X0,X0) [rectify 14]
67. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 21]
91. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 46]
143. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 91]
144. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 1]
145. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4) | ~select_4(X0))) [nnf transformation 47]
146. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) | ~select_4(X0))) [rectify 145]
147. ! [X0] : (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) => from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
148. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0))) [skolemisation 146,147]
149. ! [X0] : ((subquery_from_3(X0) | ~select_4(X0)) & (select_4(X0) | ~subquery_from_3(X0))) [nnf transformation 48]
150. ! [X0] : ((from_2(X0) | ~subquery_from_3(X0)) & (subquery_from_3(X0) | ~from_2(X0))) [nnf transformation 49]
151. ! [X0] : ((select_1(X0) | ~from_2(X0)) & (from_2(X0) | ~select_1(X0))) [nnf transformation 51]
152. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 52]
153. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 53]
154. ! [X0] : ((where_8(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_8(X0))) [nnf transformation 54]
155. ! [X0] : ((where_8(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_8(X0))) [flattening 154]
156. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4)) | ~select_6(X0))) [nnf transformation 55]
157. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) | ~select_6(X0))) [rectify 156]
158. ! [X0] : (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) => (where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
159. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & ((where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_6(X0))) [skolemisation 157,158]
160. ! [X0] : ((main_query_2(X0) | ~select_6(X0)) & (select_6(X0) | ~main_query_2(X0))) [nnf transformation 56]
163. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 67]
164. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 143]
165. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
166. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 164,165]
167. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 144]
168. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 144]
169. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [cnf transformation 148]
170. ~from_5(X0,X1,X2,X3,X4) | select_4(X0) [cnf transformation 148]
171. ~subquery_from_3(X0) | select_4(X0) [cnf transformation 149]
172. ~select_4(X0) | subquery_from_3(X0) [cnf transformation 149]
173. ~from_2(X0) | subquery_from_3(X0) [cnf transformation 150]
174. ~subquery_from_3(X0) | from_2(X0) [cnf transformation 150]
175. ~select_1(X0) | from_2(X0) [cnf transformation 151]
176. ~from_2(X0) | select_1(X0) [cnf transformation 151]
177. ~main_query_1(X0) | select_1(X0) [cnf transformation 152]
178. ~select_1(X0) | main_query_1(X0) [cnf transformation 152]
179. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 153]
180. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 153]
183. ~gte(X0,X0) | where_8(X0) | ~lte(X0,X0) [cnf transformation 155]
184. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 159]
186. ~from_7(X0,X1,X2,X3,X4) | ~where_8(X0) | select_6(X0) [cnf transformation 159]
187. ~main_query_2(X0) | select_6(X0) [cnf transformation 160]
188. ~select_6(X0) | main_query_2(X0) [cnf transformation 160]
192. lte(X0,X0) [cnf transformation 60]
199. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 163]
223. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 166]
224. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 166]
226. 0 <=> main_query_1(sK8) [avatar definition]
227. main_query_1(sK8) <- {0} [avatar component clause 226]
229. 1 <=> ~main_query_1(sK8) [avatar definition]
230. ~main_query_1(sK8) <- {1} [avatar component clause 229]
232. 2 <=> main_query_2(sK8) [avatar definition]
233. main_query_2(sK8) <- {2} [avatar component clause 232]
235. 3 <=> ~main_query_2(sK8) [avatar definition]
236. ~main_query_2(sK8) <- {3} [avatar component clause 235]
237. ~1 | ~3 [avatar split clause 224,235,229]
238. 0 | 2 [avatar split clause 223,232,226]
239. select_6(sK8) <- {2} [resolution 187,233]
248. gte(X1,X1) [resolution 199,192]
252. where_8(X0) | ~lte(X0,X0) [resolution 183,248]
254. where_8(X0) [subsumption resolution 252,192]
320. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [resolution 169,167]
322. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 184,179]
324. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [resolution 320,180]
328. from_5(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 322,168]
331. ~select_4(X1) | ~where_8(X1) | select_6(X1) [resolution 324,186]
332. ~select_4(X1) | select_6(X1) [subsumption resolution 331,254]
340. ~select_6(X9) | select_4(X9) [resolution 328,170]
343. select_4(sK8) <- {2} [resolution 340,239]
347. subquery_from_3(sK8) <- {2} [resolution 343,172]
350. from_2(sK8) <- {2} [resolution 347,174]
354. select_1(sK8) <- {2} [resolution 350,176]
358. main_query_1(sK8) <- {2} [resolution 354,178]
360. $false <- {1, 2} [subsumption resolution 358,230]
361. 1 | ~2 [avatar contradiction clause 360]
362. select_1(sK8) <- {0} [resolution 227,177]
366. from_2(sK8) <- {0} [resolution 362,175]
368. subquery_from_3(sK8) <- {0} [resolution 366,173]
370. select_4(sK8) <- {0} [resolution 368,171]
371. select_6(sK8) <- {0} [resolution 370,332]
374. main_query_2(sK8) <- {0} [resolution 371,188]
376. $false <- {0, 3} [subsumption resolution 374,236]
377. ~0 | 3 [avatar contradiction clause 376]
378. $false [avatar sat refutation 237,238,361,377]
% SZS output end Proof for query06_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.028 s
% ------------------------------
% ------------------------------
