% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query08_tptp4X
% SZS output start Proof for query00_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
6. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
11. ! [X6] : lte(X6,X6) [input]
18. ! [X6,X7] : (gte(X6,X7) <=> lte(X7,X6)) [input]
39. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
40. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 39]
41. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
42. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
43. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
47. ! [X0] : lte(X0,X0) [rectify 11]
54. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 18]
75. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 40]
121. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 75]
122. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
123. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 41]
124. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 123]
125. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
126. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 124,125]
127. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 42]
128. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
129. ! [X0] : ((where_5(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [nnf transformation 5]
130. ! [X0] : ((where_5(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [flattening 129]
131. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 6]
132. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 131]
133. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
134. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 132,133]
135. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 43]
138. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 54]
139. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 121]
140. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
141. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 139,140]
142. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 122]
143. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 122]
144. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 126]
145. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 126]
146. ~main_query_1(X0) | select_1(X0) [cnf transformation 127]
147. ~select_1(X0) | main_query_1(X0) [cnf transformation 127]
148. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 128]
149. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 128]
152. ~gte(X0,X0) | where_5(X0) | ~lte(X0,X0) [cnf transformation 130]
153. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 134]
155. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 134]
156. ~main_query_2(X0) | select_3(X0) [cnf transformation 135]
157. ~select_3(X0) | main_query_2(X0) [cnf transformation 135]
161. lte(X0,X0) [cnf transformation 47]
168. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 138]
189. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 141]
190. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 141]
192. 0 <=> main_query_1(sK8) [avatar definition]
193. main_query_1(sK8) <- {0} [avatar component clause 192]
195. 1 <=> ~main_query_1(sK8) [avatar definition]
196. ~main_query_1(sK8) <- {1} [avatar component clause 195]
198. 2 <=> main_query_2(sK8) [avatar definition]
199. main_query_2(sK8) <- {2} [avatar component clause 198]
201. 3 <=> ~main_query_2(sK8) [avatar definition]
202. ~main_query_2(sK8) <- {3} [avatar component clause 201]
203. ~1 | ~3 [avatar split clause 190,201,195]
204. 0 | 2 [avatar split clause 189,198,192]
205. select_3(sK8) <- {2} [resolution 156,199]
214. gte(X1,X1) [resolution 168,161]
218. where_5(X0) | ~lte(X0,X0) [resolution 152,214]
220. where_5(X0) [subsumption resolution 218,161]
265. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 144,142]
288. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 153,148]
291. from_4(X2,sK0(X2),sK1(X2),sK2(X2),sK3(X2)) | ~select_1(X2) [resolution 265,149]
297. from_2(X5,sK4(X5),sK5(X5),sK6(X5),sK7(X5)) | ~select_3(X5) [resolution 288,143]
301. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 291,155]
302. ~select_1(X1) | select_3(X1) [subsumption resolution 301,220]
308. ~select_3(X1) | select_1(X1) [resolution 297,145]
311. select_1(sK8) <- {2} [resolution 308,205]
315. main_query_1(sK8) <- {2} [resolution 311,147]
316. $false <- {1, 2} [subsumption resolution 315,196]
317. 1 | ~2 [avatar contradiction clause 316]
318. select_1(sK8) <- {0} [resolution 193,146]
321. select_3(sK8) <- {0} [resolution 318,302]
326. main_query_2(sK8) <- {0} [resolution 321,157]
328. $false <- {0, 3} [subsumption resolution 326,202]
329. ~0 | 3 [avatar contradiction clause 328]
330. $false [avatar sat refutation 203,204,317,329]
% SZS output end Proof for query00_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
