% Refutation found. Thanks to Tanya!
% SZS status Theorem for query02_query07_tptp4X
% SZS output start Proof for query02_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X5,X6,X7,X8,X9] : (from_4(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
5. ! [X5] : (where_5(X5) <=> (~gt(X5,X5) & ~lt(X5,X5))) [input]
6. ! [X5] : (select_3(X5) <=> ? [X6,X7,X8,X9] : (where_5(X5) & from_4(X5,X6,X7,X8,X9))) [input]
7. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
8. ! [X10] : eq(X10,X10) [input]
14. ! [X10,X11] : (lt(X10,X11) <=> (~eq(X10,X11) & lte(X10,X11))) [input]
15. ! [X10,X11] : (gt(X10,X11) <=> lt(X11,X10)) [input]
41. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
42. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
45. ! [X0] : (where_5(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [rectify 5]
46. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
47. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
48. ! [X0] : eq(X0,X0) [rectify 8]
54. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 14]
55. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 15]
132. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 42]
133. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 134]
136. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 135,136]
138. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
139. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 44]
140. ! [X0] : ((where_5(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [nnf transformation 45]
141. ! [X0] : ((where_5(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [flattening 140]
142. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 46]
143. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 142]
144. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
145. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 143,144]
146. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 47]
149. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 54]
150. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 149]
151. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 55]
152. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 132]
153. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
154. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 152,153]
155. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 133]
156. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 133]
157. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 137]
158. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 137]
159. ~main_query_1(X0) | select_1(X0) [cnf transformation 138]
160. ~select_1(X0) | main_query_1(X0) [cnf transformation 138]
161. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 139]
162. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 139]
165. gt(X0,X0) | where_5(X0) | lt(X0,X0) [cnf transformation 141]
166. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 145]
168. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 145]
169. ~main_query_2(X0) | select_3(X0) [cnf transformation 146]
170. ~select_3(X0) | main_query_2(X0) [cnf transformation 146]
171. eq(X0,X0) [cnf transformation 48]
180. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 150]
182. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 151]
207. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 154]
208. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 154]
233. 0 <=> main_query_1(sK8) [avatar definition]
234. main_query_1(sK8) <- {0} [avatar component clause 233]
236. 1 <=> ~main_query_1(sK8) [avatar definition]
237. ~main_query_1(sK8) <- {1} [avatar component clause 236]
239. 2 <=> main_query_2(sK8) [avatar definition]
240. main_query_2(sK8) <- {2} [avatar component clause 239]
242. 3 <=> ~main_query_2(sK8) [avatar definition]
243. ~main_query_2(sK8) <- {3} [avatar component clause 242]
244. ~1 | ~3 [avatar split clause 208,242,236]
245. 0 | 2 [avatar split clause 207,239,233]
246. select_3(sK8) <- {2} [resolution 169,240]
252. ~lt(X0,X0) [resolution 180,171]
254. where_5(X1) | lt(X1,X1) | lt(X1,X1) [resolution 165,182]
255. where_5(X1) | lt(X1,X1) [duplicate literal removal 254]
256. where_5(X1) [subsumption resolution 255,252]
261. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 157,155]
263. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 166,161]
265. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 261,162]
268. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 263,156]
270. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 265,168]
271. ~select_1(X1) | select_3(X1) [subsumption resolution 270,256]
273. ~select_3(X1) | select_1(X1) [resolution 268,158]
274. select_1(sK8) <- {2} [resolution 273,246]
276. main_query_1(sK8) <- {2} [resolution 274,160]
277. $false <- {1, 2} [subsumption resolution 276,237]
278. 1 | ~2 [avatar contradiction clause 277]
279. select_1(sK8) <- {0} [resolution 234,159]
280. select_3(sK8) <- {0} [resolution 279,271]
283. main_query_2(sK8) <- {0} [resolution 280,170]
285. $false <- {0, 3} [subsumption resolution 283,243]
286. ~0 | 3 [avatar contradiction clause 285]
287. $false [avatar sat refutation 244,245,278,286]
% SZS output end Proof for query02_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
