% Refutation found. Thanks to Tanya!
% SZS status Theorem for query01_query03_tptp4X
% SZS output start Proof for query01_query03_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X5,X6,X7,X8] : (from_4(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
5. ! [X0] : (select_3(X0) <=> ? [X5,X6,X7,X8] : ($true & from_4(X0,X5,X6,X7,X8))) [input]
6. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
41. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
42. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
45. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [rectify 5]
46. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 45]
132. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 42]
133. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 134]
136. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 135,136]
138. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
139. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 44]
140. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 46]
141. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 140]
142. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
143. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0))) [skolemisation 141,142]
144. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 6]
147. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 132]
148. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
149. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 147,148]
150. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 133]
151. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 133]
152. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 137]
153. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 137]
154. ~main_query_1(X0) | select_1(X0) [cnf transformation 138]
155. ~select_1(X0) | main_query_1(X0) [cnf transformation 138]
156. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 139]
157. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 139]
158. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 143]
159. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 143]
160. ~main_query_2(X0) | select_3(X0) [cnf transformation 144]
161. ~select_3(X0) | main_query_2(X0) [cnf transformation 144]
194. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 149]
195. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 149]
197. 0 <=> main_query_1(sK8) [avatar definition]
198. main_query_1(sK8) <- {0} [avatar component clause 197]
200. 1 <=> ~main_query_1(sK8) [avatar definition]
201. ~main_query_1(sK8) <- {1} [avatar component clause 200]
203. 2 <=> main_query_2(sK8) [avatar definition]
204. main_query_2(sK8) <- {2} [avatar component clause 203]
206. 3 <=> ~main_query_2(sK8) [avatar definition]
207. ~main_query_2(sK8) <- {3} [avatar component clause 206]
208. ~1 | ~3 [avatar split clause 195,206,200]
209. 0 | 2 [avatar split clause 194,203,197]
210. select_3(sK8) <- {2} [resolution 160,204]
241. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 152,150]
243. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 158,156]
245. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 241,157]
248. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 243,151]
250. ~select_1(X1) | select_3(X1) [resolution 245,159]
252. ~select_3(X1) | select_1(X1) [resolution 248,153]
255. select_1(sK8) <- {2} [resolution 252,210]
259. main_query_1(sK8) <- {2} [resolution 255,155]
260. $false <- {1, 2} [subsumption resolution 259,201]
261. 1 | ~2 [avatar contradiction clause 260]
262. select_1(sK8) <- {0} [resolution 198,154]
265. select_3(sK8) <- {0} [resolution 262,250]
270. main_query_2(sK8) <- {0} [resolution 265,161]
271. $false <- {0, 3} [subsumption resolution 270,207]
272. ~0 | 3 [avatar contradiction clause 271]
273. $false [avatar sat refutation 208,209,261,272]
% SZS output end Proof for query01_query03_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.015 s
% ------------------------------
% ------------------------------
