% Refutation found. Thanks to Tanya!
% SZS status Theorem for query11_query12_tptp4X
% SZS output start Proof for query11_query12_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [input]
5. ! [X0,X1,X5,X6,X7] : (from_5(X0,X1,X5,X6,X7) <=> tab1(X0,X1,X5,X6,X7)) [input]
6. ! [X0,X1] : (where_6(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [input]
7. ! [X0,X1] : (select_4(X0,X1) <=> ? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7))) [input]
8. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [input]
13. ! [X8,X9] : ((lte(X9,X8) & lte(X8,X9)) <=> eq(X8,X9)) [input]
54. ! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [input]
55. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [negated conjecture 54]
56. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
57. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
62. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
178. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 55]
179. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
180. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
181. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
182. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 181]
183. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
184. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 182,183]
185. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 4]
186. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 56]
187. ! [X0,X1] : ((where_6(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 6]
188. ! [X0,X1] : ((where_6(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [flattening 187]
189. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 57]
190. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 189]
191. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
192. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 190,191]
193. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 8]
194. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 62]
195. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 194]
196. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 178]
197. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
198. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 196,197]
199. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 179]
200. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 179]
201. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 180]
202. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 180]
203. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 184]
204. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 184]
205. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 184]
206. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 185]
207. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 185]
208. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 186]
209. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 186]
210. ~where_6(X0,X1) | lte(X0,X1) [cnf transformation 188]
211. ~where_6(X0,X1) | lte(X1,X0) [cnf transformation 188]
212. ~lte(X1,X0) | where_6(X0,X1) | ~lte(X0,X1) [cnf transformation 188]
213. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 192]
214. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 192]
215. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 192]
216. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 193]
217. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 193]
222. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 195]
223. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 195]
224. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 195]
261. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 198]
262. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 198]
299. 0 <=> main_query_1(sK6,sK7) [avatar definition]
300. main_query_1(sK6,sK7) <- {0} [avatar component clause 299]
302. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
303. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 302]
305. 2 <=> main_query_2(sK6,sK7) [avatar definition]
306. main_query_2(sK6,sK7) <- {2} [avatar component clause 305]
308. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
309. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 308]
310. ~1 | ~3 [avatar split clause 262,308,302]
311. 0 | 2 [avatar split clause 261,305,299]
314. select_4(sK6,sK7) <- {2} [resolution 216,306]
315. where_6(sK6,sK7) <- {2} [resolution 314,214]
317. lte(sK7,sK6) <- {2} [resolution 315,211]
318. lte(sK6,sK7) <- {2} [resolution 315,210]
331. eq(sK6,sK7) | ~lte(sK6,sK7) <- {2} [resolution 222,317]
333. eq(sK6,sK7) <- {2} [subsumption resolution 331,318]
344. where_3(sK6,sK7) <- {2} [resolution 333,202]
359. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 203,199]
363. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 213,208]
364. from_5(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 359,209]
367. from_2(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 363,200]
368. ~where_6(X0,X1) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 364,215]
370. ~select_4(X0,X1) | ~where_3(X0,X1) | select_1(X0,X1) [resolution 367,205]
388. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 370,314]
389. select_1(sK6,sK7) <- {2} [subsumption resolution 388,344]
390. main_query_1(sK6,sK7) <- {2} [resolution 389,207]
392. $false <- {1, 2} [subsumption resolution 390,303]
393. 1 | ~2 [avatar contradiction clause 392]
394. select_1(sK6,sK7) <- {0} [resolution 300,206]
396. where_3(sK6,sK7) <- {0} [resolution 394,204]
397. eq(sK6,sK7) <- {0} [resolution 396,201]
399. lte(sK7,sK6) <- {0} [resolution 397,224]
400. lte(sK6,sK7) <- {0} [resolution 397,223]
405. where_6(sK6,sK7) | ~lte(sK6,sK7) <- {0} [resolution 399,212]
406. where_6(sK6,sK7) <- {0} [subsumption resolution 405,400]
418. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 406,368]
421. select_4(sK6,sK7) <- {0} [subsumption resolution 418,394]
429. main_query_2(sK6,sK7) <- {0} [resolution 421,217]
431. $false <- {0, 3} [subsumption resolution 429,309]
432. ~0 | 3 [avatar contradiction clause 431]
433. $false [avatar sat refutation 310,311,393,432]
% SZS output end Proof for query11_query12_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.020 s
% ------------------------------
% ------------------------------
