% Refutation found. Thanks to Tanya!
% SZS status Theorem for query12_query14_tptp4X
% SZS output start Proof for query12_query14_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X5,X2,X3,X4] : (from_9(X0,X5,X2,X3,X4) <=> tab1(X0,X5,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X5,X2,X3,X4] : ($true & from_9(X0,X5,X2,X3,X4))) [input]
8. ! [X0] : (where_6_subquery_7(X0) <=> select_8(X0)) [input]
9. ! [X0,X1] : (where_6(X0,X1) <=> (where_6_subquery_7(X0) & eq(X0,X1))) [input]
10. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [input]
16. ! [X6,X7] : ((lte(X7,X6) & lte(X6,X7)) <=> eq(X6,X7)) [input]
62. ! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [input]
63. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [negated conjecture 62]
64. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 6]
65. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [rectify 7]
66. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 65]
71. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 16]
202. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 63]
203. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
204. ! [X0,X1] : ((where_3(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [nnf transformation 2]
205. ! [X0,X1] : ((where_3(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [flattening 204]
206. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
207. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 206]
208. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
209. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 207,208]
210. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 4]
211. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
212. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 64]
213. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 66]
214. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 213]
215. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
216. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 214,215]
217. ! [X0] : ((where_6_subquery_7(X0) | ~select_8(X0)) & (select_8(X0) | ~where_6_subquery_7(X0))) [nnf transformation 8]
218. ! [X0,X1] : ((where_6(X0,X1) | (~where_6_subquery_7(X0) | ~eq(X0,X1))) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 9]
219. ! [X0,X1] : ((where_6(X0,X1) | ~where_6_subquery_7(X0) | ~eq(X0,X1)) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 218]
220. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 10]
221. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 220]
222. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
223. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 221,222]
224. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 11]
225. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 71]
226. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 225]
227. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 202]
228. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
229. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 227,228]
230. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 203]
231. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 203]
232. ~where_3(X0,X1) | lte(X0,X1) [cnf transformation 205]
233. ~where_3(X0,X1) | lte(X1,X0) [cnf transformation 205]
234. ~lte(X1,X0) | where_3(X0,X1) | ~lte(X0,X1) [cnf transformation 205]
235. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 209]
236. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 209]
237. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 209]
238. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 210]
239. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 210]
240. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 211]
241. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 211]
243. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 212]
245. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 216]
247. ~select_8(X0) | where_6_subquery_7(X0) [cnf transformation 217]
248. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 219]
250. ~eq(X0,X1) | ~where_6_subquery_7(X0) | where_6(X0,X1) [cnf transformation 219]
251. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 223]
252. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 223]
253. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 223]
254. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 224]
255. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 224]
260. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 226]
261. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 226]
262. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 226]
304. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 229]
305. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 229]
347. 0 <=> main_query_1(sK10,sK11) [avatar definition]
348. main_query_1(sK10,sK11) <- {0} [avatar component clause 347]
350. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
351. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 350]
353. 2 <=> main_query_2(sK10,sK11) [avatar definition]
354. main_query_2(sK10,sK11) <- {2} [avatar component clause 353]
356. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
357. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 356]
358. ~1 | ~3 [avatar split clause 305,356,350]
359. 0 | 2 [avatar split clause 304,353,347]
360. select_4(sK10,sK11) <- {2} [resolution 254,354]
362. where_6(sK10,sK11) <- {2} [resolution 360,252]
364. eq(sK10,sK11) <- {2} [resolution 362,248]
370. lte(sK11,sK10) <- {2} [resolution 364,262]
371. lte(sK10,sK11) <- {2} [resolution 364,261]
376. where_3(sK10,sK11) | ~lte(sK10,sK11) <- {2} [resolution 234,370]
378. where_3(sK10,sK11) <- {2} [subsumption resolution 376,371]
426. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 235,230]
431. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 251,240]
434. from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 426,243]
435. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 426,241]
439. from_2(X4,X5,sK7(X4,X5),sK8(X4,X5),sK9(X4,X5)) | ~select_4(X4,X5) [resolution 431,231]
443. ~select_1(X2,X3) | select_8(X2) [resolution 434,245]
444. ~where_6(X0,X1) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 435,253]
451. ~select_4(X0,X1) | ~where_3(X0,X1) | select_1(X0,X1) [resolution 439,237]
453. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 451,360]
454. select_1(sK10,sK11) <- {2} [subsumption resolution 453,378]
456. main_query_1(sK10,sK11) <- {2} [resolution 454,239]
458. $false <- {1, 2} [subsumption resolution 456,351]
459. 1 | ~2 [avatar contradiction clause 458]
460. select_1(sK10,sK11) <- {0} [resolution 348,238]
461. select_8(sK10) <- {0} [resolution 460,443]
463. where_3(sK10,sK11) <- {0} [resolution 460,236]
464. where_6_subquery_7(sK10) <- {0} [resolution 461,247]
465. lte(sK11,sK10) <- {0} [resolution 463,233]
466. lte(sK10,sK11) <- {0} [resolution 463,232]
469. eq(sK10,sK11) | ~lte(sK10,sK11) <- {0} [resolution 465,260]
471. eq(sK10,sK11) <- {0} [subsumption resolution 469,466]
478. ~where_6_subquery_7(sK10) | where_6(sK10,sK11) <- {0} [resolution 471,250]
482. where_6(sK10,sK11) <- {0} [subsumption resolution 478,464]
492. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {0} [resolution 482,444]
495. select_4(sK10,sK11) <- {0} [subsumption resolution 492,460]
500. main_query_2(sK10,sK11) <- {0} [resolution 495,255]
502. $false <- {0, 3} [subsumption resolution 500,357]
503. ~0 | 3 [avatar contradiction clause 502]
504. $false [avatar sat refutation 358,359,459,503]
% SZS output end Proof for query12_query14_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.020 s
% ------------------------------
% ------------------------------
