% Refutation found. Thanks to Tanya!
% SZS status Theorem for query11_query12_tptp4X
% SZS output start Proof for query11_query12_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X1,X0] : (where_3(X1,X0) <=> eq(X0,X1)) [input]
3. ! [X1,X0] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X1,X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X8,X7] : (where_6(X8,X7) <=> (lte(X8,X7) & lte(X7,X8))) [input]
7. ! [X8,X7] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X8,X7) & from_5(X7,X8,X9,X10,X11))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
13. ! [X12,X13] : ((lte(X13,X12) & lte(X12,X13)) <=> eq(X12,X13)) [input]
54. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
55. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 54]
56. ! [X0,X1] : (where_3(X0,X1) <=> eq(X1,X0)) [rectify 2]
57. ! [X0,X1] : (select_1(X1,X0) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X1,X0,X2,X3,X4))) [rectify 3]
58. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
59. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
60. ! [X0,X1] : (where_6(X0,X1) <=> (lte(X0,X1) & lte(X1,X0))) [rectify 6]
61. ! [X0,X1] : (select_4(X1,X0) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X1,X0,X2,X3,X4))) [rectify 7]
62. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
67. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
108. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 55]
184. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 108]
185. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
186. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X1,X0)) & (eq(X1,X0) | ~where_3(X0,X1))) [nnf transformation 56]
187. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X1,X0,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X1,X0,X2,X3,X4)) | ~select_1(X1,X0))) [nnf transformation 57]
188. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X1,X0,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X1,X0,X5,X6,X7)) | ~select_1(X1,X0))) [rectify 187]
189. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X1,X0,X5,X6,X7)) => (where_3(X0,X1) & from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
190. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X1,X0,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X1,X0))) [skolemisation 188,189]
191. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 58]
192. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 59]
193. ! [X0,X1] : ((where_6(X0,X1) | (~lte(X0,X1) | ~lte(X1,X0))) & ((lte(X0,X1) & lte(X1,X0)) | ~where_6(X0,X1))) [nnf transformation 60]
194. ! [X0,X1] : ((where_6(X0,X1) | ~lte(X0,X1) | ~lte(X1,X0)) & ((lte(X0,X1) & lte(X1,X0)) | ~where_6(X0,X1))) [flattening 193]
195. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X1,X0,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X1,X0,X2,X3,X4)) | ~select_4(X1,X0))) [nnf transformation 61]
196. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X1,X0,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X1,X0,X5,X6,X7)) | ~select_4(X1,X0))) [rectify 195]
197. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X1,X0,X5,X6,X7)) => (where_6(X0,X1) & from_5(X1,X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
198. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X1,X0,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X1,X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X1,X0))) [skolemisation 196,197]
199. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 62]
200. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 67]
201. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 200]
202. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 184]
203. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
204. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 202,203]
205. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 185]
206. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 185]
207. ~where_3(X0,X1) | eq(X1,X0) [cnf transformation 186]
208. ~eq(X1,X0) | where_3(X0,X1) [cnf transformation 186]
209. from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X1,X0) [cnf transformation 190]
210. ~select_1(X1,X0) | where_3(X0,X1) [cnf transformation 190]
211. ~from_2(X1,X0,X2,X3,X4) | ~where_3(X0,X1) | select_1(X1,X0) [cnf transformation 190]
212. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 191]
213. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 191]
214. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 192]
215. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 192]
216. ~where_6(X0,X1) | lte(X1,X0) [cnf transformation 194]
217. ~where_6(X0,X1) | lte(X0,X1) [cnf transformation 194]
218. ~lte(X1,X0) | ~lte(X0,X1) | where_6(X0,X1) [cnf transformation 194]
219. from_5(X1,X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X1,X0) [cnf transformation 198]
220. ~select_4(X1,X0) | where_6(X0,X1) [cnf transformation 198]
221. ~from_5(X1,X0,X2,X3,X4) | ~where_6(X0,X1) | select_4(X1,X0) [cnf transformation 198]
222. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 199]
223. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 199]
228. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 201]
229. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 201]
230. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 201]
267. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 204]
268. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 204]
305. 0 <=> main_query_1(sK6,sK7) [avatar definition]
306. main_query_1(sK6,sK7) <- {0} [avatar component clause 305]
308. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
309. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 308]
311. 2 <=> main_query_2(sK6,sK7) [avatar definition]
312. main_query_2(sK6,sK7) <- {2} [avatar component clause 311]
314. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
315. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 314]
316. ~1 | ~3 [avatar split clause 268,314,308]
317. 0 | 2 [avatar split clause 267,311,305]
320. select_4(sK6,sK7) <- {2} [resolution 222,312]
321. where_6(sK7,sK6) <- {2} [resolution 320,220]
323. lte(sK7,sK6) <- {2} [resolution 321,217]
324. lte(sK6,sK7) <- {2} [resolution 321,216]
337. eq(sK6,sK7) | ~lte(sK6,sK7) <- {2} [resolution 228,323]
339. eq(sK6,sK7) <- {2} [subsumption resolution 337,324]
350. where_3(sK7,sK6) <- {2} [resolution 339,208]
365. tab1(X2,X3,sK0(X3,X2),sK1(X3,X2),sK2(X3,X2)) | ~select_1(X2,X3) [resolution 209,205]
369. tab1(X2,X3,sK3(X3,X2),sK4(X3,X2),sK5(X3,X2)) | ~select_4(X2,X3) [resolution 219,214]
370. from_5(X0,X1,sK0(X1,X0),sK1(X1,X0),sK2(X1,X0)) | ~select_1(X0,X1) [resolution 365,215]
373. from_2(X2,X3,sK3(X3,X2),sK4(X3,X2),sK5(X3,X2)) | ~select_4(X2,X3) [resolution 369,206]
374. ~where_6(X1,X0) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 370,221]
376. ~select_4(X0,X1) | ~where_3(X1,X0) | select_1(X0,X1) [resolution 373,211]
394. ~where_3(sK7,sK6) | select_1(sK6,sK7) <- {2} [resolution 376,320]
395. select_1(sK6,sK7) <- {2} [subsumption resolution 394,350]
396. main_query_1(sK6,sK7) <- {2} [resolution 395,213]
398. $false <- {1, 2} [subsumption resolution 396,309]
399. 1 | ~2 [avatar contradiction clause 398]
400. select_1(sK6,sK7) <- {0} [resolution 306,212]
402. where_3(sK7,sK6) <- {0} [resolution 400,210]
403. eq(sK6,sK7) <- {0} [resolution 402,207]
405. lte(sK7,sK6) <- {0} [resolution 403,230]
406. lte(sK6,sK7) <- {0} [resolution 403,229]
415. ~lte(sK7,sK6) | where_6(sK7,sK6) <- {0} [resolution 406,218]
416. where_6(sK7,sK6) <- {0} [subsumption resolution 415,405]
427. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 416,374]
430. select_4(sK6,sK7) <- {0} [subsumption resolution 427,400]
435. main_query_2(sK6,sK7) <- {0} [resolution 430,223]
437. $false <- {0, 3} [subsumption resolution 435,315]
438. ~0 | 3 [avatar contradiction clause 437]
439. $false [avatar sat refutation 316,317,399,438]
% SZS output end Proof for query11_query12_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.010 s
% ------------------------------
% ------------------------------
