% Refutation found. Thanks to Tanya!
% SZS status Theorem for query08_query09_tptp4X
% SZS output start Proof for query08_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X6,X7,X8,X9,X10] : (from_7(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
6. ! [X6] : (select_6(X6) <=> ? [X7,X8,X9,X10] : ($true & from_7(X6,X7,X8,X9,X10))) [input]
7. ! [X11] : (query_5(X11) <=> select_6(X11)) [input]
8. ! [X12] : (from_8(X12) <=> query_5(X12)) [input]
9. ! [X12] : (select_4(X12) <=> ($true & from_8(X12))) [input]
10. ! [X11] : (main_query_2(X11) <=> select_4(X11)) [input]
14. ! [X13] : lte(X13,X13) [input]
20. ! [X13,X14] : (gte(X13,X14) <=> lte(X14,X13)) [input]
44. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
45. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 44]
46. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
47. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
48. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [rectify 6]
49. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 48]
50. ! [X0] : (query_5(X0) <=> select_6(X0)) [rectify 7]
51. ! [X0] : (from_8(X0) <=> query_5(X0)) [rectify 8]
52. ! [X0] : (select_4(X0) <=> ($true & from_8(X0))) [rectify 9]
53. ! [X0] : (select_4(X0) <=> from_8(X0)) [true and false elimination 52]
54. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 10]
58. ! [X0] : lte(X0,X0) [rectify 14]
64. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 20]
88. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 45]
140. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 88]
141. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
142. ! [X0] : ((where_3(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
143. ! [X0] : ((where_3(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_3(X0))) [flattening 142]
144. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
145. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 144]
146. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
147. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 145,146]
148. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 46]
149. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 47]
150. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 49]
151. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 150]
152. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
153. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 151,152]
154. ! [X0] : ((query_5(X0) | ~select_6(X0)) & (select_6(X0) | ~query_5(X0))) [nnf transformation 50]
155. ! [X0] : ((from_8(X0) | ~query_5(X0)) & (query_5(X0) | ~from_8(X0))) [nnf transformation 51]
156. ! [X0] : ((select_4(X0) | ~from_8(X0)) & (from_8(X0) | ~select_4(X0))) [nnf transformation 53]
157. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 54]
160. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 64]
161. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 140]
162. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
163. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 161,162]
164. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 141]
165. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 141]
168. ~gte(X0,X0) | where_3(X0) | ~lte(X0,X0) [cnf transformation 143]
169. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 147]
171. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 147]
172. ~main_query_1(X0) | select_1(X0) [cnf transformation 148]
173. ~select_1(X0) | main_query_1(X0) [cnf transformation 148]
174. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 149]
175. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 149]
176. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 153]
177. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 153]
178. ~query_5(X0) | select_6(X0) [cnf transformation 154]
179. ~select_6(X0) | query_5(X0) [cnf transformation 154]
180. ~from_8(X0) | query_5(X0) [cnf transformation 155]
181. ~query_5(X0) | from_8(X0) [cnf transformation 155]
182. ~select_4(X0) | from_8(X0) [cnf transformation 156]
183. ~from_8(X0) | select_4(X0) [cnf transformation 156]
184. ~main_query_2(X0) | select_4(X0) [cnf transformation 157]
185. ~select_4(X0) | main_query_2(X0) [cnf transformation 157]
189. lte(X0,X0) [cnf transformation 58]
195. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 160]
219. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 163]
220. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 163]
245. 0 <=> main_query_1(sK8) [avatar definition]
246. main_query_1(sK8) <- {0} [avatar component clause 245]
248. 1 <=> ~main_query_1(sK8) [avatar definition]
249. ~main_query_1(sK8) <- {1} [avatar component clause 248]
251. 2 <=> main_query_2(sK8) [avatar definition]
252. main_query_2(sK8) <- {2} [avatar component clause 251]
254. 3 <=> ~main_query_2(sK8) [avatar definition]
255. ~main_query_2(sK8) <- {3} [avatar component clause 254]
256. ~1 | ~3 [avatar split clause 220,254,248]
257. 0 | 2 [avatar split clause 219,251,245]
258. select_4(sK8) <- {2} [resolution 184,252]
259. from_8(sK8) <- {2} [resolution 258,182]
262. query_5(sK8) <- {2} [resolution 259,180]
264. select_6(sK8) <- {2} [resolution 262,178]
271. gte(X1,X1) [resolution 195,189]
273. where_3(X0) | ~lte(X0,X0) [resolution 168,271]
275. where_3(X0) [subsumption resolution 273,189]
281. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 169,164]
283. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 176,174]
285. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 281,175]
288. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 283,165]
290. ~select_1(X1) | select_6(X1) [resolution 285,177]
292. ~select_6(X1) | ~where_3(X1) | select_1(X1) [resolution 288,171]
293. ~select_6(X1) | select_1(X1) [subsumption resolution 292,275]
294. select_1(sK8) <- {2} [resolution 293,264]
296. main_query_1(sK8) <- {2} [resolution 294,173]
298. $false <- {1, 2} [subsumption resolution 296,249]
299. 1 | ~2 [avatar contradiction clause 298]
300. select_1(sK8) <- {0} [resolution 246,172]
301. select_6(sK8) <- {0} [resolution 300,290]
305. query_5(sK8) <- {0} [resolution 301,179]
306. from_8(sK8) <- {0} [resolution 305,181]
308. select_4(sK8) <- {0} [resolution 306,183]
310. main_query_2(sK8) <- {0} [resolution 308,185]
312. $false <- {0, 3} [subsumption resolution 310,255]
313. ~0 | 3 [avatar contradiction clause 312]
314. $false [avatar sat refutation 256,257,299,313]
% SZS output end Proof for query08_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.016 s
% ------------------------------
% ------------------------------
