% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_07_query_08_tptp4X
% SZS output start Proof for query_07_query_08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_8(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8] : (select_7(X7,X8) <=> ? [X9,X10,X11] : ($true & from_8(X7,X8,X9,X10,X11))) [input]
7. ! [X5,X6] : (subquery_from_6(X5,X6) <=> select_7(X5,X6)) [input]
8. ! [X12,X13] : (from_5(X12,X13) <=> subquery_from_6(X12,X13)) [input]
9. ! [X12,X13] : (select_4(X12,X13) <=> ($true & from_5(X12,X13))) [input]
10. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
11. ! [X14] : eq(X14,X14) [input]
17. ! [X14,X15] : (lt(X14,X15) <=> (~eq(X14,X15) & lte(X14,X15))) [input]
18. ! [X14,X15] : (gt(X14,X15) <=> lt(X15,X14)) [input]
55. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
56. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 55]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
58. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
59. ! [X0,X1] : (select_7(X0,X1) <=> ? [X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
60. ! [X0,X1] : (select_7(X0,X1) <=> ? [X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 59]
61. ! [X0,X1] : (subquery_from_6(X0,X1) <=> select_7(X0,X1)) [rectify 7]
62. ! [X0,X1] : (from_5(X0,X1) <=> subquery_from_6(X0,X1)) [rectify 8]
63. ! [X0,X1] : (select_4(X0,X1) <=> ($true & from_5(X0,X1))) [rectify 9]
64. ! [X0,X1] : (select_4(X0,X1) <=> from_5(X0,X1)) [true and false elimination 63]
65. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 10]
66. ! [X0] : eq(X0,X0) [rectify 11]
72. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 17]
73. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 18]
110. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 56]
184. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 110]
185. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
186. ! [X0] : ((where_3(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
187. ! [X0] : ((where_3(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [flattening 186]
188. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
189. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 188]
190. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
191. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 189,190]
192. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
193. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 58]
194. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0,X1))) [nnf transformation 60]
195. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_8(X0,X1,X5,X6,X7) | ~select_7(X0,X1))) [rectify 194]
196. ! [X1,X0] : (? [X5,X6,X7] : from_8(X0,X1,X5,X6,X7) => from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) [choice axiom]
197. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1))) [skolemisation 195,196]
198. ! [X0,X1] : ((subquery_from_6(X0,X1) | ~select_7(X0,X1)) & (select_7(X0,X1) | ~subquery_from_6(X0,X1))) [nnf transformation 61]
199. ! [X0,X1] : ((from_5(X0,X1) | ~subquery_from_6(X0,X1)) & (subquery_from_6(X0,X1) | ~from_5(X0,X1))) [nnf transformation 62]
200. ! [X0,X1] : ((select_4(X0,X1) | ~from_5(X0,X1)) & (from_5(X0,X1) | ~select_4(X0,X1))) [nnf transformation 64]
201. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 65]
204. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 72]
205. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 204]
206. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 73]
207. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 184]
208. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
209. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 207,208]
210. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 185]
211. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 185]
214. gt(X0,X0) | where_3(X0) | lt(X0,X0) [cnf transformation 187]
215. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 191]
217. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0,X1) [cnf transformation 191]
218. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 192]
219. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 192]
220. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 193]
221. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 193]
222. from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1) [cnf transformation 197]
223. ~from_8(X0,X1,X2,X3,X4) | select_7(X0,X1) [cnf transformation 197]
224. ~subquery_from_6(X0,X1) | select_7(X0,X1) [cnf transformation 198]
225. ~select_7(X0,X1) | subquery_from_6(X0,X1) [cnf transformation 198]
226. ~from_5(X0,X1) | subquery_from_6(X0,X1) [cnf transformation 199]
227. ~subquery_from_6(X0,X1) | from_5(X0,X1) [cnf transformation 199]
228. ~select_4(X0,X1) | from_5(X0,X1) [cnf transformation 200]
229. ~from_5(X0,X1) | select_4(X0,X1) [cnf transformation 200]
230. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 201]
231. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 201]
232. eq(X0,X0) [cnf transformation 66]
241. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 205]
243. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 206]
279. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 209]
280. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 209]
282. 0 <=> main_query_1(sK6,sK7) [avatar definition]
283. main_query_1(sK6,sK7) <- {0} [avatar component clause 282]
285. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
286. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 285]
288. 2 <=> main_query_2(sK6,sK7) [avatar definition]
289. main_query_2(sK6,sK7) <- {2} [avatar component clause 288]
291. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
292. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 291]
293. ~1 | ~3 [avatar split clause 280,291,285]
294. 0 | 2 [avatar split clause 279,288,282]
295. select_4(sK6,sK7) <- {2} [resolution 230,289]
296. from_5(sK6,sK7) <- {2} [resolution 295,228]
299. subquery_from_6(sK6,sK7) <- {2} [resolution 296,226]
302. select_7(sK6,sK7) <- {2} [resolution 299,224]
306. ~lt(X0,X0) [resolution 241,232]
308. where_3(X1) | lt(X1,X1) | lt(X1,X1) [resolution 214,243]
309. where_3(X1) | lt(X1,X1) [duplicate literal removal 308]
310. where_3(X1) [subsumption resolution 309,306]
329. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 215,210]
331. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1) [resolution 222,220]
344. from_8(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 329,221]
350. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1) [resolution 331,211]
365. ~select_1(X2,X3) | select_7(X2,X3) [resolution 344,223]
386. ~select_7(X2,X3) | ~where_3(X2) | select_1(X2,X3) [resolution 350,217]
392. ~select_7(X2,X3) | select_1(X2,X3) [subsumption resolution 386,310]
400. select_1(sK6,sK7) <- {2} [resolution 392,302]
409. main_query_1(sK6,sK7) <- {2} [resolution 400,219]
411. $false <- {1, 2} [subsumption resolution 409,286]
412. 1 | ~2 [avatar contradiction clause 411]
413. select_1(sK6,sK7) <- {0} [resolution 283,218]
414. select_7(sK6,sK7) <- {0} [resolution 413,365]
418. subquery_from_6(sK6,sK7) <- {0} [resolution 414,225]
426. from_5(sK6,sK7) <- {0} [resolution 418,227]
435. select_4(sK6,sK7) <- {0} [resolution 426,229]
444. main_query_2(sK6,sK7) <- {0} [resolution 435,231]
446. $false <- {0, 3} [subsumption resolution 444,292]
447. ~0 | 3 [avatar contradiction clause 446]
448. $false [avatar sat refutation 293,294,412,447]
% SZS output end Proof for query_07_query_08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.036 s
% ------------------------------
% ------------------------------
