% Refutation found. Thanks to Tanya!
% SZS status Theorem for query03_query07_tptp4X
% SZS output start Proof for query03_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X5,X6,X7,X8] : (from_4(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
5. ! [X0] : (where_5(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
6. ! [X0] : (select_3(X0) <=> ? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8))) [input]
7. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
8. ! [X9] : eq(X9,X9) [input]
14. ! [X9,X10] : (lt(X9,X10) <=> (~eq(X9,X10) & lte(X9,X10))) [input]
15. ! [X9,X10] : (gt(X9,X10) <=> lt(X10,X9)) [input]
41. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
42. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
45. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
46. ! [X0] : eq(X0,X0) [rectify 8]
52. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 14]
53. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 15]
130. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 42]
131. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
132. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
133. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 132]
134. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 133,134]
136. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
137. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 44]
138. ! [X0] : ((where_5(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [nnf transformation 5]
139. ! [X0] : ((where_5(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [flattening 138]
140. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 45]
141. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 140]
142. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
143. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 141,142]
144. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 7]
147. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 52]
148. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 147]
149. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 53]
150. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 130]
151. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
152. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 150,151]
153. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 131]
154. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 131]
155. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 135]
156. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 135]
157. ~main_query_1(X0) | select_1(X0) [cnf transformation 136]
158. ~select_1(X0) | main_query_1(X0) [cnf transformation 136]
159. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 137]
160. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 137]
163. gt(X0,X0) | where_5(X0) | lt(X0,X0) [cnf transformation 139]
164. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 143]
166. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 143]
167. ~main_query_2(X0) | select_3(X0) [cnf transformation 144]
168. ~select_3(X0) | main_query_2(X0) [cnf transformation 144]
169. eq(X0,X0) [cnf transformation 46]
178. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 148]
180. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 149]
205. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 152]
206. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 152]
231. 0 <=> main_query_1(sK8) [avatar definition]
232. main_query_1(sK8) <- {0} [avatar component clause 231]
234. 1 <=> ~main_query_1(sK8) [avatar definition]
235. ~main_query_1(sK8) <- {1} [avatar component clause 234]
237. 2 <=> main_query_2(sK8) [avatar definition]
238. main_query_2(sK8) <- {2} [avatar component clause 237]
240. 3 <=> ~main_query_2(sK8) [avatar definition]
241. ~main_query_2(sK8) <- {3} [avatar component clause 240]
242. ~1 | ~3 [avatar split clause 206,240,234]
243. 0 | 2 [avatar split clause 205,237,231]
244. select_3(sK8) <- {2} [resolution 167,238]
250. ~lt(X0,X0) [resolution 178,169]
252. where_5(X1) | lt(X1,X1) | lt(X1,X1) [resolution 163,180]
253. where_5(X1) | lt(X1,X1) [duplicate literal removal 252]
254. where_5(X1) [subsumption resolution 253,250]
259. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 155,153]
261. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 164,159]
263. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 259,160]
266. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 261,154]
268. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 263,166]
269. ~select_1(X1) | select_3(X1) [subsumption resolution 268,254]
271. ~select_3(X1) | select_1(X1) [resolution 266,156]
272. select_1(sK8) <- {2} [resolution 271,244]
274. main_query_1(sK8) <- {2} [resolution 272,158]
275. $false <- {1, 2} [subsumption resolution 274,235]
276. 1 | ~2 [avatar contradiction clause 275]
277. select_1(sK8) <- {0} [resolution 232,157]
278. select_3(sK8) <- {0} [resolution 277,269]
281. main_query_2(sK8) <- {0} [resolution 278,168]
283. $false <- {0, 3} [subsumption resolution 281,241]
284. ~0 | 3 [avatar contradiction clause 283]
285. $false [avatar sat refutation 242,243,276,284]
% SZS output end Proof for query03_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.012 s
% ------------------------------
% ------------------------------
