% Refutation found. Thanks to Tanya!
% SZS status Theorem for query04_query05_tptp4X
% SZS output start Proof for query04_query05_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X5,X1,X2,X3,X4] : (from_5(X5,X1,X2,X3,X4) <=> tab1(X5,X1,X2,X3,X4)) [input]
6. ! [X5] : (where_6(X5) <=> eq(X5,X5)) [input]
7. ! [X5] : (select_4(X5) <=> ? [X1,X2,X3,X4] : (where_6(X5) & from_5(X5,X1,X2,X3,X4))) [input]
8. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
43. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
44. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 43]
45. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
46. ! [X0] : (where_6(X0) <=> eq(X0,X0)) [rectify 6]
47. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
48. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
49. ! [X0] : eq(X0,X0) [rectify 9]
134. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 44]
135. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
136. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 137]
139. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
140. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 138,139]
141. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
142. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 45]
143. ! [X0] : ((where_6(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_6(X0))) [nnf transformation 46]
144. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 47]
145. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 144]
146. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
147. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 145,146]
148. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 48]
151. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 134]
152. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
153. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 151,152]
154. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 135]
155. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 135]
156. eq(X0,X0) | ~where_3(X0) [cnf transformation 136]
157. ~eq(X0,X0) | where_3(X0) [cnf transformation 136]
158. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 140]
160. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 140]
161. ~main_query_1(X0) | select_1(X0) [cnf transformation 141]
162. ~select_1(X0) | main_query_1(X0) [cnf transformation 141]
163. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 142]
164. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 142]
166. ~eq(X0,X0) | where_6(X0) [cnf transformation 143]
167. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 147]
169. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 147]
170. ~main_query_2(X0) | select_4(X0) [cnf transformation 148]
171. ~select_4(X0) | main_query_2(X0) [cnf transformation 148]
172. eq(X0,X0) [cnf transformation 49]
204. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 153]
205. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 153]
207. 0 <=> main_query_1(sK8) [avatar definition]
208. main_query_1(sK8) <- {0} [avatar component clause 207]
210. 1 <=> ~main_query_1(sK8) [avatar definition]
211. ~main_query_1(sK8) <- {1} [avatar component clause 210]
213. 2 <=> main_query_2(sK8) [avatar definition]
214. main_query_2(sK8) <- {2} [avatar component clause 213]
216. 3 <=> ~main_query_2(sK8) [avatar definition]
217. ~main_query_2(sK8) <- {3} [avatar component clause 216]
218. ~1 | ~3 [avatar split clause 205,216,210]
219. 0 | 2 [avatar split clause 204,213,207]
220. select_4(sK8) <- {2} [resolution 170,214]
224. where_3(X1) [resolution 157,172]
227. where_6(X1) | ~where_3(X1) [resolution 166,156]
229. where_6(X1) [subsumption resolution 227,224]
267. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 158,154]
269. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 167,163]
271. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 267,164]
275. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 269,155]
280. ~select_1(X5) | ~where_6(X5) | select_4(X5) [resolution 271,169]
281. ~select_1(X5) | select_4(X5) [subsumption resolution 280,229]
287. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 275,160]
288. ~select_4(X1) | select_1(X1) [subsumption resolution 287,224]
291. select_1(sK8) <- {2} [resolution 288,220]
293. main_query_1(sK8) <- {2} [resolution 291,162]
295. $false <- {1, 2} [subsumption resolution 293,211]
296. 1 | ~2 [avatar contradiction clause 295]
297. select_1(sK8) <- {0} [resolution 208,161]
298. select_4(sK8) <- {0} [resolution 297,281]
302. main_query_2(sK8) <- {0} [resolution 298,171]
304. $false <- {0, 3} [subsumption resolution 302,217]
305. ~0 | 3 [avatar contradiction clause 304]
306. $false [avatar sat refutation 218,219,296,305]
% SZS output end Proof for query04_query05_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.013 s
% ------------------------------
% ------------------------------
