% Refutation found. Thanks to Tanya!
% SZS status Theorem for query10_query13_tptp4X
% SZS output start Proof for query10_query13_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1] : (where_6(X0,X1) <=> (~gt(X0,X1) & ~lt(X0,X1))) [input]
7. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
13. ! [X7,X8] : ((lte(X8,X7) & lte(X7,X8)) <=> eq(X7,X8)) [input]
15. ! [X7,X8] : (lte(X8,X7) | lte(X7,X8)) [input]
16. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
17. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
48. ! [X7,X0,X1] : ((where_6(X0,X1) & eq(X0,X7)) => where_6(X7,X1)) [input]
52. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
53. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 52]
54. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
55. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
56. ! [X0] : eq(X0,X0) [rectify 9]
60. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
62. ! [X0,X1] : (lte(X1,X0) | lte(X0,X1)) [rectify 15]
63. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 16]
64. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 17]
95. ! [X0,X1,X2] : ((where_6(X1,X2) & eq(X1,X0)) => where_6(X0,X2)) [rectify 48]
99. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 53]
161. ! [X0,X1,X2] : (where_6(X0,X2) | (~where_6(X1,X2) | ~eq(X1,X0))) [ennf transformation 95]
162. ! [X0,X1,X2] : (where_6(X0,X2) | ~where_6(X1,X2) | ~eq(X1,X0)) [flattening 161]
169. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 99]
170. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
171. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
172. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
173. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 172]
174. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
175. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 173,174]
176. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 54]
177. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
178. ! [X0,X1] : ((where_6(X0,X1) | (gt(X0,X1) | lt(X0,X1))) & ((~gt(X0,X1) & ~lt(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 6]
179. ! [X0,X1] : ((where_6(X0,X1) | gt(X0,X1) | lt(X0,X1)) & ((~gt(X0,X1) & ~lt(X0,X1)) | ~where_6(X0,X1))) [flattening 178]
180. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 7]
181. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 180]
182. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
183. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 181,182]
184. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 55]
185. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 60]
186. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 185]
187. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 63]
188. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 187]
189. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 64]
190. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 169]
191. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
192. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 190,191]
193. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 170]
194. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 170]
195. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 171]
196. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 171]
197. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 175]
198. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 175]
199. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 175]
200. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 176]
201. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 176]
202. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 177]
203. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 177]
204. ~lt(X0,X1) | ~where_6(X0,X1) [cnf transformation 179]
205. ~gt(X0,X1) | ~where_6(X0,X1) [cnf transformation 179]
206. gt(X0,X1) | where_6(X0,X1) | lt(X0,X1) [cnf transformation 179]
207. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 183]
208. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 183]
209. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 183]
210. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 184]
211. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 184]
212. eq(X0,X0) [cnf transformation 56]
218. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 186]
220. lte(X1,X0) | lte(X0,X1) [cnf transformation 62]
222. ~lt(X0,X1) | ~eq(X0,X1) [cnf transformation 188]
223. ~lte(X0,X1) | eq(X0,X1) | lt(X0,X1) [cnf transformation 188]
224. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 189]
225. ~lt(X1,X0) | gt(X0,X1) [cnf transformation 189]
254. ~where_6(X1,X2) | where_6(X0,X2) | ~eq(X1,X0) [cnf transformation 162]
258. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 192]
259. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 192]
261. 0 <=> main_query_1(sK6,sK7) [avatar definition]
262. main_query_1(sK6,sK7) <- {0} [avatar component clause 261]
264. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
265. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 264]
267. 2 <=> main_query_2(sK6,sK7) [avatar definition]
268. main_query_2(sK6,sK7) <- {2} [avatar component clause 267]
270. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
271. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 270]
272. ~1 | ~3 [avatar split clause 259,270,264]
273. 0 | 2 [avatar split clause 258,267,261]
276. select_4(sK6,sK7) <- {2} [resolution 210,268]
277. where_6(sK6,sK7) <- {2} [resolution 276,208]
283. lt(X1,X0) | lt(X0,X1) | where_6(X0,X1) [resolution 206,224]
293. eq(X1,X2) | lt(X1,X2) | lte(X2,X1) [resolution 223,220]
295. lte(X2,X1) | lt(X1,X2) [subsumption resolution 293,218]
296. lt(X1,X0) | lt(X0,X1) | eq(X1,X0) [resolution 295,223]
318. ~eq(X3,X2) | where_6(X3,X2) | lt(X2,X3) [resolution 283,222]
321. lt(X0,X0) | where_6(X0,X0) [factoring 283]
323. where_6(X1,X1) | ~eq(X1,X1) [resolution 321,222]
326. where_6(X1,X1) [subsumption resolution 323,212]
328. ~eq(X3,X2) | where_6(X2,X3) [resolution 326,254]
334. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 197,193]
340. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 207,202]
361. 5 <=> ~lt(sK7,sK6) [avatar definition]
362. ~lt(sK7,sK6) <- {5} [avatar component clause 361]
364. 4 <=> lt(sK7,sK6) [avatar definition]
365. lt(sK7,sK6) <- {4} [avatar component clause 364]
367. 7 <=> ~where_6(sK7,sK6) [avatar definition]
370. 6 <=> where_6(sK7,sK6) [avatar definition]
374. gt(sK6,sK7) <- {4} [resolution 365,225]
377. ~where_6(sK7,sK6) <- {4} [resolution 365,204]
378. ~7 | ~4 [avatar split clause 377,364,367]
381. ~where_6(sK6,sK7) <- {4} [resolution 374,205]
382. $false <- {2, 4} [subsumption resolution 381,277]
383. ~2 | ~4 [avatar contradiction clause 382]
384. lt(sK6,sK7) | eq(sK6,sK7) <- {5} [resolution 362,296]
389. 9 <=> ~eq(sK6,sK7) [avatar definition]
390. ~eq(sK6,sK7) <- {9} [avatar component clause 389]
392. 8 <=> eq(sK6,sK7) [avatar definition]
393. eq(sK6,sK7) <- {8} [avatar component clause 392]
398. 10 <=> lt(sK6,sK7) [avatar definition]
399. lt(sK6,sK7) <- {10} [avatar component clause 398]
400. 8 | 10 | 5 [avatar split clause 384,361,398,392]
414. ~where_6(sK6,sK7) <- {10} [resolution 399,204]
416. $false <- {2, 10} [subsumption resolution 414,277]
417. ~2 | ~10 [avatar contradiction clause 416]
421. where_3(sK6,sK7) <- {8} [resolution 393,196]
430. where_6(sK7,sK6) <- {8} [resolution 393,328]
480. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 334,203]
488. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 340,194]
597. where_6(sK6,sK7) | lt(sK7,sK6) <- {8} [resolution 318,393]
647. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 480,209]
664. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 488,199]
960. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 664,276]
964. select_1(sK6,sK7) <- {2, 8} [subsumption resolution 960,421]
978. main_query_1(sK6,sK7) <- {2, 8} [resolution 964,201]
980. $false <- {1, 2, 8} [subsumption resolution 978,265]
981. 1 | ~2 | ~8 [avatar contradiction clause 980]
995. where_6(sK6,sK7) <- {5, 8} [subsumption resolution 597,362]
998. select_1(sK6,sK7) <- {0} [resolution 262,200]
1003. where_3(sK6,sK7) <- {0} [resolution 998,198]
1013. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {5, 8} [resolution 995,647]
1016. select_4(sK6,sK7) <- {0, 5, 8} [subsumption resolution 1013,998]
1030. main_query_2(sK6,sK7) <- {0, 5, 8} [resolution 1016,211]
1036. $false <- {0, 3, 5, 8} [subsumption resolution 1030,271]
1037. ~0 | 3 | 5 | ~8 [avatar contradiction clause 1036]
1042. 6 | ~8 [avatar split clause 430,392,370]
1044. eq(sK6,sK7) <- {0} [resolution 1003,195]
1045. $false <- {0, 9} [subsumption resolution 1044,390]
1046. ~0 | 9 [avatar contradiction clause 1045]
1047. $false [avatar sat refutation 272,273,378,383,400,417,981,1037,1042,1046]
% SZS output end Proof for query10_query13_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.024 s
% ------------------------------
% ------------------------------
