% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_05_query_10_tptp4X
% SZS output start Proof for query_05_query_10_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X0,X5,X6,X7,X8] : (from_8(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
6. ! [X0] : (select_7(X0) <=> ? [X5,X6,X7,X8] : ($true & from_8(X0,X5,X6,X7,X8))) [input]
7. ! [X0] : (subquery_from_6(X0) <=> select_7(X0)) [input]
8. ! [X9] : (from_5(X9) <=> subquery_from_6(X9)) [input]
9. ! [X9] : (select_4(X9) <=> ($true & from_5(X9))) [input]
10. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [input]
11. ! [X10] : eq(X10,X10) [input]
42. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
43. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 42]
44. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
45. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
46. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 45]
47. ! [X0] : (from_5(X0) <=> subquery_from_6(X0)) [rectify 8]
48. ! [X0] : (select_4(X0) <=> ($true & from_5(X0))) [rectify 9]
49. ! [X0] : (select_4(X0) <=> from_5(X0)) [true and false elimination 48]
50. ! [X0] : eq(X0,X0) [rectify 11]
126. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 43]
127. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
128. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
129. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
130. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 129]
131. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
132. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 130,131]
133. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
134. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 44]
135. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 46]
136. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 135]
137. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
138. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0))) [skolemisation 136,137]
139. ! [X0] : ((subquery_from_6(X0) | ~select_7(X0)) & (select_7(X0) | ~subquery_from_6(X0))) [nnf transformation 7]
140. ! [X0] : ((from_5(X0) | ~subquery_from_6(X0)) & (subquery_from_6(X0) | ~from_5(X0))) [nnf transformation 47]
141. ! [X0] : ((select_4(X0) | ~from_5(X0)) & (from_5(X0) | ~select_4(X0))) [nnf transformation 49]
142. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 10]
145. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 126]
146. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
147. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 145,146]
148. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 127]
149. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 127]
151. ~eq(X0,X0) | where_3(X0) [cnf transformation 128]
152. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 132]
154. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 132]
155. ~main_query_1(X0) | select_1(X0) [cnf transformation 133]
156. ~select_1(X0) | main_query_1(X0) [cnf transformation 133]
157. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 134]
158. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 134]
159. from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [cnf transformation 138]
160. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 138]
161. ~subquery_from_6(X0) | select_7(X0) [cnf transformation 139]
162. ~select_7(X0) | subquery_from_6(X0) [cnf transformation 139]
163. ~from_5(X0) | subquery_from_6(X0) [cnf transformation 140]
164. ~subquery_from_6(X0) | from_5(X0) [cnf transformation 140]
165. ~select_4(X0) | from_5(X0) [cnf transformation 141]
166. ~from_5(X0) | select_4(X0) [cnf transformation 141]
167. ~main_query_2(X0) | select_4(X0) [cnf transformation 142]
168. ~select_4(X0) | main_query_2(X0) [cnf transformation 142]
169. eq(X0,X0) [cnf transformation 50]
198. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 147]
199. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 147]
201. 0 <=> main_query_1(sK8) [avatar definition]
202. main_query_1(sK8) <- {0} [avatar component clause 201]
204. 1 <=> ~main_query_1(sK8) [avatar definition]
205. ~main_query_1(sK8) <- {1} [avatar component clause 204]
207. 2 <=> main_query_2(sK8) [avatar definition]
208. main_query_2(sK8) <- {2} [avatar component clause 207]
210. 3 <=> ~main_query_2(sK8) [avatar definition]
211. ~main_query_2(sK8) <- {3} [avatar component clause 210]
212. ~1 | ~3 [avatar split clause 199,210,204]
213. 0 | 2 [avatar split clause 198,207,201]
214. select_4(sK8) <- {2} [resolution 167,208]
215. from_5(sK8) <- {2} [resolution 214,165]
218. subquery_from_6(sK8) <- {2} [resolution 215,163]
220. select_7(sK8) <- {2} [resolution 218,161]
222. where_3(X1) [resolution 151,169]
257. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 152,148]
259. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [resolution 159,157]
262. from_8(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 257,158]
266. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_7(X1) [resolution 259,149]
272. ~select_1(X7) | select_7(X7) [resolution 262,160]
278. ~select_7(X1) | ~where_3(X1) | select_1(X1) [resolution 266,154]
279. ~select_7(X1) | select_1(X1) [subsumption resolution 278,222]
280. select_1(sK8) <- {2} [resolution 279,220]
282. main_query_1(sK8) <- {2} [resolution 280,156]
284. $false <- {1, 2} [subsumption resolution 282,205]
285. 1 | ~2 [avatar contradiction clause 284]
286. select_1(sK8) <- {0} [resolution 202,155]
287. select_7(sK8) <- {0} [resolution 286,272]
291. subquery_from_6(sK8) <- {0} [resolution 287,162]
292. from_5(sK8) <- {0} [resolution 291,164]
296. select_4(sK8) <- {0} [resolution 292,166]
300. main_query_2(sK8) <- {0} [resolution 296,168]
302. $false <- {0, 3} [subsumption resolution 300,211]
303. ~0 | 3 [avatar contradiction clause 302]
304. $false [avatar sat refutation 212,213,285,303]
% SZS output end Proof for query_05_query_10_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.027 s
% ------------------------------
% ------------------------------
