% Refutation found. Thanks to Tanya!
% SZS status Theorem for query04_query07_tptp4X
% SZS output start Proof for query04_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
7. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [input]
9. ! [X5] : eq(X5,X5) [input]
16. ! [X5,X6] : (lt(X5,X6) <=> (~eq(X5,X6) & lte(X5,X6))) [input]
17. ! [X5,X6] : (gt(X5,X6) <=> lt(X6,X5)) [input]
40. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
41. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 40]
42. ! [X0] : eq(X0,X0) [rectify 9]
49. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 16]
50. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 17]
118. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 41]
119. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
120. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
121. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
122. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 121]
123. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
124. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 122,123]
125. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
126. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
127. ! [X0] : ((where_6(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [nnf transformation 6]
128. ! [X0] : ((where_6(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [flattening 127]
129. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 7]
130. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 129]
131. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
132. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 130,131]
133. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 8]
136. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 49]
137. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 136]
138. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 50]
139. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 118]
140. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
141. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 139,140]
142. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 119]
143. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 119]
145. ~eq(X0,X0) | where_3(X0) [cnf transformation 120]
146. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 124]
148. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 124]
149. ~main_query_1(X0) | select_1(X0) [cnf transformation 125]
150. ~select_1(X0) | main_query_1(X0) [cnf transformation 125]
151. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 126]
152. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 126]
155. gt(X0,X0) | where_6(X0) | lt(X0,X0) [cnf transformation 128]
156. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 132]
158. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 132]
159. ~main_query_2(X0) | select_4(X0) [cnf transformation 133]
160. ~select_4(X0) | main_query_2(X0) [cnf transformation 133]
161. eq(X0,X0) [cnf transformation 42]
171. ~lt(X0,X1) | ~eq(X0,X1) [cnf transformation 137]
173. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 138]
195. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 141]
196. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 141]
198. 0 <=> main_query_1(sK8) [avatar definition]
199. main_query_1(sK8) <- {0} [avatar component clause 198]
201. 1 <=> ~main_query_1(sK8) [avatar definition]
202. ~main_query_1(sK8) <- {1} [avatar component clause 201]
204. 2 <=> main_query_2(sK8) [avatar definition]
205. main_query_2(sK8) <- {2} [avatar component clause 204]
207. 3 <=> ~main_query_2(sK8) [avatar definition]
208. ~main_query_2(sK8) <- {3} [avatar component clause 207]
209. ~1 | ~3 [avatar split clause 196,207,201]
210. 0 | 2 [avatar split clause 195,204,198]
211. select_4(sK8) <- {2} [resolution 159,205]
215. where_3(X1) [resolution 145,161]
224. where_6(X1) | lt(X1,X1) | lt(X1,X1) [resolution 155,173]
225. lt(X1,X1) | where_6(X1) [duplicate literal removal 224]
228. where_6(X2) | ~eq(X2,X2) [resolution 225,171]
230. where_6(X2) [subsumption resolution 228,161]
265. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 146,142]
305. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 156,151]
431. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 265,152]
436. from_2(X3,sK4(X3),sK5(X3),sK6(X3),sK7(X3)) | ~select_4(X3) [resolution 305,143]
440. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 431,158]
446. ~select_1(X1) | select_4(X1) [subsumption resolution 440,230]
452. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 436,148]
458. ~select_4(X1) | select_1(X1) [subsumption resolution 452,215]
461. select_1(sK8) <- {2} [resolution 458,211]
471. main_query_1(sK8) <- {2} [resolution 461,150]
473. $false <- {1, 2} [subsumption resolution 471,202]
474. 1 | ~2 [avatar contradiction clause 473]
475. select_1(sK8) <- {0} [resolution 199,149]
476. select_4(sK8) <- {0} [resolution 475,446]
480. main_query_2(sK8) <- {0} [resolution 476,160]
482. $false <- {0, 3} [subsumption resolution 480,208]
483. ~0 | 3 [avatar contradiction clause 482]
484. $false [avatar sat refutation 209,210,474,483]
% SZS output end Proof for query04_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.019 s
% ------------------------------
% ------------------------------
