% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query08_tptp4X
% SZS output start Proof for query00_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
6. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
11. ! [X5] : lte(X5,X5) [input]
17. ! [X5,X6] : (gte(X5,X6) <=> lte(X6,X5)) [input]
41. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
42. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
47. ! [X0] : lte(X0,X0) [rectify 11]
53. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 17]
128. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 42]
129. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
130. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
131. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 130]
132. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
133. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 131,132]
134. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
135. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
136. ! [X0] : ((where_5(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [nnf transformation 5]
137. ! [X0] : ((where_5(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [flattening 136]
138. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 6]
139. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 138]
140. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
141. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 139,140]
142. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 7]
145. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 53]
146. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 128]
147. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
148. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 146,147]
149. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 129]
150. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 129]
151. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 133]
152. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 133]
153. ~main_query_1(X0) | select_1(X0) [cnf transformation 134]
154. ~select_1(X0) | main_query_1(X0) [cnf transformation 134]
155. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 135]
156. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 135]
159. ~gte(X0,X0) | where_5(X0) | ~lte(X0,X0) [cnf transformation 137]
160. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 141]
162. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 141]
163. ~main_query_2(X0) | select_3(X0) [cnf transformation 142]
164. ~select_3(X0) | main_query_2(X0) [cnf transformation 142]
168. lte(X0,X0) [cnf transformation 47]
174. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 145]
198. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 148]
199. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 148]
224. 0 <=> main_query_1(sK8) [avatar definition]
225. main_query_1(sK8) <- {0} [avatar component clause 224]
227. 1 <=> ~main_query_1(sK8) [avatar definition]
228. ~main_query_1(sK8) <- {1} [avatar component clause 227]
230. 2 <=> main_query_2(sK8) [avatar definition]
231. main_query_2(sK8) <- {2} [avatar component clause 230]
233. 3 <=> ~main_query_2(sK8) [avatar definition]
234. ~main_query_2(sK8) <- {3} [avatar component clause 233]
235. ~1 | ~3 [avatar split clause 199,233,227]
236. 0 | 2 [avatar split clause 198,230,224]
237. select_3(sK8) <- {2} [resolution 163,231]
245. gte(X1,X1) [resolution 174,168]
247. where_5(X0) | ~lte(X0,X0) [resolution 159,245]
249. where_5(X0) [subsumption resolution 247,168]
255. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 151,149]
257. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 160,155]
259. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 255,156]
262. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 257,150]
264. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 259,162]
265. ~select_1(X1) | select_3(X1) [subsumption resolution 264,249]
267. ~select_3(X1) | select_1(X1) [resolution 262,152]
268. select_1(sK8) <- {2} [resolution 267,237]
270. main_query_1(sK8) <- {2} [resolution 268,154]
271. $false <- {1, 2} [subsumption resolution 270,228]
272. 1 | ~2 [avatar contradiction clause 271]
273. select_1(sK8) <- {0} [resolution 225,153]
274. select_3(sK8) <- {0} [resolution 273,265]
277. main_query_2(sK8) <- {0} [resolution 274,164]
279. $false <- {0, 3} [subsumption resolution 277,234]
280. ~0 | 3 [avatar contradiction clause 279]
281. $false [avatar sat refutation 235,236,272,280]
% SZS output end Proof for query00_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.014 s
% ------------------------------
% ------------------------------
