% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_12_query_28_tptp4X
% SZS output start Proof for query_12_query_28_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X12,X9,X10,X11] : (from_9(X7,X12,X9,X10,X11) <=> tab1(X7,X12,X9,X10,X11)) [input]
7. ! [X7,X8] : (where_10(X7,X8) <=> eq(X7,X8)) [input]
8. ! [X7,X8] : (select_8(X7,X8) <=> ? [X12,X9,X10,X11] : (where_10(X7,X8) & from_9(X7,X12,X9,X10,X11))) [input]
9. ! [X5,X8] : (where_6_subquery_7(X5,X8) <=> select_8(X5,X8)) [input]
10. ! [X7,X8] : (where_6(X7,X8) <=> where_6_subquery_7(X7,X8)) [input]
11. ! [X7,X8] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X7,X8) & from_5(X7,X8,X9,X10,X11))) [input]
12. ! [X5,X8] : (main_query_2(X5,X8) <=> select_4(X5,X8)) [input]
14. ! [X13,X14] : (eq(X13,X14) => eq(X14,X13)) [input]
40. ! [X13,X5,X6] : ((select_1(X5,X6) & eq(X6,X13)) => select_1(X5,X13)) [input]
66. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
67. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 66]
68. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
69. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
70. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 6]
71. ! [X0,X1] : (where_10(X0,X1) <=> eq(X0,X1)) [rectify 7]
72. ! [X0,X1] : (select_8(X0,X1) <=> ? [X2,X3,X4,X5] : (where_10(X0,X1) & from_9(X0,X2,X3,X4,X5))) [rectify 8]
73. ! [X0,X1] : (where_6_subquery_7(X0,X1) <=> select_8(X0,X1)) [rectify 9]
74. ! [X0,X1] : (where_6(X0,X1) <=> where_6_subquery_7(X0,X1)) [rectify 10]
75. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 11]
76. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 12]
78. ! [X0,X1] : (eq(X0,X1) => eq(X1,X0)) [rectify 14]
104. ! [X0,X1,X2] : ((select_1(X1,X2) & eq(X2,X0)) => select_1(X1,X0)) [rectify 40]
130. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 67]
131. ! [X0,X1] : (eq(X1,X0) | ~eq(X0,X1)) [ennf transformation 78]
170. ! [X0,X1,X2] : (select_1(X1,X0) | (~select_1(X1,X2) | ~eq(X2,X0))) [ennf transformation 104]
171. ! [X0,X1,X2] : (select_1(X1,X0) | ~select_1(X1,X2) | ~eq(X2,X0)) [flattening 170]
222. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 130]
223. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
224. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
225. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
226. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 225]
227. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
228. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 226,227]
229. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 68]
230. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 69]
231. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 70]
232. ! [X0,X1] : ((where_10(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_10(X0,X1))) [nnf transformation 71]
233. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5] : (~where_10(X0,X1) | ~from_9(X0,X2,X3,X4,X5))) & (? [X2,X3,X4,X5] : (where_10(X0,X1) & from_9(X0,X2,X3,X4,X5)) | ~select_8(X0,X1))) [nnf transformation 72]
234. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5] : (~where_10(X0,X1) | ~from_9(X0,X2,X3,X4,X5))) & (? [X6,X7,X8,X9] : (where_10(X0,X1) & from_9(X0,X6,X7,X8,X9)) | ~select_8(X0,X1))) [rectify 233]
235. ! [X1,X0] : (? [X6,X7,X8,X9] : (where_10(X0,X1) & from_9(X0,X6,X7,X8,X9)) => (where_10(X0,X1) & from_9(X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1),sK6(X0,X1)))) [choice axiom]
236. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5] : (~where_10(X0,X1) | ~from_9(X0,X2,X3,X4,X5))) & ((where_10(X0,X1) & from_9(X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1),sK6(X0,X1))) | ~select_8(X0,X1))) [skolemisation 234,235]
237. ! [X0,X1] : ((where_6_subquery_7(X0,X1) | ~select_8(X0,X1)) & (select_8(X0,X1) | ~where_6_subquery_7(X0,X1))) [nnf transformation 73]
238. ! [X0,X1] : ((where_6(X0,X1) | ~where_6_subquery_7(X0,X1)) & (where_6_subquery_7(X0,X1) | ~where_6(X0,X1))) [nnf transformation 74]
239. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 75]
240. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 239]
241. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
242. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 240,241]
243. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 76]
246. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 222]
247. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
248. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 246,247]
249. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 223]
250. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 223]
251. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 224]
252. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 224]
253. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 228]
254. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 228]
255. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 228]
256. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 229]
257. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 229]
258. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 230]
259. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 230]
261. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 231]
262. ~where_10(X0,X1) | eq(X0,X1) [cnf transformation 232]
263. ~eq(X0,X1) | where_10(X0,X1) [cnf transformation 232]
265. ~select_8(X0,X1) | where_10(X0,X1) [cnf transformation 236]
266. ~from_9(X0,X2,X3,X4,X5) | ~where_10(X0,X1) | select_8(X0,X1) [cnf transformation 236]
267. ~where_6_subquery_7(X0,X1) | select_8(X0,X1) [cnf transformation 237]
268. ~select_8(X0,X1) | where_6_subquery_7(X0,X1) [cnf transformation 237]
269. ~where_6(X0,X1) | where_6_subquery_7(X0,X1) [cnf transformation 238]
270. ~where_6_subquery_7(X0,X1) | where_6(X0,X1) [cnf transformation 238]
271. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 242]
272. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 242]
273. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 242]
274. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 243]
275. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 243]
277. ~eq(X0,X1) | eq(X1,X0) [cnf transformation 131]
301. ~select_1(X1,X2) | select_1(X1,X0) | ~eq(X2,X0) [cnf transformation 171]
327. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 248]
328. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 248]
330. 0 <=> main_query_1(sK10,sK11) [avatar definition]
331. main_query_1(sK10,sK11) <- {0} [avatar component clause 330]
333. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
334. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 333]
336. 2 <=> main_query_2(sK10,sK11) [avatar definition]
337. main_query_2(sK10,sK11) <- {2} [avatar component clause 336]
339. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
340. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 339]
341. ~1 | ~3 [avatar split clause 328,339,333]
342. 0 | 2 [avatar split clause 327,336,330]
347. select_4(sK10,sK11) <- {2} [resolution 274,337]
348. where_6(sK10,sK11) <- {2} [resolution 347,272]
350. where_6_subquery_7(sK10,sK11) <- {2} [resolution 348,269]
353. select_8(sK10,sK11) <- {2} [resolution 350,267]
356. where_10(sK10,sK11) <- {2} [resolution 353,265]
358. eq(sK10,sK11) <- {2} [resolution 356,262]
365. where_3(sK10,sK11) <- {2} [resolution 358,252]
467. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 253,249]
471. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 271,258]
558. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 467,259]
559. from_9(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 467,261]
567. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 471,250]
730. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 558,273]
739. ~where_10(X2,X4) | ~select_1(X2,X3) | select_8(X2,X4) [resolution 559,266]
750. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 567,255]
773. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 750,347]
777. select_1(sK10,sK11) <- {2} [subsumption resolution 773,365]
791. main_query_1(sK10,sK11) <- {2} [resolution 777,257]
793. $false <- {1, 2} [subsumption resolution 791,334]
794. 1 | ~2 [avatar contradiction clause 793]
797. select_1(sK10,sK11) <- {0} [resolution 331,256]
799. ~eq(sK11,X0) | select_1(sK10,X0) <- {0} [resolution 797,301]
802. where_3(sK10,sK11) <- {0} [resolution 797,254]
803. eq(sK10,sK11) <- {0} [resolution 802,251]
806. where_10(sK10,sK11) <- {0} [resolution 803,263]
807. eq(sK11,sK10) <- {0} [resolution 803,277]
816. ~select_1(sK10,X0) | select_8(sK10,sK11) <- {0} [resolution 806,739]
824. 4 <=> select_8(sK10,sK11) [avatar definition]
825. select_8(sK10,sK11) <- {4} [avatar component clause 824]
827. 6 <=> ! [X0] : ~select_1(sK10,X0) [avatar definition]
828. ~select_1(sK10,X0) <- {6} [avatar component clause 827]
829. 4 | 6 | ~0 [avatar split clause 816,330,827,824]
831. select_1(sK10,sK10) <- {0} [resolution 807,799]
842. $false <- {0, 6} [subsumption resolution 831,828]
843. ~0 | ~6 [avatar contradiction clause 842]
931. where_6_subquery_7(sK10,sK11) <- {4} [resolution 825,268]
1026. where_6(sK10,sK11) <- {4} [resolution 931,270]
1066. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {4} [resolution 1026,730]
1070. select_4(sK10,sK11) <- {0, 4} [subsumption resolution 1066,797]
1125. main_query_2(sK10,sK11) <- {0, 4} [resolution 1070,275]
1127. $false <- {0, 3, 4} [subsumption resolution 1125,340]
1128. ~0 | 3 | ~4 [avatar contradiction clause 1127]
1129. $false [avatar sat refutation 341,342,794,829,843,1128]
% SZS output end Proof for query_12_query_28_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.019 s
% ------------------------------
% ------------------------------
