% Refutation found. Thanks to Tanya!
% SZS status Theorem for query11_query14_tptp4X
% SZS output start Proof for query11_query14_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8,X9,X10,X11] : (from_9(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
7. ! [X7] : (select_8(X7) <=> ? [X8,X9,X10,X11] : ($true & from_9(X7,X8,X9,X10,X11))) [input]
8. ! [X5] : (where_6_subquery_7(X5) <=> select_8(X5)) [input]
9. ! [X7,X8] : (where_6(X7,X8) <=> (where_6_subquery_7(X7) & eq(X7,X8))) [input]
10. ! [X7,X8] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X7,X8) & from_5(X7,X8,X9,X10,X11))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
63. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
64. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 63]
65. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
66. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
67. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 6]
68. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [rectify 7]
69. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 68]
70. ! [X0] : (where_6_subquery_7(X0) <=> select_8(X0)) [rectify 8]
71. ! [X0,X1] : (where_6(X0,X1) <=> (where_6_subquery_7(X0) & eq(X0,X1))) [rectify 9]
72. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 10]
73. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
125. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 64]
211. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 125]
212. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
213. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
214. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
215. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 214]
216. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
217. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 215,216]
218. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 65]
219. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 66]
220. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 67]
221. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 69]
222. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 221]
223. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
224. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 222,223]
225. ! [X0] : ((where_6_subquery_7(X0) | ~select_8(X0)) & (select_8(X0) | ~where_6_subquery_7(X0))) [nnf transformation 70]
226. ! [X0,X1] : ((where_6(X0,X1) | (~where_6_subquery_7(X0) | ~eq(X0,X1))) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 71]
227. ! [X0,X1] : ((where_6(X0,X1) | ~where_6_subquery_7(X0) | ~eq(X0,X1)) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 226]
228. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 72]
229. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 228]
230. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
231. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 229,230]
232. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 73]
235. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 211]
236. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
237. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 235,236]
238. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 212]
239. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 212]
240. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 213]
241. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 213]
242. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 217]
243. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 217]
244. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 217]
245. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 218]
246. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 218]
247. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 219]
248. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 219]
250. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 220]
252. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 224]
254. ~select_8(X0) | where_6_subquery_7(X0) [cnf transformation 225]
255. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 227]
257. ~eq(X0,X1) | ~where_6_subquery_7(X0) | where_6(X0,X1) [cnf transformation 227]
258. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 231]
259. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 231]
260. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 231]
261. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 232]
262. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 232]
312. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 237]
313. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 237]
315. 0 <=> main_query_1(sK10,sK11) [avatar definition]
316. main_query_1(sK10,sK11) <- {0} [avatar component clause 315]
318. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
319. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 318]
321. 2 <=> main_query_2(sK10,sK11) [avatar definition]
322. main_query_2(sK10,sK11) <- {2} [avatar component clause 321]
324. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
325. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 324]
326. ~1 | ~3 [avatar split clause 313,324,318]
327. 0 | 2 [avatar split clause 312,321,315]
399. select_1(sK10,sK11) <- {0} [resolution 316,245]
403. where_3(sK10,sK11) <- {0} [resolution 399,243]
405. eq(sK10,sK11) <- {0} [resolution 403,240]
408. ~where_6_subquery_7(sK10) | where_6(sK10,sK11) <- {0} [resolution 405,257]
417. 8 <=> where_6(sK10,sK11) [avatar definition]
418. where_6(sK10,sK11) <- {8} [avatar component clause 417]
420. 10 <=> where_6_subquery_7(sK10) [avatar definition]
423. 11 <=> ~where_6_subquery_7(sK10) [avatar definition]
425. 8 | ~11 | ~0 [avatar split clause 408,315,423,417]
475. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 242,238]
479. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 258,247]
632. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 475,248]
633. from_9(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 475,250]
652. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 479,239]
688. 14 <=> select_8(sK10) [avatar definition]
689. select_8(sK10) <- {14} [avatar component clause 688]
1272. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 632,260]
1280. ~select_1(X2,X3) | select_8(X2) [resolution 633,252]
1286. select_8(sK10) <- {0} [resolution 1280,399]
1301. 14 | ~0 [avatar split clause 1286,315,688]
1303. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 652,244]
1321. where_6_subquery_7(sK10) <- {14} [resolution 689,254]
1324. 10 | ~14 [avatar split clause 1321,688,420]
1351. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {8} [resolution 418,1272]
1356. select_4(sK10,sK11) <- {0, 8} [subsumption resolution 1351,399]
1381. main_query_2(sK10,sK11) <- {0, 8} [resolution 1356,262]
1383. $false <- {0, 3, 8} [subsumption resolution 1381,325]
1384. ~0 | 3 | ~8 [avatar contradiction clause 1383]
1400. select_4(sK10,sK11) <- {2} [resolution 322,261]
1404. where_6(sK10,sK11) <- {2} [resolution 1400,259]
1405. 8 | ~2 [avatar split clause 1404,321,417]
1409. eq(sK10,sK11) <- {8} [resolution 418,255]
1412. where_3(sK10,sK11) <- {8} [resolution 1409,241]
1830. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 1303,1400]
1834. select_1(sK10,sK11) <- {2, 8} [subsumption resolution 1830,1412]
1843. main_query_1(sK10,sK11) <- {2, 8} [resolution 1834,246]
1849. $false <- {1, 2, 8} [subsumption resolution 1843,319]
1850. 1 | ~2 | ~8 [avatar contradiction clause 1849]
1851. $false [avatar sat refutation 326,327,425,1301,1324,1384,1405,1850]
% SZS output end Proof for query11_query14_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5628
% Time elapsed: 0.232 s
% ------------------------------
% ------------------------------
