% Refutation found. Thanks to Tanya!
% SZS status Theorem for query20_query21_tptp4X
% SZS output start Proof for query20_query21_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
3. ! [X5,X6] : (where_9(X5,X6) <=> eq(X5,X6)) [input]
4. ! [X5,X6] : (select_7(X5,X6) <=> ? [X0,X1,X2,X3,X4] : (where_9(X5,X6) & from_8(X0,X1,X2,X3,X4))) [input]
5. ! [X5,X6] : (query_6(X5,X6) <=> select_7(X5,X6)) [input]
6. ! [X7,X8] : (from_10(X7,X8) <=> query_6(X7,X8)) [input]
7. ! [X7] : (select_5(X7) <=> ? [X8] : ($true & from_10(X7,X8))) [input]
8. ! [X5] : (query_4(X5) <=> select_5(X5)) [input]
9. ! [X5] : (where_3(X5) <=> query_4(X5)) [input]
10. ! [X5] : (select_1(X5) <=> ? [X0,X1,X2,X3,X4] : (where_3(X5) & from_2(X0,X1,X2,X3,X4))) [input]
11. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
12. ! [X0,X1,X2,X3,X4] : (from_14(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
13. ! [X5,X6] : (where_15(X5,X6) <=> eq(X5,X6)) [input]
14. ! [X5,X6] : (select_13(X5,X6) <=> ? [X0,X1,X2,X3,X4] : (where_15(X5,X6) & from_14(X0,X1,X2,X3,X4))) [input]
15. ! [X5,X6] : (query_12(X5,X6) <=> select_13(X5,X6)) [input]
16. ! [X7,X8] : (from_16(X7,X8) <=> query_12(X7,X8)) [input]
17. ! [X7] : (select_11(X7) <=> ? [X8] : ($true & from_16(X7,X8))) [input]
18. ! [X5] : (main_query_2(X5) <=> select_11(X5)) [input]
19. ! [X9] : eq(X9,X9) [input]
73. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
74. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 73]
75. ! [X0,X1] : (where_9(X0,X1) <=> eq(X0,X1)) [rectify 3]
76. ! [X0,X1] : (select_7(X0,X1) <=> ? [X2,X3,X4,X5,X6] : (where_9(X0,X1) & from_8(X2,X3,X4,X5,X6))) [rectify 4]
77. ! [X0,X1] : (query_6(X0,X1) <=> select_7(X0,X1)) [rectify 5]
78. ! [X0,X1] : (from_10(X0,X1) <=> query_6(X0,X1)) [rectify 6]
79. ! [X0] : (select_5(X0) <=> ? [X1] : ($true & from_10(X0,X1))) [rectify 7]
80. ! [X0] : (select_5(X0) <=> ? [X1] : from_10(X0,X1)) [true and false elimination 79]
81. ! [X0] : (query_4(X0) <=> select_5(X0)) [rectify 8]
82. ! [X0] : (where_3(X0) <=> query_4(X0)) [rectify 9]
83. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : (where_3(X0) & from_2(X1,X2,X3,X4,X5))) [rectify 10]
84. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 11]
85. ! [X0,X1] : (where_15(X0,X1) <=> eq(X0,X1)) [rectify 13]
86. ! [X0,X1] : (select_13(X0,X1) <=> ? [X2,X3,X4,X5,X6] : (where_15(X0,X1) & from_14(X2,X3,X4,X5,X6))) [rectify 14]
87. ! [X0,X1] : (query_12(X0,X1) <=> select_13(X0,X1)) [rectify 15]
88. ! [X0,X1] : (from_16(X0,X1) <=> query_12(X0,X1)) [rectify 16]
89. ! [X0] : (select_11(X0) <=> ? [X1] : ($true & from_16(X0,X1))) [rectify 17]
90. ! [X0] : (select_11(X0) <=> ? [X1] : from_16(X0,X1)) [true and false elimination 89]
91. ! [X0] : (main_query_2(X0) <=> select_11(X0)) [rectify 18]
92. ! [X0] : eq(X0,X0) [rectify 19]
146. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 74]
240. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 146]
241. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
242. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 2]
243. ! [X0,X1] : ((where_9(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_9(X0,X1))) [nnf transformation 75]
244. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_9(X0,X1) | ~from_8(X2,X3,X4,X5,X6))) & (? [X2,X3,X4,X5,X6] : (where_9(X0,X1) & from_8(X2,X3,X4,X5,X6)) | ~select_7(X0,X1))) [nnf transformation 76]
245. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_9(X0,X1) | ~from_8(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_9(X0,X1) & from_8(X7,X8,X9,X10,X11)) | ~select_7(X0,X1))) [rectify 244]
246. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_9(X0,X1) & from_8(X7,X8,X9,X10,X11)) => (where_9(X0,X1) & from_8(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)))) [choice axiom]
247. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_9(X0,X1) | ~from_8(X2,X3,X4,X5,X6))) & ((where_9(X0,X1) & from_8(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1))) | ~select_7(X0,X1))) [skolemisation 245,246]
248. ! [X0,X1] : ((query_6(X0,X1) | ~select_7(X0,X1)) & (select_7(X0,X1) | ~query_6(X0,X1))) [nnf transformation 77]
249. ! [X0,X1] : ((from_10(X0,X1) | ~query_6(X0,X1)) & (query_6(X0,X1) | ~from_10(X0,X1))) [nnf transformation 78]
250. ! [X0] : ((select_5(X0) | ! [X1] : ~from_10(X0,X1)) & (? [X1] : from_10(X0,X1) | ~select_5(X0))) [nnf transformation 80]
251. ! [X0] : ((select_5(X0) | ! [X1] : ~from_10(X0,X1)) & (? [X2] : from_10(X0,X2) | ~select_5(X0))) [rectify 250]
252. ! [X0] : (? [X2] : from_10(X0,X2) => from_10(X0,sK5(X0))) [choice axiom]
253. ! [X0] : ((select_5(X0) | ! [X1] : ~from_10(X0,X1)) & (from_10(X0,sK5(X0)) | ~select_5(X0))) [skolemisation 251,252]
254. ! [X0] : ((query_4(X0) | ~select_5(X0)) & (select_5(X0) | ~query_4(X0))) [nnf transformation 81]
255. ! [X0] : ((where_3(X0) | ~query_4(X0)) & (query_4(X0) | ~where_3(X0))) [nnf transformation 82]
256. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : (~where_3(X0) | ~from_2(X1,X2,X3,X4,X5))) & (? [X1,X2,X3,X4,X5] : (where_3(X0) & from_2(X1,X2,X3,X4,X5)) | ~select_1(X0))) [nnf transformation 83]
257. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : (~where_3(X0) | ~from_2(X1,X2,X3,X4,X5))) & (? [X6,X7,X8,X9,X10] : (where_3(X0) & from_2(X6,X7,X8,X9,X10)) | ~select_1(X0))) [rectify 256]
258. ! [X0] : (? [X6,X7,X8,X9,X10] : (where_3(X0) & from_2(X6,X7,X8,X9,X10)) => (where_3(X0) & from_2(sK6(X0),sK7(X0),sK8(X0),sK9(X0),sK10(X0)))) [choice axiom]
259. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : (~where_3(X0) | ~from_2(X1,X2,X3,X4,X5))) & ((where_3(X0) & from_2(sK6(X0),sK7(X0),sK8(X0),sK9(X0),sK10(X0))) | ~select_1(X0))) [skolemisation 257,258]
260. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 84]
261. ! [X0,X1,X2,X3,X4] : ((from_14(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_14(X0,X1,X2,X3,X4))) [nnf transformation 12]
262. ! [X0,X1] : ((where_15(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_15(X0,X1))) [nnf transformation 85]
263. ! [X0,X1] : ((select_13(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_15(X0,X1) | ~from_14(X2,X3,X4,X5,X6))) & (? [X2,X3,X4,X5,X6] : (where_15(X0,X1) & from_14(X2,X3,X4,X5,X6)) | ~select_13(X0,X1))) [nnf transformation 86]
264. ! [X0,X1] : ((select_13(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_15(X0,X1) | ~from_14(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_15(X0,X1) & from_14(X7,X8,X9,X10,X11)) | ~select_13(X0,X1))) [rectify 263]
265. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_15(X0,X1) & from_14(X7,X8,X9,X10,X11)) => (where_15(X0,X1) & from_14(sK11(X0,X1),sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1)))) [choice axiom]
266. ! [X0,X1] : ((select_13(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_15(X0,X1) | ~from_14(X2,X3,X4,X5,X6))) & ((where_15(X0,X1) & from_14(sK11(X0,X1),sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1))) | ~select_13(X0,X1))) [skolemisation 264,265]
267. ! [X0,X1] : ((query_12(X0,X1) | ~select_13(X0,X1)) & (select_13(X0,X1) | ~query_12(X0,X1))) [nnf transformation 87]
268. ! [X0,X1] : ((from_16(X0,X1) | ~query_12(X0,X1)) & (query_12(X0,X1) | ~from_16(X0,X1))) [nnf transformation 88]
269. ! [X0] : ((select_11(X0) | ! [X1] : ~from_16(X0,X1)) & (? [X1] : from_16(X0,X1) | ~select_11(X0))) [nnf transformation 90]
270. ! [X0] : ((select_11(X0) | ! [X1] : ~from_16(X0,X1)) & (? [X2] : from_16(X0,X2) | ~select_11(X0))) [rectify 269]
271. ! [X0] : (? [X2] : from_16(X0,X2) => from_16(X0,sK16(X0))) [choice axiom]
272. ! [X0] : ((select_11(X0) | ! [X1] : ~from_16(X0,X1)) & (from_16(X0,sK16(X0)) | ~select_11(X0))) [skolemisation 270,271]
273. ! [X0] : ((main_query_2(X0) | ~select_11(X0)) & (select_11(X0) | ~main_query_2(X0))) [nnf transformation 91]
276. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 240]
277. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK17) | ~main_query_1(sK17)) & (main_query_2(sK17) | main_query_1(sK17))) [choice axiom]
278. (~main_query_2(sK17) | ~main_query_1(sK17)) & (main_query_2(sK17) | main_query_1(sK17)) [skolemisation 276,277]
279. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 241]
280. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 241]
281. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 242]
282. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 242]
284. ~eq(X0,X1) | where_9(X0,X1) [cnf transformation 243]
285. from_8(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_7(X0,X1) [cnf transformation 247]
287. select_7(X0,X1) | ~where_9(X0,X1) | ~from_8(X2,X3,X4,X5,X6) [cnf transformation 247]
289. ~select_7(X0,X1) | query_6(X0,X1) [cnf transformation 248]
291. ~query_6(X0,X1) | from_10(X0,X1) [cnf transformation 249]
293. ~from_10(X0,X1) | select_5(X0) [cnf transformation 253]
295. ~select_5(X0) | query_4(X0) [cnf transformation 254]
297. ~query_4(X0) | where_3(X0) [cnf transformation 255]
298. from_2(sK6(X0),sK7(X0),sK8(X0),sK9(X0),sK10(X0)) | ~select_1(X0) [cnf transformation 259]
300. select_1(X0) | ~where_3(X0) | ~from_2(X1,X2,X3,X4,X5) [cnf transformation 259]
301. ~main_query_1(X0) | select_1(X0) [cnf transformation 260]
302. ~select_1(X0) | main_query_1(X0) [cnf transformation 260]
303. ~from_14(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 261]
304. ~tab1(X0,X1,X2,X3,X4) | from_14(X0,X1,X2,X3,X4) [cnf transformation 261]
306. ~eq(X0,X1) | where_15(X0,X1) [cnf transformation 262]
307. from_14(sK11(X0,X1),sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1)) | ~select_13(X0,X1) [cnf transformation 266]
309. select_13(X0,X1) | ~where_15(X0,X1) | ~from_14(X2,X3,X4,X5,X6) [cnf transformation 266]
310. ~query_12(X0,X1) | select_13(X0,X1) [cnf transformation 267]
311. ~select_13(X0,X1) | query_12(X0,X1) [cnf transformation 267]
312. ~from_16(X0,X1) | query_12(X0,X1) [cnf transformation 268]
313. ~query_12(X0,X1) | from_16(X0,X1) [cnf transformation 268]
314. from_16(X0,sK16(X0)) | ~select_11(X0) [cnf transformation 272]
315. ~from_16(X0,X1) | select_11(X0) [cnf transformation 272]
316. ~main_query_2(X0) | select_11(X0) [cnf transformation 273]
317. ~select_11(X0) | main_query_2(X0) [cnf transformation 273]
318. eq(X0,X0) [cnf transformation 92]
370. main_query_2(sK17) | main_query_1(sK17) [cnf transformation 278]
371. ~main_query_2(sK17) | ~main_query_1(sK17) [cnf transformation 278]
417. 0 <=> ! [X3,X5,X2,X4,X6] : ~from_8(X2,X3,X4,X5,X6) [avatar definition]
418. ~from_8(X2,X3,X4,X5,X6) <- {0} [avatar component clause 417]
420. 2 <=> ! [X1,X0] : (select_7(X0,X1) | ~where_9(X0,X1)) [avatar definition]
421. ~where_9(X0,X1) | select_7(X0,X1) <- {2} [avatar component clause 420]
422. 0 | 2 [avatar split clause 287,420,417]
424. 4 <=> ! [X1,X3,X5,X2,X4] : ~from_2(X1,X2,X3,X4,X5) [avatar definition]
425. ~from_2(X1,X2,X3,X4,X5) <- {4} [avatar component clause 424]
427. 6 <=> ! [X0] : (select_1(X0) | ~where_3(X0)) [avatar definition]
428. ~where_3(X0) | select_1(X0) <- {6} [avatar component clause 427]
429. 4 | 6 [avatar split clause 300,427,424]
431. 8 <=> ! [X3,X5,X2,X4,X6] : ~from_14(X2,X3,X4,X5,X6) [avatar definition]
432. ~from_14(X2,X3,X4,X5,X6) <- {8} [avatar component clause 431]
434. 10 <=> ! [X1,X0] : (select_13(X0,X1) | ~where_15(X0,X1)) [avatar definition]
435. ~where_15(X0,X1) | select_13(X0,X1) <- {10} [avatar component clause 434]
436. 8 | 10 [avatar split clause 309,434,431]
438. 12 <=> main_query_1(sK17) [avatar definition]
439. main_query_1(sK17) <- {12} [avatar component clause 438]
441. 13 <=> ~main_query_1(sK17) [avatar definition]
442. ~main_query_1(sK17) <- {13} [avatar component clause 441]
444. 14 <=> main_query_2(sK17) [avatar definition]
445. main_query_2(sK17) <- {14} [avatar component clause 444]
447. 15 <=> ~main_query_2(sK17) [avatar definition]
448. ~main_query_2(sK17) <- {15} [avatar component clause 447]
449. ~13 | ~15 [avatar split clause 371,447,441]
450. 12 | 14 [avatar split clause 370,444,438]
458. where_9(X0,X0) [resolution 284,318]
462. where_15(X0,X0) [resolution 306,318]
464. query_12(X0,sK16(X0)) | ~select_11(X0) [resolution 314,312]
469. select_7(X0,X0) <- {2} [resolution 421,458]
470. query_6(X0,X0) <- {2} [resolution 469,289]
472. from_10(X0,X0) <- {2} [resolution 470,291]
476. select_5(X1) <- {2} [resolution 472,293]
478. query_4(X0) <- {2} [resolution 476,295]
480. where_3(X0) <- {2} [resolution 478,297]
494. tab1(sK6(X0),sK7(X0),sK8(X0),sK9(X0),sK10(X0)) | ~select_1(X0) [resolution 298,279]
498. select_11(sK17) <- {14} [resolution 445,316]
499. tab1(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_7(X0,X1) [resolution 285,281]
501. tab1(sK11(X0,X1),sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1)) | ~select_13(X0,X1) [resolution 307,303]
502. ~select_13(X2,X3) <- {8} [resolution 307,432]
504. select_13(X1,sK16(X1)) | ~select_11(X1) [resolution 464,310]
505. ~select_11(X1) <- {8} [subsumption resolution 504,502]
506. $false <- {8, 14} [resolution 505,498]
507. ~8 | ~14 [avatar contradiction clause 506]
508. select_13(X0,X0) <- {10} [resolution 435,462]
509. query_12(X0,X0) <- {10} [resolution 508,311]
511. from_14(sK6(X0),sK7(X0),sK8(X0),sK9(X0),sK10(X0)) | ~select_1(X0) [resolution 494,304]
514. from_16(X0,X0) <- {10} [resolution 509,313]
518. from_2(sK0(X4,X5),sK1(X4,X5),sK2(X4,X5),sK3(X4,X5),sK4(X4,X5)) | ~select_7(X4,X5) [resolution 499,280]
519. ~select_7(X4,X5) <- {4} [subsumption resolution 518,425]
520. $false <- {2, 4} [resolution 519,469]
523. ~2 | ~4 [avatar contradiction clause 520]
524. select_1(X0) <- {2, 6} [subsumption resolution 428,480]
525. main_query_1(X0) <- {2, 6} [resolution 524,302]
528. $false <- {2, 6, 13} [resolution 525,442]
529. ~2 | ~6 | 13 [avatar contradiction clause 528]
531. select_11(X1) <- {10} [resolution 514,315]
533. from_8(sK11(X2,X3),sK12(X2,X3),sK13(X2,X3),sK14(X2,X3),sK15(X2,X3)) | ~select_13(X2,X3) [resolution 501,282]
535. main_query_2(X0) <- {10} [resolution 531,317]
540. $false <- {10, 15} [resolution 535,448]
541. ~10 | 15 [avatar contradiction clause 540]
567. ~select_13(X0,X1) <- {0} [resolution 533,418]
571. ~select_11(X0) <- {0} [resolution 567,504]
574. $false <- {0, 10} [subsumption resolution 571,531]
575. ~0 | ~10 [avatar contradiction clause 574]
577. select_1(sK17) <- {12} [resolution 439,301]
586. ~select_1(X2) <- {8} [resolution 432,511]
591. $false <- {8, 12} [resolution 586,577]
592. ~8 | ~12 [avatar contradiction clause 591]
593. $false [avatar sat refutation 422,429,436,449,450,507,523,529,541,575,592]
% SZS output end Proof for query20_query21_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
