% Refutation found. Thanks to Tanya!
% SZS status Theorem for query07_query09_tptp4X
% SZS output start Proof for query07_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X6,X7,X8,X9,X10] : (from_7(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
6. ! [X6] : (select_6(X6) <=> ? [X7,X8,X9,X10] : ($true & from_7(X6,X7,X8,X9,X10))) [input]
7. ! [X5] : (query_5(X5) <=> select_6(X5)) [input]
8. ! [X11] : (from_8(X11) <=> query_5(X11)) [input]
9. ! [X11] : (select_4(X11) <=> ($true & from_8(X11))) [input]
10. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
11. ! [X12] : eq(X12,X12) [input]
17. ! [X12,X13] : (lt(X12,X13) <=> (~eq(X12,X13) & lte(X12,X13))) [input]
18. ! [X12,X13] : (gt(X12,X13) <=> lt(X13,X12)) [input]
44. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
45. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 44]
46. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
47. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
48. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [rectify 6]
49. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 48]
50. ! [X0] : (query_5(X0) <=> select_6(X0)) [rectify 7]
51. ! [X0] : (from_8(X0) <=> query_5(X0)) [rectify 8]
52. ! [X0] : (select_4(X0) <=> ($true & from_8(X0))) [rectify 9]
53. ! [X0] : (select_4(X0) <=> from_8(X0)) [true and false elimination 52]
54. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 10]
55. ! [X0] : eq(X0,X0) [rectify 11]
61. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 17]
62. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 18]
88. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 45]
140. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 88]
141. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
142. ! [X0] : ((where_3(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
143. ! [X0] : ((where_3(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [flattening 142]
144. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
145. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 144]
146. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
147. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 145,146]
148. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 46]
149. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 47]
150. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 49]
151. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 150]
152. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
153. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 151,152]
154. ! [X0] : ((query_5(X0) | ~select_6(X0)) & (select_6(X0) | ~query_5(X0))) [nnf transformation 50]
155. ! [X0] : ((from_8(X0) | ~query_5(X0)) & (query_5(X0) | ~from_8(X0))) [nnf transformation 51]
156. ! [X0] : ((select_4(X0) | ~from_8(X0)) & (from_8(X0) | ~select_4(X0))) [nnf transformation 53]
157. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 54]
160. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 61]
161. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 160]
162. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 62]
163. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 140]
164. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
165. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 163,164]
166. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 141]
167. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 141]
170. gt(X0,X0) | where_3(X0) | lt(X0,X0) [cnf transformation 143]
171. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 147]
173. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 147]
174. ~main_query_1(X0) | select_1(X0) [cnf transformation 148]
175. ~select_1(X0) | main_query_1(X0) [cnf transformation 148]
176. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 149]
177. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 149]
178. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 153]
179. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 153]
180. ~query_5(X0) | select_6(X0) [cnf transformation 154]
181. ~select_6(X0) | query_5(X0) [cnf transformation 154]
182. ~from_8(X0) | query_5(X0) [cnf transformation 155]
183. ~query_5(X0) | from_8(X0) [cnf transformation 155]
184. ~select_4(X0) | from_8(X0) [cnf transformation 156]
185. ~from_8(X0) | select_4(X0) [cnf transformation 156]
186. ~main_query_2(X0) | select_4(X0) [cnf transformation 157]
187. ~select_4(X0) | main_query_2(X0) [cnf transformation 157]
188. eq(X0,X0) [cnf transformation 55]
197. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 161]
199. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 162]
224. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 165]
225. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 165]
250. 0 <=> main_query_1(sK8) [avatar definition]
251. main_query_1(sK8) <- {0} [avatar component clause 250]
253. 1 <=> ~main_query_1(sK8) [avatar definition]
254. ~main_query_1(sK8) <- {1} [avatar component clause 253]
256. 2 <=> main_query_2(sK8) [avatar definition]
257. main_query_2(sK8) <- {2} [avatar component clause 256]
259. 3 <=> ~main_query_2(sK8) [avatar definition]
260. ~main_query_2(sK8) <- {3} [avatar component clause 259]
261. ~1 | ~3 [avatar split clause 225,259,253]
262. 0 | 2 [avatar split clause 224,256,250]
263. select_4(sK8) <- {2} [resolution 186,257]
264. from_8(sK8) <- {2} [resolution 263,184]
267. query_5(sK8) <- {2} [resolution 264,182]
269. select_6(sK8) <- {2} [resolution 267,180]
274. ~lt(X0,X0) [resolution 197,188]
276. where_3(X1) | lt(X1,X1) | lt(X1,X1) [resolution 170,199]
277. where_3(X1) | lt(X1,X1) [duplicate literal removal 276]
278. where_3(X1) [subsumption resolution 277,274]
283. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 171,166]
285. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 178,176]
287. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 283,177]
290. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 285,167]
292. ~select_1(X1) | select_6(X1) [resolution 287,179]
294. ~select_6(X1) | ~where_3(X1) | select_1(X1) [resolution 290,173]
295. ~select_6(X1) | select_1(X1) [subsumption resolution 294,278]
296. select_1(sK8) <- {2} [resolution 295,269]
298. main_query_1(sK8) <- {2} [resolution 296,175]
300. $false <- {1, 2} [subsumption resolution 298,254]
301. 1 | ~2 [avatar contradiction clause 300]
302. select_1(sK8) <- {0} [resolution 251,174]
303. select_6(sK8) <- {0} [resolution 302,292]
307. query_5(sK8) <- {0} [resolution 303,181]
308. from_8(sK8) <- {0} [resolution 307,183]
310. select_4(sK8) <- {0} [resolution 308,185]
312. main_query_2(sK8) <- {0} [resolution 310,187]
314. $false <- {0, 3} [subsumption resolution 312,260]
315. ~0 | 3 [avatar contradiction clause 314]
316. $false [avatar sat refutation 261,262,301,315]
% SZS output end Proof for query07_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.015 s
% ------------------------------
% ------------------------------
