% Refutation found. Thanks to Tanya!
% SZS status Theorem for query02_query08_tptp4X
% SZS output start Proof for query02_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_4(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (where_5(X6) <=> (gte(X6,X6) & lte(X6,X6))) [input]
6. ! [X6] : (select_3(X6) <=> ? [X7,X8,X9,X10] : (where_5(X6) & from_4(X6,X7,X8,X9,X10))) [input]
7. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
11. ! [X11] : lte(X11,X11) [input]
18. ! [X11,X12] : (gte(X11,X12) <=> lte(X12,X11)) [input]
42. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
43. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 42]
44. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
45. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
46. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
47. ! [X0] : (where_5(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 5]
48. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
49. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
53. ! [X0] : lte(X0,X0) [rectify 11]
60. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 18]
84. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 43]
136. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 84]
137. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 44]
139. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 138]
140. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
141. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 139,140]
142. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 45]
143. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 46]
144. ! [X0] : ((where_5(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [nnf transformation 47]
145. ! [X0] : ((where_5(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [flattening 144]
146. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 48]
147. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 146]
148. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
149. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 147,148]
150. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 49]
153. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 60]
154. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 136]
155. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
156. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 154,155]
157. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 137]
158. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 137]
159. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 141]
160. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 141]
161. ~main_query_1(X0) | select_1(X0) [cnf transformation 142]
162. ~select_1(X0) | main_query_1(X0) [cnf transformation 142]
163. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 143]
164. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 143]
167. ~gte(X0,X0) | where_5(X0) | ~lte(X0,X0) [cnf transformation 145]
168. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 149]
170. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 149]
171. ~main_query_2(X0) | select_3(X0) [cnf transformation 150]
172. ~select_3(X0) | main_query_2(X0) [cnf transformation 150]
176. lte(X0,X0) [cnf transformation 53]
183. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 153]
207. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 156]
208. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 156]
210. 0 <=> main_query_1(sK8) [avatar definition]
211. main_query_1(sK8) <- {0} [avatar component clause 210]
213. 1 <=> ~main_query_1(sK8) [avatar definition]
214. ~main_query_1(sK8) <- {1} [avatar component clause 213]
216. 2 <=> main_query_2(sK8) [avatar definition]
217. main_query_2(sK8) <- {2} [avatar component clause 216]
219. 3 <=> ~main_query_2(sK8) [avatar definition]
220. ~main_query_2(sK8) <- {3} [avatar component clause 219]
221. ~1 | ~3 [avatar split clause 208,219,213]
222. 0 | 2 [avatar split clause 207,216,210]
223. select_3(sK8) <- {2} [resolution 171,217]
232. gte(X1,X1) [resolution 183,176]
236. where_5(X0) | ~lte(X0,X0) [resolution 167,232]
238. where_5(X0) [subsumption resolution 236,176]
304. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 159,157]
306. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 168,163]
308. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 304,164]
312. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 306,158]
317. ~select_1(X5) | ~where_5(X5) | select_3(X5) [resolution 308,170]
318. ~select_1(X5) | select_3(X5) [subsumption resolution 317,238]
324. ~select_3(X1) | select_1(X1) [resolution 312,160]
327. select_1(sK8) <- {2} [resolution 324,223]
331. main_query_1(sK8) <- {2} [resolution 327,162]
332. $false <- {1, 2} [subsumption resolution 331,214]
333. 1 | ~2 [avatar contradiction clause 332]
334. select_1(sK8) <- {0} [resolution 211,161]
337. select_3(sK8) <- {0} [resolution 334,318]
342. main_query_2(sK8) <- {0} [resolution 337,172]
344. $false <- {0, 3} [subsumption resolution 342,220]
345. ~0 | 3 [avatar contradiction clause 344]
346. $false [avatar sat refutation 221,222,333,345]
% SZS output end Proof for query02_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.017 s
% ------------------------------
% ------------------------------
