% Refutation found. Thanks to Tanya!
% SZS status Theorem for query01_query04_tptp4X
% SZS output start Proof for query01_query04_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [input]
6. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
8. ! [X5] : eq(X5,X5) [input]
39. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
40. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 39]
41. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
42. ! [X0] : eq(X0,X0) [rectify 8]
118. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 40]
119. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
120. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 41]
121. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 120]
122. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
123. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 121,122]
124. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
125. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
126. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 5]
127. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 6]
128. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 127]
129. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
130. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 128,129]
131. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 7]
134. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 118]
135. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
136. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 134,135]
137. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 119]
138. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 119]
139. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 123]
140. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 123]
141. ~main_query_1(X0) | select_1(X0) [cnf transformation 124]
142. ~select_1(X0) | main_query_1(X0) [cnf transformation 124]
143. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 125]
144. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 125]
146. ~eq(X0,X0) | where_5(X0) [cnf transformation 126]
147. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 130]
149. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 130]
150. ~main_query_2(X0) | select_3(X0) [cnf transformation 131]
151. ~select_3(X0) | main_query_2(X0) [cnf transformation 131]
152. eq(X0,X0) [cnf transformation 42]
181. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 136]
182. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 136]
184. 0 <=> main_query_1(sK8) [avatar definition]
185. main_query_1(sK8) <- {0} [avatar component clause 184]
187. 1 <=> ~main_query_1(sK8) [avatar definition]
188. ~main_query_1(sK8) <- {1} [avatar component clause 187]
190. 2 <=> main_query_2(sK8) [avatar definition]
191. main_query_2(sK8) <- {2} [avatar component clause 190]
193. 3 <=> ~main_query_2(sK8) [avatar definition]
194. ~main_query_2(sK8) <- {3} [avatar component clause 193]
195. ~1 | ~3 [avatar split clause 182,193,187]
196. 0 | 2 [avatar split clause 181,190,184]
197. select_3(sK8) <- {2} [resolution 150,191]
201. where_5(X1) [resolution 146,152]
235. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 139,137]
237. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 147,143]
240. from_4(X2,sK0(X2),sK1(X2),sK2(X2),sK3(X2)) | ~select_1(X2) [resolution 235,144]
246. from_2(X5,sK4(X5),sK5(X5),sK6(X5),sK7(X5)) | ~select_3(X5) [resolution 237,138]
250. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 240,149]
251. ~select_1(X1) | select_3(X1) [subsumption resolution 250,201]
257. ~select_3(X1) | select_1(X1) [resolution 246,140]
260. select_1(sK8) <- {2} [resolution 257,197]
264. main_query_1(sK8) <- {2} [resolution 260,142]
265. $false <- {1, 2} [subsumption resolution 264,188]
266. 1 | ~2 [avatar contradiction clause 265]
267. select_1(sK8) <- {0} [resolution 185,141]
270. select_3(sK8) <- {0} [resolution 267,251]
275. main_query_2(sK8) <- {0} [resolution 270,151]
277. $false <- {0, 3} [subsumption resolution 275,194]
278. ~0 | 3 [avatar contradiction clause 277]
279. $false [avatar sat refutation 195,196,266,278]
% SZS output end Proof for query01_query04_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
