% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_04_query_08_tptp4X
% SZS output start Proof for query_04_query_08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [input]
4. ! [X5,X6,X2,X3,X4] : (from_7(X5,X6,X2,X3,X4) <=> tab1(X5,X6,X2,X3,X4)) [input]
5. ! [X5,X6] : (select_6(X5,X6) <=> ? [X2,X3,X4] : ($true & from_7(X5,X6,X2,X3,X4))) [input]
6. ! [X7,X8] : (subquery_from_5(X7,X8) <=> select_6(X7,X8)) [input]
7. ! [X9,X10] : (from_4(X9,X10) <=> subquery_from_5(X9,X10)) [input]
8. ! [X9,X10] : (select_3(X9,X10) <=> ($true & from_4(X9,X10))) [input]
9. ! [X7,X8] : (main_query_2(X7,X8) <=> select_3(X7,X8)) [input]
54. ! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [input]
55. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [negated conjecture 54]
56. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
57. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
58. ! [X0,X1] : (select_6(X0,X1) <=> ? [X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [rectify 5]
59. ! [X0,X1] : (select_6(X0,X1) <=> ? [X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 58]
60. ! [X0,X1] : (subquery_from_5(X0,X1) <=> select_6(X0,X1)) [rectify 6]
61. ! [X0,X1] : (from_4(X0,X1) <=> subquery_from_5(X0,X1)) [rectify 7]
62. ! [X0,X1] : (select_3(X0,X1) <=> ($true & from_4(X0,X1))) [rectify 8]
63. ! [X0,X1] : (select_3(X0,X1) <=> from_4(X0,X1)) [true and false elimination 62]
64. ! [X0,X1] : (main_query_2(X0,X1) <=> select_3(X0,X1)) [rectify 9]
182. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 55]
183. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
184. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0,X1))) [nnf transformation 56]
185. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) | ~select_1(X0,X1))) [rectify 184]
186. ! [X1,X0] : (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) => from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) [choice axiom]
187. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1))) [skolemisation 185,186]
188. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 3]
189. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 57]
190. ! [X0,X1] : ((select_6(X0,X1) | ! [X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0,X1))) [nnf transformation 59]
191. ! [X0,X1] : ((select_6(X0,X1) | ! [X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_7(X0,X1,X5,X6,X7) | ~select_6(X0,X1))) [rectify 190]
192. ! [X1,X0] : (? [X5,X6,X7] : from_7(X0,X1,X5,X6,X7) => from_7(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) [choice axiom]
193. ! [X0,X1] : ((select_6(X0,X1) | ! [X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_6(X0,X1))) [skolemisation 191,192]
194. ! [X0,X1] : ((subquery_from_5(X0,X1) | ~select_6(X0,X1)) & (select_6(X0,X1) | ~subquery_from_5(X0,X1))) [nnf transformation 60]
195. ! [X0,X1] : ((from_4(X0,X1) | ~subquery_from_5(X0,X1)) & (subquery_from_5(X0,X1) | ~from_4(X0,X1))) [nnf transformation 61]
196. ! [X0,X1] : ((select_3(X0,X1) | ~from_4(X0,X1)) & (from_4(X0,X1) | ~select_3(X0,X1))) [nnf transformation 63]
197. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 64]
200. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 182]
201. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
202. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 200,201]
203. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 183]
204. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 183]
205. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 187]
206. ~from_2(X0,X1,X2,X3,X4) | select_1(X0,X1) [cnf transformation 187]
207. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 188]
208. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 188]
209. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 189]
210. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 189]
211. from_7(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_6(X0,X1) [cnf transformation 193]
212. ~from_7(X0,X1,X2,X3,X4) | select_6(X0,X1) [cnf transformation 193]
213. ~subquery_from_5(X0,X1) | select_6(X0,X1) [cnf transformation 194]
214. ~select_6(X0,X1) | subquery_from_5(X0,X1) [cnf transformation 194]
215. ~from_4(X0,X1) | subquery_from_5(X0,X1) [cnf transformation 195]
216. ~subquery_from_5(X0,X1) | from_4(X0,X1) [cnf transformation 195]
217. ~select_3(X0,X1) | from_4(X0,X1) [cnf transformation 196]
218. ~from_4(X0,X1) | select_3(X0,X1) [cnf transformation 196]
219. ~main_query_2(X0,X1) | select_3(X0,X1) [cnf transformation 197]
220. ~select_3(X0,X1) | main_query_2(X0,X1) [cnf transformation 197]
263. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 202]
264. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 202]
266. 0 <=> main_query_1(sK6,sK7) [avatar definition]
267. main_query_1(sK6,sK7) <- {0} [avatar component clause 266]
269. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
270. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 269]
272. 2 <=> main_query_2(sK6,sK7) [avatar definition]
273. main_query_2(sK6,sK7) <- {2} [avatar component clause 272]
275. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
276. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 275]
277. ~1 | ~3 [avatar split clause 264,275,269]
278. 0 | 2 [avatar split clause 263,272,266]
279. select_3(sK6,sK7) <- {2} [resolution 219,273]
281. from_4(sK6,sK7) <- {2} [resolution 279,217]
284. subquery_from_5(sK6,sK7) <- {2} [resolution 281,215]
287. select_6(sK6,sK7) <- {2} [resolution 284,213]
307. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 205,203]
309. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_6(X0,X1) [resolution 211,209]
322. from_7(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 307,210]
328. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_6(X0,X1) [resolution 309,204]
336. ~select_1(X2,X3) | select_6(X2,X3) [resolution 322,212]
357. ~select_6(X2,X3) | select_1(X2,X3) [resolution 328,206]
370. select_1(sK6,sK7) <- {2} [resolution 357,287]
379. main_query_1(sK6,sK7) <- {2} [resolution 370,208]
380. $false <- {1, 2} [subsumption resolution 379,270]
381. 1 | ~2 [avatar contradiction clause 380]
382. select_1(sK6,sK7) <- {0} [resolution 267,207]
383. select_6(sK6,sK7) <- {0} [resolution 382,336]
386. subquery_from_5(sK6,sK7) <- {0} [resolution 383,214]
394. from_4(sK6,sK7) <- {0} [resolution 386,216]
403. select_3(sK6,sK7) <- {0} [resolution 394,218]
412. main_query_2(sK6,sK7) <- {0} [resolution 403,220]
414. $false <- {0, 3} [subsumption resolution 412,276]
415. ~0 | 3 [avatar contradiction clause 414]
416. $false [avatar sat refutation 277,278,381,415]
% SZS output end Proof for query_04_query_08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
