% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_13_query_28_tptp4X
% SZS output start Proof for query_13_query_28_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X7,X2,X3,X4] : (from_5(X0,X7,X2,X3,X4) <=> tab1(X0,X7,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0,X7] : (where_10(X0,X7) <=> eq(X0,X7)) [input]
8. ! [X0,X7] : (select_8(X0,X7) <=> ? [X1,X2,X3,X4] : (where_10(X0,X7) & from_9(X0,X1,X2,X3,X4))) [input]
9. ! [X5,X7] : (where_6_subquery_7(X5,X7) <=> select_8(X5,X7)) [input]
10. ! [X0,X7] : (where_6(X0,X7) <=> where_6_subquery_7(X0,X7)) [input]
11. ! [X0,X7] : (select_4(X0,X7) <=> ? [X2,X3,X4] : (where_6(X0,X7) & from_5(X0,X7,X2,X3,X4))) [input]
12. ! [X5,X7] : (main_query_2(X5,X7) <=> select_4(X5,X7)) [input]
17. ! [X8,X9] : ((lte(X9,X8) & lte(X8,X9)) <=> eq(X8,X9)) [input]
66. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
67. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 66]
68. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
69. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
70. ! [X0,X1] : (where_10(X0,X1) <=> eq(X0,X1)) [rectify 7]
71. ! [X0,X1] : (select_8(X0,X1) <=> ? [X2,X3,X4,X5] : (where_10(X0,X1) & from_9(X0,X2,X3,X4,X5))) [rectify 8]
72. ! [X0,X1] : (where_6_subquery_7(X0,X1) <=> select_8(X0,X1)) [rectify 9]
73. ! [X0,X1] : (where_6(X0,X1) <=> where_6_subquery_7(X0,X1)) [rectify 10]
74. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 11]
75. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 12]
80. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 17]
129. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 67]
221. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 129]
222. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
223. ! [X0,X1] : ((where_3(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [nnf transformation 2]
224. ! [X0,X1] : ((where_3(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [flattening 223]
225. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
226. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 225]
227. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
228. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 226,227]
229. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 68]
230. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 69]
231. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
232. ! [X0,X1] : ((where_10(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_10(X0,X1))) [nnf transformation 70]
233. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5] : (~where_10(X0,X1) | ~from_9(X0,X2,X3,X4,X5))) & (? [X2,X3,X4,X5] : (where_10(X0,X1) & from_9(X0,X2,X3,X4,X5)) | ~select_8(X0,X1))) [nnf transformation 71]
234. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5] : (~where_10(X0,X1) | ~from_9(X0,X2,X3,X4,X5))) & (? [X6,X7,X8,X9] : (where_10(X0,X1) & from_9(X0,X6,X7,X8,X9)) | ~select_8(X0,X1))) [rectify 233]
235. ! [X1,X0] : (? [X6,X7,X8,X9] : (where_10(X0,X1) & from_9(X0,X6,X7,X8,X9)) => (where_10(X0,X1) & from_9(X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1),sK6(X0,X1)))) [choice axiom]
236. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5] : (~where_10(X0,X1) | ~from_9(X0,X2,X3,X4,X5))) & ((where_10(X0,X1) & from_9(X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1),sK6(X0,X1))) | ~select_8(X0,X1))) [skolemisation 234,235]
237. ! [X0,X1] : ((where_6_subquery_7(X0,X1) | ~select_8(X0,X1)) & (select_8(X0,X1) | ~where_6_subquery_7(X0,X1))) [nnf transformation 72]
238. ! [X0,X1] : ((where_6(X0,X1) | ~where_6_subquery_7(X0,X1)) & (where_6_subquery_7(X0,X1) | ~where_6(X0,X1))) [nnf transformation 73]
239. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 74]
240. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 239]
241. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
242. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 240,241]
243. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 75]
244. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 80]
245. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 244]
246. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 221]
247. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
248. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 246,247]
249. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 222]
250. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 222]
251. ~where_3(X0,X1) | lte(X0,X1) [cnf transformation 224]
252. ~where_3(X0,X1) | lte(X1,X0) [cnf transformation 224]
253. ~lte(X1,X0) | where_3(X0,X1) | ~lte(X0,X1) [cnf transformation 224]
254. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 228]
255. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 228]
256. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 228]
257. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 229]
258. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 229]
259. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 230]
260. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 230]
262. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 231]
263. ~where_10(X0,X1) | eq(X0,X1) [cnf transformation 232]
264. ~eq(X0,X1) | where_10(X0,X1) [cnf transformation 232]
266. ~select_8(X0,X1) | where_10(X0,X1) [cnf transformation 236]
267. ~from_9(X0,X2,X3,X4,X5) | ~where_10(X0,X1) | select_8(X0,X1) [cnf transformation 236]
268. ~where_6_subquery_7(X0,X1) | select_8(X0,X1) [cnf transformation 237]
269. ~select_8(X0,X1) | where_6_subquery_7(X0,X1) [cnf transformation 237]
270. ~where_6(X0,X1) | where_6_subquery_7(X0,X1) [cnf transformation 238]
271. ~where_6_subquery_7(X0,X1) | where_6(X0,X1) [cnf transformation 238]
272. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 242]
273. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 242]
274. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 242]
275. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 243]
276. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 243]
281. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 245]
282. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 245]
283. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 245]
328. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 248]
329. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 248]
331. 0 <=> main_query_1(sK10,sK11) [avatar definition]
332. main_query_1(sK10,sK11) <- {0} [avatar component clause 331]
334. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
335. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 334]
337. 2 <=> main_query_2(sK10,sK11) [avatar definition]
338. main_query_2(sK10,sK11) <- {2} [avatar component clause 337]
340. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
341. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 340]
342. ~1 | ~3 [avatar split clause 329,340,334]
343. 0 | 2 [avatar split clause 328,337,331]
346. select_4(sK10,sK11) <- {2} [resolution 275,338]
348. where_6(sK10,sK11) <- {2} [resolution 346,273]
350. where_6_subquery_7(sK10,sK11) <- {2} [resolution 348,270]
353. select_8(sK10,sK11) <- {2} [resolution 350,268]
356. where_10(sK10,sK11) <- {2} [resolution 353,266]
362. eq(sK10,sK11) <- {2} [resolution 356,263]
365. lte(sK11,sK10) <- {2} [resolution 362,283]
366. lte(sK10,sK11) <- {2} [resolution 362,282]
372. where_3(sK10,sK11) | ~lte(sK10,sK11) <- {2} [resolution 365,253]
373. where_3(sK10,sK11) <- {2} [subsumption resolution 372,366]
481. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 254,249]
483. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 272,259]
558. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 481,260]
559. from_9(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 481,262]
567. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 483,250]
729. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 558,274]
738. ~where_10(X2,X4) | ~select_1(X2,X3) | select_8(X2,X4) [resolution 559,267]
747. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 567,256]
831. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 747,346]
835. select_1(sK10,sK11) <- {2} [subsumption resolution 831,373]
850. main_query_1(sK10,sK11) <- {2} [resolution 835,258]
852. $false <- {1, 2} [subsumption resolution 850,335]
853. 1 | ~2 [avatar contradiction clause 852]
854. select_1(sK10,sK11) <- {0} [resolution 332,257]
857. where_3(sK10,sK11) <- {0} [resolution 854,255]
858. lte(sK11,sK10) <- {0} [resolution 857,252]
859. lte(sK10,sK11) <- {0} [resolution 857,251]
869. eq(sK10,sK11) | ~lte(sK10,sK11) <- {0} [resolution 858,281]
871. eq(sK10,sK11) <- {0} [subsumption resolution 869,859]
893. where_10(sK10,sK11) <- {0} [resolution 871,264]
959. ~select_1(sK10,X0) | select_8(sK10,sK11) <- {0} [resolution 893,738]
965. 4 <=> select_8(sK10,sK11) [avatar definition]
966. select_8(sK10,sK11) <- {4} [avatar component clause 965]
968. 6 <=> ! [X0] : ~select_1(sK10,X0) [avatar definition]
969. ~select_1(sK10,X0) <- {6} [avatar component clause 968]
970. 4 | 6 | ~0 [avatar split clause 959,331,968,965]
1006. $false <- {0, 6} [resolution 969,854]
1007. ~0 | ~6 [avatar contradiction clause 1006]
1019. where_6_subquery_7(sK10,sK11) <- {4} [resolution 966,269]
1067. where_6(sK10,sK11) <- {4} [resolution 1019,271]
1090. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {4} [resolution 1067,729]
1094. select_4(sK10,sK11) <- {0, 4} [subsumption resolution 1090,854]
1123. main_query_2(sK10,sK11) <- {0, 4} [resolution 1094,276]
1125. $false <- {0, 3, 4} [subsumption resolution 1123,341]
1126. ~0 | 3 | ~4 [avatar contradiction clause 1125]
1127. $false [avatar sat refutation 342,343,853,970,1007,1126]
% SZS output end Proof for query_13_query_28_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.016 s
% ------------------------------
% ------------------------------
