% Refutation found. Thanks to Tanya!
% SZS status Theorem for query06_query08_tptp4X
% SZS output start Proof for query06_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (query_2(X5) <=> select_3(X5)) [input]
4. ! [X6] : (from_5(X6) <=> query_2(X6)) [input]
5. ! [X6] : (select_1(X6) <=> ($true & from_5(X6))) [input]
6. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
7. ! [X7,X8,X9,X10,X11] : (from_7(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
8. ! [X7] : (where_8(X7) <=> (gte(X7,X7) & lte(X7,X7))) [input]
9. ! [X7] : (select_6(X7) <=> ? [X8,X9,X10,X11] : (where_8(X7) & from_7(X7,X8,X9,X10,X11))) [input]
10. ! [X5] : (main_query_2(X5) <=> select_6(X5)) [input]
14. ! [X12] : lte(X12,X12) [input]
20. ! [X12,X13] : (gte(X12,X13) <=> lte(X13,X12)) [input]
52. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
53. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 52]
54. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 2]
55. ! [X0] : (query_2(X0) <=> select_3(X0)) [rectify 3]
56. ! [X0] : (from_5(X0) <=> query_2(X0)) [rectify 4]
57. ! [X0] : (select_1(X0) <=> ($true & from_5(X0))) [rectify 5]
58. ! [X0] : (select_1(X0) <=> from_5(X0)) [true and false elimination 57]
59. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 6]
60. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 7]
61. ! [X0] : (where_8(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 8]
62. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4))) [rectify 9]
63. ! [X0] : (main_query_2(X0) <=> select_6(X0)) [rectify 10]
67. ! [X0] : lte(X0,X0) [rectify 14]
73. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 20]
105. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 53]
173. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 105]
174. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 1]
175. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 54]
176. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 175]
177. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
178. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0))) [skolemisation 176,177]
179. ! [X0] : ((query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~query_2(X0))) [nnf transformation 55]
180. ! [X0] : ((from_5(X0) | ~query_2(X0)) & (query_2(X0) | ~from_5(X0))) [nnf transformation 56]
181. ! [X0] : ((select_1(X0) | ~from_5(X0)) & (from_5(X0) | ~select_1(X0))) [nnf transformation 58]
182. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 59]
183. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 60]
184. ! [X0] : ((where_8(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_8(X0))) [nnf transformation 61]
185. ! [X0] : ((where_8(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_8(X0))) [flattening 184]
186. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4)) | ~select_6(X0))) [nnf transformation 62]
187. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) | ~select_6(X0))) [rectify 186]
188. ! [X0] : (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) => (where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
189. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & ((where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_6(X0))) [skolemisation 187,188]
190. ! [X0] : ((main_query_2(X0) | ~select_6(X0)) & (select_6(X0) | ~main_query_2(X0))) [nnf transformation 63]
193. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 73]
194. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 173]
195. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
196. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 194,195]
197. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 174]
198. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 174]
199. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [cnf transformation 178]
200. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 178]
201. ~query_2(X0) | select_3(X0) [cnf transformation 179]
202. ~select_3(X0) | query_2(X0) [cnf transformation 179]
203. ~from_5(X0) | query_2(X0) [cnf transformation 180]
204. ~query_2(X0) | from_5(X0) [cnf transformation 180]
205. ~select_1(X0) | from_5(X0) [cnf transformation 181]
206. ~from_5(X0) | select_1(X0) [cnf transformation 181]
207. ~main_query_1(X0) | select_1(X0) [cnf transformation 182]
208. ~select_1(X0) | main_query_1(X0) [cnf transformation 182]
209. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 183]
210. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 183]
213. ~gte(X0,X0) | where_8(X0) | ~lte(X0,X0) [cnf transformation 185]
214. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 189]
216. ~from_7(X0,X1,X2,X3,X4) | ~where_8(X0) | select_6(X0) [cnf transformation 189]
217. ~main_query_2(X0) | select_6(X0) [cnf transformation 190]
218. ~select_6(X0) | main_query_2(X0) [cnf transformation 190]
222. lte(X0,X0) [cnf transformation 67]
228. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 193]
260. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 196]
261. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 196]
294. 0 <=> main_query_1(sK8) [avatar definition]
295. main_query_1(sK8) <- {0} [avatar component clause 294]
297. 1 <=> ~main_query_1(sK8) [avatar definition]
298. ~main_query_1(sK8) <- {1} [avatar component clause 297]
300. 2 <=> main_query_2(sK8) [avatar definition]
301. main_query_2(sK8) <- {2} [avatar component clause 300]
303. 3 <=> ~main_query_2(sK8) [avatar definition]
304. ~main_query_2(sK8) <- {3} [avatar component clause 303]
305. ~1 | ~3 [avatar split clause 261,303,297]
306. 0 | 2 [avatar split clause 260,300,294]
307. select_6(sK8) <- {2} [resolution 217,301]
315. gte(X1,X1) [resolution 228,222]
317. where_8(X0) | ~lte(X0,X0) [resolution 213,315]
319. where_8(X0) [subsumption resolution 317,222]
325. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [resolution 199,197]
327. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 214,209]
329. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [resolution 325,210]
332. from_4(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 327,198]
334. ~select_3(X1) | ~where_8(X1) | select_6(X1) [resolution 329,216]
335. ~select_3(X1) | select_6(X1) [subsumption resolution 334,319]
337. ~select_6(X1) | select_3(X1) [resolution 332,200]
338. select_3(sK8) <- {2} [resolution 337,307]
340. query_2(sK8) <- {2} [resolution 338,202]
341. from_5(sK8) <- {2} [resolution 340,204]
343. select_1(sK8) <- {2} [resolution 341,206]
345. main_query_1(sK8) <- {2} [resolution 343,208]
347. $false <- {1, 2} [subsumption resolution 345,298]
348. 1 | ~2 [avatar contradiction clause 347]
349. select_1(sK8) <- {0} [resolution 295,207]
351. from_5(sK8) <- {0} [resolution 349,205]
353. query_2(sK8) <- {0} [resolution 351,203]
355. select_3(sK8) <- {0} [resolution 353,201]
356. select_6(sK8) <- {0} [resolution 355,335]
359. main_query_2(sK8) <- {0} [resolution 356,218]
361. $false <- {0, 3} [subsumption resolution 359,304]
362. ~0 | 3 [avatar contradiction clause 361]
363. $false [avatar sat refutation 305,306,348,362]
% SZS output end Proof for query06_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
