% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query06_tptp4X
% SZS output start Proof for query00_query06_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_6(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (select_5(X6) <=> ? [X7,X8,X9,X10] : ($true & from_6(X6,X7,X8,X9,X10))) [input]
6. ! [X5] : (query_4(X5) <=> select_5(X5)) [input]
7. ! [X11] : (from_7(X11) <=> query_4(X11)) [input]
8. ! [X11] : (select_3(X11) <=> ($true & from_7(X11))) [input]
9. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
50. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
51. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 50]
52. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
53. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
54. ! [X0,X1,X2,X3,X4] : (from_6(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
55. ! [X0] : (select_5(X0) <=> ? [X1,X2,X3,X4] : ($true & from_6(X0,X1,X2,X3,X4))) [rectify 5]
56. ! [X0] : (select_5(X0) <=> ? [X1,X2,X3,X4] : from_6(X0,X1,X2,X3,X4)) [true and false elimination 55]
57. ! [X0] : (query_4(X0) <=> select_5(X0)) [rectify 6]
58. ! [X0] : (from_7(X0) <=> query_4(X0)) [rectify 7]
59. ! [X0] : (select_3(X0) <=> ($true & from_7(X0))) [rectify 8]
60. ! [X0] : (select_3(X0) <=> from_7(X0)) [true and false elimination 59]
61. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 9]
102. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 51]
168. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 102]
169. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
170. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 52]
171. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 170]
172. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
173. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 171,172]
174. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 53]
175. ! [X0,X1,X2,X3,X4] : ((from_6(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_6(X0,X1,X2,X3,X4))) [nnf transformation 54]
176. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_6(X0,X1,X2,X3,X4) | ~select_5(X0))) [nnf transformation 56]
177. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_6(X0,X5,X6,X7,X8) | ~select_5(X0))) [rectify 176]
178. ! [X0] : (? [X5,X6,X7,X8] : from_6(X0,X5,X6,X7,X8) => from_6(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
179. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (from_6(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_5(X0))) [skolemisation 177,178]
180. ! [X0] : ((query_4(X0) | ~select_5(X0)) & (select_5(X0) | ~query_4(X0))) [nnf transformation 57]
181. ! [X0] : ((from_7(X0) | ~query_4(X0)) & (query_4(X0) | ~from_7(X0))) [nnf transformation 58]
182. ! [X0] : ((select_3(X0) | ~from_7(X0)) & (from_7(X0) | ~select_3(X0))) [nnf transformation 60]
183. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 61]
186. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 168]
187. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
188. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 186,187]
189. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 169]
190. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 169]
191. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 173]
192. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 173]
193. ~main_query_1(X0) | select_1(X0) [cnf transformation 174]
194. ~select_1(X0) | main_query_1(X0) [cnf transformation 174]
195. ~from_6(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 175]
196. ~tab1(X0,X1,X2,X3,X4) | from_6(X0,X1,X2,X3,X4) [cnf transformation 175]
197. from_6(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_5(X0) [cnf transformation 179]
198. ~from_6(X0,X1,X2,X3,X4) | select_5(X0) [cnf transformation 179]
199. ~query_4(X0) | select_5(X0) [cnf transformation 180]
200. ~select_5(X0) | query_4(X0) [cnf transformation 180]
201. ~from_7(X0) | query_4(X0) [cnf transformation 181]
202. ~query_4(X0) | from_7(X0) [cnf transformation 181]
203. ~select_3(X0) | from_7(X0) [cnf transformation 182]
204. ~from_7(X0) | select_3(X0) [cnf transformation 182]
205. ~main_query_2(X0) | select_3(X0) [cnf transformation 183]
206. ~select_3(X0) | main_query_2(X0) [cnf transformation 183]
245. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 188]
246. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 188]
278. 0 <=> main_query_1(sK8) [avatar definition]
279. main_query_1(sK8) <- {0} [avatar component clause 278]
281. 1 <=> ~main_query_1(sK8) [avatar definition]
282. ~main_query_1(sK8) <- {1} [avatar component clause 281]
284. 2 <=> main_query_2(sK8) [avatar definition]
285. main_query_2(sK8) <- {2} [avatar component clause 284]
287. 3 <=> ~main_query_2(sK8) [avatar definition]
288. ~main_query_2(sK8) <- {3} [avatar component clause 287]
289. ~1 | ~3 [avatar split clause 246,287,281]
290. 0 | 2 [avatar split clause 245,284,278]
291. select_3(sK8) <- {2} [resolution 205,285]
293. from_7(sK8) <- {2} [resolution 291,203]
296. query_4(sK8) <- {2} [resolution 293,201]
299. select_5(sK8) <- {2} [resolution 296,199]
305. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 191,189]
307. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_5(X0) [resolution 197,195]
309. from_6(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 305,196]
312. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_5(X1) [resolution 307,190]
314. ~select_1(X1) | select_5(X1) [resolution 309,198]
316. ~select_5(X1) | select_1(X1) [resolution 312,192]
317. select_1(sK8) <- {2} [resolution 316,299]
319. main_query_1(sK8) <- {2} [resolution 317,194]
320. $false <- {1, 2} [subsumption resolution 319,282]
321. 1 | ~2 [avatar contradiction clause 320]
322. select_1(sK8) <- {0} [resolution 279,193]
323. select_5(sK8) <- {0} [resolution 322,314]
326. query_4(sK8) <- {0} [resolution 323,200]
327. from_7(sK8) <- {0} [resolution 326,202]
329. select_3(sK8) <- {0} [resolution 327,204]
331. main_query_2(sK8) <- {0} [resolution 329,206]
333. $false <- {0, 3} [subsumption resolution 331,288]
334. ~0 | 3 [avatar contradiction clause 333]
335. $false [avatar sat refutation 289,290,321,334]
% SZS output end Proof for query00_query06_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.014 s
% ------------------------------
% ------------------------------
