% Refutation found. Thanks to Tanya!
% SZS status Theorem for query04_query08_tptp4X
% SZS output start Proof for query04_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X5,X1,X2,X3,X4] : (from_5(X5,X1,X2,X3,X4) <=> tab1(X5,X1,X2,X3,X4)) [input]
6. ! [X5] : (where_6(X5) <=> (gte(X5,X5) & lte(X5,X5))) [input]
7. ! [X5] : (select_4(X5) <=> ? [X1,X2,X3,X4] : (where_6(X5) & from_5(X5,X1,X2,X3,X4))) [input]
8. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
12. ! [X7] : lte(X7,X7) [input]
19. ! [X7,X8] : (gte(X7,X8) <=> lte(X8,X7)) [input]
43. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
44. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 43]
45. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
46. ! [X0] : (where_6(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 6]
47. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
48. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
49. ! [X0] : eq(X0,X0) [rectify 9]
52. ! [X0] : lte(X0,X0) [rectify 12]
59. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 19]
134. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 44]
135. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
136. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 137]
139. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
140. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 138,139]
141. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
142. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 45]
143. ! [X0] : ((where_6(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [nnf transformation 46]
144. ! [X0] : ((where_6(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [flattening 143]
145. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 47]
146. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 145]
147. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
148. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 146,147]
149. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 48]
152. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 59]
153. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 134]
154. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
155. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 153,154]
156. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 135]
157. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 135]
159. ~eq(X0,X0) | where_3(X0) [cnf transformation 136]
160. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 140]
162. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 140]
163. ~main_query_1(X0) | select_1(X0) [cnf transformation 141]
164. ~select_1(X0) | main_query_1(X0) [cnf transformation 141]
165. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 142]
166. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 142]
169. ~gte(X0,X0) | where_6(X0) | ~lte(X0,X0) [cnf transformation 144]
170. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 148]
172. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 148]
173. ~main_query_2(X0) | select_4(X0) [cnf transformation 149]
174. ~select_4(X0) | main_query_2(X0) [cnf transformation 149]
175. eq(X0,X0) [cnf transformation 49]
178. lte(X0,X0) [cnf transformation 52]
185. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 152]
209. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 155]
210. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 155]
212. 0 <=> main_query_1(sK8) [avatar definition]
213. main_query_1(sK8) <- {0} [avatar component clause 212]
215. 1 <=> ~main_query_1(sK8) [avatar definition]
216. ~main_query_1(sK8) <- {1} [avatar component clause 215]
218. 2 <=> main_query_2(sK8) [avatar definition]
219. main_query_2(sK8) <- {2} [avatar component clause 218]
221. 3 <=> ~main_query_2(sK8) [avatar definition]
222. ~main_query_2(sK8) <- {3} [avatar component clause 221]
223. ~1 | ~3 [avatar split clause 210,221,215]
224. 0 | 2 [avatar split clause 209,218,212]
225. select_4(sK8) <- {2} [resolution 173,219]
229. where_3(X1) [resolution 159,175]
239. gte(X1,X1) [resolution 185,178]
244. where_6(X1) | ~lte(X1,X1) [resolution 239,169]
245. where_6(X1) [subsumption resolution 244,178]
312. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 160,156]
314. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 170,165]
316. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 312,166]
321. from_2(X3,sK4(X3),sK5(X3),sK6(X3),sK7(X3)) | ~select_4(X3) [resolution 314,157]
325. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 316,172]
326. ~select_1(X1) | select_4(X1) [subsumption resolution 325,245]
332. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 321,162]
333. ~select_4(X1) | select_1(X1) [subsumption resolution 332,229]
336. select_1(sK8) <- {2} [resolution 333,225]
338. main_query_1(sK8) <- {2} [resolution 336,164]
340. $false <- {1, 2} [subsumption resolution 338,216]
341. 1 | ~2 [avatar contradiction clause 340]
342. select_1(sK8) <- {0} [resolution 213,163]
343. select_4(sK8) <- {0} [resolution 342,326]
347. main_query_2(sK8) <- {0} [resolution 343,174]
349. $false <- {0, 3} [subsumption resolution 347,222]
350. ~0 | 3 [avatar contradiction clause 349]
351. $false [avatar sat refutation 223,224,341,350]
% SZS output end Proof for query04_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.017 s
% ------------------------------
% ------------------------------
