% Refutation found. Thanks to Tanya!
% SZS status Theorem for query06_query09_tptp4X
% SZS output start Proof for query06_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : ($true & from_5(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (subquery_from_3(X5) <=> select_4(X5)) [input]
4. ! [X6] : (from_2(X6) <=> subquery_from_3(X6)) [input]
5. ! [X6] : (select_1(X6) <=> ($true & from_2(X6))) [input]
6. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
7. ! [X7,X1,X2,X3,X4] : (from_10(X7,X1,X2,X3,X4) <=> tab1(X7,X1,X2,X3,X4)) [input]
8. ! [X7] : (select_9(X7) <=> ? [X1,X2,X3,X4] : ($true & from_10(X7,X1,X2,X3,X4))) [input]
9. ! [X7] : (subquery_from_8(X7) <=> select_9(X7)) [input]
10. ! [X8] : (from_7(X8) <=> subquery_from_8(X8)) [input]
11. ! [X8] : (select_6(X8) <=> ($true & from_7(X8))) [input]
12. ! [X7] : (main_query_2(X7) <=> select_6(X7)) [input]
44. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
45. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 44]
46. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4)) [true and false elimination 2]
47. ! [X0] : (subquery_from_3(X0) <=> select_4(X0)) [rectify 3]
48. ! [X0] : (from_2(X0) <=> subquery_from_3(X0)) [rectify 4]
49. ! [X0] : (select_1(X0) <=> ($true & from_2(X0))) [rectify 5]
50. ! [X0] : (select_1(X0) <=> from_2(X0)) [true and false elimination 49]
51. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 6]
52. ! [X0,X1,X2,X3,X4] : (from_10(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 7]
53. ! [X0] : (select_9(X0) <=> ? [X1,X2,X3,X4] : ($true & from_10(X0,X1,X2,X3,X4))) [rectify 8]
54. ! [X0] : (select_9(X0) <=> ? [X1,X2,X3,X4] : from_10(X0,X1,X2,X3,X4)) [true and false elimination 53]
55. ! [X0] : (subquery_from_8(X0) <=> select_9(X0)) [rectify 9]
56. ! [X0] : (from_7(X0) <=> subquery_from_8(X0)) [rectify 10]
57. ! [X0] : (select_6(X0) <=> ($true & from_7(X0))) [rectify 11]
58. ! [X0] : (select_6(X0) <=> from_7(X0)) [true and false elimination 57]
59. ! [X0] : (main_query_2(X0) <=> select_6(X0)) [rectify 12]
91. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 45]
137. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 91]
138. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 1]
139. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4) | ~select_4(X0))) [nnf transformation 46]
140. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) | ~select_4(X0))) [rectify 139]
141. ! [X0] : (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) => from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
142. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0))) [skolemisation 140,141]
143. ! [X0] : ((subquery_from_3(X0) | ~select_4(X0)) & (select_4(X0) | ~subquery_from_3(X0))) [nnf transformation 47]
144. ! [X0] : ((from_2(X0) | ~subquery_from_3(X0)) & (subquery_from_3(X0) | ~from_2(X0))) [nnf transformation 48]
145. ! [X0] : ((select_1(X0) | ~from_2(X0)) & (from_2(X0) | ~select_1(X0))) [nnf transformation 50]
146. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 51]
147. ! [X0,X1,X2,X3,X4] : ((from_10(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_10(X0,X1,X2,X3,X4))) [nnf transformation 52]
148. ! [X0] : ((select_9(X0) | ! [X1,X2,X3,X4] : ~from_10(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_10(X0,X1,X2,X3,X4) | ~select_9(X0))) [nnf transformation 54]
149. ! [X0] : ((select_9(X0) | ! [X1,X2,X3,X4] : ~from_10(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_10(X0,X5,X6,X7,X8) | ~select_9(X0))) [rectify 148]
150. ! [X0] : (? [X5,X6,X7,X8] : from_10(X0,X5,X6,X7,X8) => from_10(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
151. ! [X0] : ((select_9(X0) | ! [X1,X2,X3,X4] : ~from_10(X0,X1,X2,X3,X4)) & (from_10(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_9(X0))) [skolemisation 149,150]
152. ! [X0] : ((subquery_from_8(X0) | ~select_9(X0)) & (select_9(X0) | ~subquery_from_8(X0))) [nnf transformation 55]
153. ! [X0] : ((from_7(X0) | ~subquery_from_8(X0)) & (subquery_from_8(X0) | ~from_7(X0))) [nnf transformation 56]
154. ! [X0] : ((select_6(X0) | ~from_7(X0)) & (from_7(X0) | ~select_6(X0))) [nnf transformation 58]
155. ! [X0] : ((main_query_2(X0) | ~select_6(X0)) & (select_6(X0) | ~main_query_2(X0))) [nnf transformation 59]
158. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 137]
159. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
160. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 158,159]
161. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 138]
162. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 138]
163. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [cnf transformation 142]
164. ~from_5(X0,X1,X2,X3,X4) | select_4(X0) [cnf transformation 142]
165. ~subquery_from_3(X0) | select_4(X0) [cnf transformation 143]
166. ~select_4(X0) | subquery_from_3(X0) [cnf transformation 143]
167. ~from_2(X0) | subquery_from_3(X0) [cnf transformation 144]
168. ~subquery_from_3(X0) | from_2(X0) [cnf transformation 144]
169. ~select_1(X0) | from_2(X0) [cnf transformation 145]
170. ~from_2(X0) | select_1(X0) [cnf transformation 145]
171. ~main_query_1(X0) | select_1(X0) [cnf transformation 146]
172. ~select_1(X0) | main_query_1(X0) [cnf transformation 146]
173. ~from_10(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 147]
174. ~tab1(X0,X1,X2,X3,X4) | from_10(X0,X1,X2,X3,X4) [cnf transformation 147]
175. from_10(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_9(X0) [cnf transformation 151]
176. ~from_10(X0,X1,X2,X3,X4) | select_9(X0) [cnf transformation 151]
177. ~subquery_from_8(X0) | select_9(X0) [cnf transformation 152]
178. ~select_9(X0) | subquery_from_8(X0) [cnf transformation 152]
179. ~from_7(X0) | subquery_from_8(X0) [cnf transformation 153]
180. ~subquery_from_8(X0) | from_7(X0) [cnf transformation 153]
181. ~select_6(X0) | from_7(X0) [cnf transformation 154]
182. ~from_7(X0) | select_6(X0) [cnf transformation 154]
183. ~main_query_2(X0) | select_6(X0) [cnf transformation 155]
184. ~select_6(X0) | main_query_2(X0) [cnf transformation 155]
214. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 160]
215. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 160]
217. 0 <=> main_query_1(sK8) [avatar definition]
218. main_query_1(sK8) <- {0} [avatar component clause 217]
220. 1 <=> ~main_query_1(sK8) [avatar definition]
221. ~main_query_1(sK8) <- {1} [avatar component clause 220]
223. 2 <=> main_query_2(sK8) [avatar definition]
224. main_query_2(sK8) <- {2} [avatar component clause 223]
226. 3 <=> ~main_query_2(sK8) [avatar definition]
227. ~main_query_2(sK8) <- {3} [avatar component clause 226]
228. ~1 | ~3 [avatar split clause 215,226,220]
229. 0 | 2 [avatar split clause 214,223,217]
230. select_6(sK8) <- {2} [resolution 183,224]
232. from_7(sK8) <- {2} [resolution 230,181]
235. subquery_from_8(sK8) <- {2} [resolution 232,179]
238. select_9(sK8) <- {2} [resolution 235,177]
267. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [resolution 163,161]
269. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_9(X0) [resolution 175,173]
272. from_10(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [resolution 267,174]
276. from_5(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_9(X1) [resolution 269,162]
279. ~select_4(X1) | select_9(X1) [resolution 272,176]
283. ~select_9(X1) | select_4(X1) [resolution 276,164]
291. select_4(sK8) <- {2} [resolution 283,238]
295. subquery_from_3(sK8) <- {2} [resolution 291,166]
298. from_2(sK8) <- {2} [resolution 295,168]
302. select_1(sK8) <- {2} [resolution 298,170]
306. main_query_1(sK8) <- {2} [resolution 302,172]
308. $false <- {1, 2} [subsumption resolution 306,221]
309. 1 | ~2 [avatar contradiction clause 308]
310. select_1(sK8) <- {0} [resolution 218,171]
312. from_2(sK8) <- {0} [resolution 310,169]
314. subquery_from_3(sK8) <- {0} [resolution 312,167]
316. select_4(sK8) <- {0} [resolution 314,165]
317. select_9(sK8) <- {0} [resolution 316,279]
320. subquery_from_8(sK8) <- {0} [resolution 317,178]
323. from_7(sK8) <- {0} [resolution 320,180]
327. select_6(sK8) <- {0} [resolution 323,182]
331. main_query_2(sK8) <- {0} [resolution 327,184]
333. $false <- {0, 3} [subsumption resolution 331,227]
334. ~0 | 3 [avatar contradiction clause 333]
335. $false [avatar sat refutation 228,229,309,334]
% SZS output end Proof for query06_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.018 s
% ------------------------------
% ------------------------------
