% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_15_query_16_tptp4X
% SZS output start Proof for query_15_query_16_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1,X2,X3,X4] : (from_6(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
3. ! [X0] : (select_5(X0) <=> ? [X1,X2,X3,X4] : ($true & from_6(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (where_3_subquery_4(X5) <=> select_5(X5)) [input]
5. ! [X0,X1] : (where_3(X0,X1) <=> (where_3_subquery_4(X0) & eq(X0,X1))) [input]
6. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
7. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
8. ! [X0,X1,X2,X3,X4] : (from_11(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
9. ! [X0,X1] : (where_12(X0,X1) <=> eq(X0,X1)) [input]
10. ! [X0,X1] : (select_10(X0,X1) <=> ? [X2,X3,X4] : (where_12(X0,X1) & from_11(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X6] : (subquery_from_9(X5,X6) <=> select_10(X5,X6)) [input]
12. ! [X7,X8] : (from_8(X7,X8) <=> subquery_from_9(X7,X8)) [input]
13. ! [X7,X8] : (select_7(X7,X8) <=> ($true & from_8(X7,X8))) [input]
14. ! [X5,X6] : (main_query_2(X5,X6) <=> select_7(X5,X6)) [input]
68. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
69. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 68]
70. ! [X0] : (select_5(X0) <=> ? [X1,X2,X3,X4] : from_6(X0,X1,X2,X3,X4)) [true and false elimination 3]
71. ! [X0] : (where_3_subquery_4(X0) <=> select_5(X0)) [rectify 4]
72. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 7]
73. ! [X0,X1] : (subquery_from_9(X0,X1) <=> select_10(X0,X1)) [rectify 11]
74. ! [X0,X1] : (from_8(X0,X1) <=> subquery_from_9(X0,X1)) [rectify 12]
75. ! [X0,X1] : (select_7(X0,X1) <=> ($true & from_8(X0,X1))) [rectify 13]
76. ! [X0,X1] : (select_7(X0,X1) <=> from_8(X0,X1)) [true and false elimination 75]
77. ! [X0,X1] : (main_query_2(X0,X1) <=> select_7(X0,X1)) [rectify 14]
131. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 69]
223. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 131]
224. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
225. ! [X0,X1,X2,X3,X4] : ((from_6(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_6(X0,X1,X2,X3,X4))) [nnf transformation 2]
226. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_6(X0,X1,X2,X3,X4) | ~select_5(X0))) [nnf transformation 70]
227. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_6(X0,X5,X6,X7,X8) | ~select_5(X0))) [rectify 226]
228. ! [X0] : (? [X5,X6,X7,X8] : from_6(X0,X5,X6,X7,X8) => from_6(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
229. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (from_6(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_5(X0))) [skolemisation 227,228]
230. ! [X0] : ((where_3_subquery_4(X0) | ~select_5(X0)) & (select_5(X0) | ~where_3_subquery_4(X0))) [nnf transformation 71]
231. ! [X0,X1] : ((where_3(X0,X1) | (~where_3_subquery_4(X0) | ~eq(X0,X1))) & ((where_3_subquery_4(X0) & eq(X0,X1)) | ~where_3(X0,X1))) [nnf transformation 5]
232. ! [X0,X1] : ((where_3(X0,X1) | ~where_3_subquery_4(X0) | ~eq(X0,X1)) & ((where_3_subquery_4(X0) & eq(X0,X1)) | ~where_3(X0,X1))) [flattening 231]
233. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 6]
234. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 233]
235. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1)))) [choice axiom]
236. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1))) | ~select_1(X0,X1))) [skolemisation 234,235]
237. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 72]
238. ! [X0,X1,X2,X3,X4] : ((from_11(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_11(X0,X1,X2,X3,X4))) [nnf transformation 8]
239. ! [X0,X1] : ((where_12(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_12(X0,X1))) [nnf transformation 9]
240. ! [X0,X1] : ((select_10(X0,X1) | ! [X2,X3,X4] : (~where_12(X0,X1) | ~from_11(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_12(X0,X1) & from_11(X0,X1,X2,X3,X4)) | ~select_10(X0,X1))) [nnf transformation 10]
241. ! [X0,X1] : ((select_10(X0,X1) | ! [X2,X3,X4] : (~where_12(X0,X1) | ~from_11(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_12(X0,X1) & from_11(X0,X1,X5,X6,X7)) | ~select_10(X0,X1))) [rectify 240]
242. ! [X1,X0] : (? [X5,X6,X7] : (where_12(X0,X1) & from_11(X0,X1,X5,X6,X7)) => (where_12(X0,X1) & from_11(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
243. ! [X0,X1] : ((select_10(X0,X1) | ! [X2,X3,X4] : (~where_12(X0,X1) | ~from_11(X0,X1,X2,X3,X4))) & ((where_12(X0,X1) & from_11(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_10(X0,X1))) [skolemisation 241,242]
244. ! [X0,X1] : ((subquery_from_9(X0,X1) | ~select_10(X0,X1)) & (select_10(X0,X1) | ~subquery_from_9(X0,X1))) [nnf transformation 73]
245. ! [X0,X1] : ((from_8(X0,X1) | ~subquery_from_9(X0,X1)) & (subquery_from_9(X0,X1) | ~from_8(X0,X1))) [nnf transformation 74]
246. ! [X0,X1] : ((select_7(X0,X1) | ~from_8(X0,X1)) & (from_8(X0,X1) | ~select_7(X0,X1))) [nnf transformation 76]
247. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_7(X0,X1)) & (select_7(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 77]
250. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 223]
251. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
252. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 250,251]
253. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 224]
254. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 224]
256. ~tab1(X0,X1,X2,X3,X4) | from_6(X0,X1,X2,X3,X4) [cnf transformation 225]
258. ~from_6(X0,X1,X2,X3,X4) | select_5(X0) [cnf transformation 229]
260. ~select_5(X0) | where_3_subquery_4(X0) [cnf transformation 230]
261. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 232]
263. ~eq(X0,X1) | ~where_3_subquery_4(X0) | where_3(X0,X1) [cnf transformation 232]
264. from_2(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1)) | ~select_1(X0,X1) [cnf transformation 236]
265. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 236]
266. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 236]
267. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 237]
268. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 237]
269. ~from_11(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 238]
270. ~tab1(X0,X1,X2,X3,X4) | from_11(X0,X1,X2,X3,X4) [cnf transformation 238]
271. ~where_12(X0,X1) | eq(X0,X1) [cnf transformation 239]
272. ~eq(X0,X1) | where_12(X0,X1) [cnf transformation 239]
273. from_11(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_10(X0,X1) [cnf transformation 243]
274. ~select_10(X0,X1) | where_12(X0,X1) [cnf transformation 243]
275. ~from_11(X0,X1,X2,X3,X4) | ~where_12(X0,X1) | select_10(X0,X1) [cnf transformation 243]
276. ~subquery_from_9(X0,X1) | select_10(X0,X1) [cnf transformation 244]
277. ~select_10(X0,X1) | subquery_from_9(X0,X1) [cnf transformation 244]
278. ~from_8(X0,X1) | subquery_from_9(X0,X1) [cnf transformation 245]
279. ~subquery_from_9(X0,X1) | from_8(X0,X1) [cnf transformation 245]
280. ~select_7(X0,X1) | from_8(X0,X1) [cnf transformation 246]
281. ~from_8(X0,X1) | select_7(X0,X1) [cnf transformation 246]
282. ~main_query_2(X0,X1) | select_7(X0,X1) [cnf transformation 247]
283. ~select_7(X0,X1) | main_query_2(X0,X1) [cnf transformation 247]
335. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 252]
336. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 252]
338. 0 <=> main_query_1(sK10,sK11) [avatar definition]
339. main_query_1(sK10,sK11) <- {0} [avatar component clause 338]
341. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
342. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 341]
344. 2 <=> main_query_2(sK10,sK11) [avatar definition]
345. main_query_2(sK10,sK11) <- {2} [avatar component clause 344]
347. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
348. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 347]
349. ~1 | ~3 [avatar split clause 336,347,341]
350. 0 | 2 [avatar split clause 335,344,338]
353. select_7(sK10,sK11) <- {2} [resolution 282,345]
354. from_8(sK10,sK11) <- {2} [resolution 353,280]
357. subquery_from_9(sK10,sK11) <- {2} [resolution 354,278]
360. select_10(sK10,sK11) <- {2} [resolution 357,276]
363. where_12(sK10,sK11) <- {2} [resolution 360,274]
365. eq(sK10,sK11) <- {2} [resolution 363,271]
371. ~where_3_subquery_4(sK10) | where_3(sK10,sK11) <- {2} [resolution 263,365]
376. 4 <=> where_3(sK10,sK11) [avatar definition]
377. where_3(sK10,sK11) <- {4} [avatar component clause 376]
382. 7 <=> ~where_3_subquery_4(sK10) [avatar definition]
383. ~where_3_subquery_4(sK10) <- {7} [avatar component clause 382]
384. 4 | ~7 | ~2 [avatar split clause 371,344,382,376]
507. tab1(X2,X3,sK4(X2,X3),sK5(X2,X3),sK6(X2,X3)) | ~select_1(X2,X3) [resolution 264,253]
510. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_10(X2,X3) [resolution 273,269]
596. from_11(X4,X5,sK4(X4,X5),sK5(X4,X5),sK6(X4,X5)) | ~select_1(X4,X5) [resolution 507,270]
604. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_10(X0,X1) [resolution 510,254]
605. from_6(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_10(X2,X3) [resolution 510,256]
757. 14 <=> select_5(sK10) [avatar definition]
758. select_5(sK10) <- {14} [avatar component clause 757]
826. ~where_12(X2,X3) | ~select_1(X2,X3) | select_10(X2,X3) [resolution 596,275]
833. ~select_10(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 604,266]
843. ~select_10(X2,X3) | select_5(X2) [resolution 605,258]
849. select_5(sK10) <- {2} [resolution 843,360]
864. 14 | ~2 [avatar split clause 849,344,757]
878. select_1(sK10,sK11) <- {0} [resolution 339,267]
879. where_3_subquery_4(sK10) <- {14} [resolution 758,260]
880. $false <- {7, 14} [subsumption resolution 879,383]
881. 7 | ~14 [avatar contradiction clause 880]
886. where_3(sK10,sK11) <- {0} [resolution 878,265]
889. 4 | ~0 [avatar split clause 886,338,376]
895. eq(sK10,sK11) <- {4} [resolution 377,261]
900. where_12(sK10,sK11) <- {4} [resolution 895,272]
944. ~select_1(sK10,sK11) | select_10(sK10,sK11) <- {4} [resolution 900,826]
948. select_10(sK10,sK11) <- {0, 4} [subsumption resolution 944,878]
1045. subquery_from_9(sK10,sK11) <- {0, 4} [resolution 948,277]
1088. from_8(sK10,sK11) <- {0, 4} [resolution 1045,279]
1108. select_7(sK10,sK11) <- {0, 4} [resolution 1088,281]
1119. main_query_2(sK10,sK11) <- {0, 4} [resolution 1108,283]
1121. $false <- {0, 3, 4} [subsumption resolution 1119,348]
1122. ~0 | 3 | ~4 [avatar contradiction clause 1121]
1127. 20 <=> select_10(sK10,sK11) [avatar definition]
1128. select_10(sK10,sK11) <- {20} [avatar component clause 1127]
1130. 22 <=> select_1(sK10,sK11) [avatar definition]
1131. select_1(sK10,sK11) <- {22} [avatar component clause 1130]
1138. select_7(sK10,sK11) <- {2} [resolution 345,282]
1142. from_8(sK10,sK11) <- {2} [resolution 1138,280]
1146. subquery_from_9(sK10,sK11) <- {2} [resolution 1142,278]
1152. select_10(sK10,sK11) <- {2} [resolution 1146,276]
1153. 20 | ~2 [avatar split clause 1152,344,1127]
1450. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {20} [resolution 833,1128]
1454. select_1(sK10,sK11) <- {4, 20} [subsumption resolution 1450,377]
1463. 22 | ~4 | ~20 [avatar split clause 1454,1127,376,1130]
1479. main_query_1(sK10,sK11) <- {22} [resolution 1131,268]
1485. $false <- {1, 22} [subsumption resolution 1479,342]
1486. 1 | ~22 [avatar contradiction clause 1485]
1487. $false [avatar sat refutation 349,350,384,864,881,889,1122,1153,1463,1486]
% SZS output end Proof for query_15_query_16_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.018 s
% ------------------------------
% ------------------------------
