% Refutation found. Thanks to Tanya!
% SZS status Theorem for query06_query09_tptp4X
% SZS output start Proof for query06_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (query_2(X5) <=> select_3(X5)) [input]
4. ! [X6] : (from_5(X6) <=> query_2(X6)) [input]
5. ! [X6] : (select_1(X6) <=> ($true & from_5(X6))) [input]
6. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
7. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
8. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
9. ! [X7] : (query_7(X7) <=> select_8(X7)) [input]
10. ! [X8] : (from_10(X8) <=> query_7(X8)) [input]
11. ! [X8] : (select_6(X8) <=> ($true & from_10(X8))) [input]
12. ! [X7] : (main_query_2(X7) <=> select_6(X7)) [input]
46. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
47. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 46]
48. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0] : (query_2(X0) <=> select_3(X0)) [rectify 3]
50. ! [X0] : (from_5(X0) <=> query_2(X0)) [rectify 4]
51. ! [X0] : (select_1(X0) <=> ($true & from_5(X0))) [rectify 5]
52. ! [X0] : (select_1(X0) <=> from_5(X0)) [true and false elimination 51]
53. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 6]
54. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 8]
55. ! [X0] : (query_7(X0) <=> select_8(X0)) [rectify 9]
56. ! [X0] : (from_10(X0) <=> query_7(X0)) [rectify 10]
57. ! [X0] : (select_6(X0) <=> ($true & from_10(X0))) [rectify 11]
58. ! [X0] : (select_6(X0) <=> from_10(X0)) [true and false elimination 57]
59. ! [X0] : (main_query_2(X0) <=> select_6(X0)) [rectify 12]
93. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 47]
145. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 93]
146. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 1]
147. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 48]
148. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 147]
149. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
150. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0))) [skolemisation 148,149]
151. ! [X0] : ((query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~query_2(X0))) [nnf transformation 49]
152. ! [X0] : ((from_5(X0) | ~query_2(X0)) & (query_2(X0) | ~from_5(X0))) [nnf transformation 50]
153. ! [X0] : ((select_1(X0) | ~from_5(X0)) & (from_5(X0) | ~select_1(X0))) [nnf transformation 52]
154. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 53]
155. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 7]
156. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 54]
157. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 156]
158. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
159. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_8(X0))) [skolemisation 157,158]
160. ! [X0] : ((query_7(X0) | ~select_8(X0)) & (select_8(X0) | ~query_7(X0))) [nnf transformation 55]
161. ! [X0] : ((from_10(X0) | ~query_7(X0)) & (query_7(X0) | ~from_10(X0))) [nnf transformation 56]
162. ! [X0] : ((select_6(X0) | ~from_10(X0)) & (from_10(X0) | ~select_6(X0))) [nnf transformation 58]
163. ! [X0] : ((main_query_2(X0) | ~select_6(X0)) & (select_6(X0) | ~main_query_2(X0))) [nnf transformation 59]
166. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 145]
167. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
168. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 166,167]
169. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 146]
170. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 146]
171. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [cnf transformation 150]
172. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 150]
173. ~query_2(X0) | select_3(X0) [cnf transformation 151]
174. ~select_3(X0) | query_2(X0) [cnf transformation 151]
175. ~from_5(X0) | query_2(X0) [cnf transformation 152]
176. ~query_2(X0) | from_5(X0) [cnf transformation 152]
177. ~select_1(X0) | from_5(X0) [cnf transformation 153]
178. ~from_5(X0) | select_1(X0) [cnf transformation 153]
179. ~main_query_1(X0) | select_1(X0) [cnf transformation 154]
180. ~select_1(X0) | main_query_1(X0) [cnf transformation 154]
181. ~from_9(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 155]
182. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 155]
183. from_9(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_8(X0) [cnf transformation 159]
184. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 159]
185. ~query_7(X0) | select_8(X0) [cnf transformation 160]
186. ~select_8(X0) | query_7(X0) [cnf transformation 160]
187. ~from_10(X0) | query_7(X0) [cnf transformation 161]
188. ~query_7(X0) | from_10(X0) [cnf transformation 161]
189. ~select_6(X0) | from_10(X0) [cnf transformation 162]
190. ~from_10(X0) | select_6(X0) [cnf transformation 162]
191. ~main_query_2(X0) | select_6(X0) [cnf transformation 163]
192. ~select_6(X0) | main_query_2(X0) [cnf transformation 163]
224. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 168]
225. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 168]
250. 0 <=> main_query_1(sK8) [avatar definition]
251. main_query_1(sK8) <- {0} [avatar component clause 250]
253. 1 <=> ~main_query_1(sK8) [avatar definition]
254. ~main_query_1(sK8) <- {1} [avatar component clause 253]
256. 2 <=> main_query_2(sK8) [avatar definition]
257. main_query_2(sK8) <- {2} [avatar component clause 256]
259. 3 <=> ~main_query_2(sK8) [avatar definition]
260. ~main_query_2(sK8) <- {3} [avatar component clause 259]
261. ~1 | ~3 [avatar split clause 225,259,253]
262. 0 | 2 [avatar split clause 224,256,250]
263. select_6(sK8) <- {2} [resolution 191,257]
265. from_10(sK8) <- {2} [resolution 263,189]
268. query_7(sK8) <- {2} [resolution 265,187]
271. select_8(sK8) <- {2} [resolution 268,185]
277. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [resolution 171,169]
279. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_8(X0) [resolution 183,181]
281. from_9(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [resolution 277,182]
284. from_4(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_8(X1) [resolution 279,170]
286. ~select_3(X1) | select_8(X1) [resolution 281,184]
288. ~select_8(X1) | select_3(X1) [resolution 284,172]
289. select_3(sK8) <- {2} [resolution 288,271]
291. query_2(sK8) <- {2} [resolution 289,174]
292. from_5(sK8) <- {2} [resolution 291,176]
294. select_1(sK8) <- {2} [resolution 292,178]
296. main_query_1(sK8) <- {2} [resolution 294,180]
298. $false <- {1, 2} [subsumption resolution 296,254]
299. 1 | ~2 [avatar contradiction clause 298]
300. select_1(sK8) <- {0} [resolution 251,179]
302. from_5(sK8) <- {0} [resolution 300,177]
304. query_2(sK8) <- {0} [resolution 302,175]
306. select_3(sK8) <- {0} [resolution 304,173]
307. select_8(sK8) <- {0} [resolution 306,286]
310. query_7(sK8) <- {0} [resolution 307,186]
311. from_10(sK8) <- {0} [resolution 310,188]
313. select_6(sK8) <- {0} [resolution 311,190]
315. main_query_2(sK8) <- {0} [resolution 313,192]
317. $false <- {0, 3} [subsumption resolution 315,260]
318. ~0 | 3 [avatar contradiction clause 317]
319. $false [avatar sat refutation 261,262,299,318]
% SZS output end Proof for query06_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.013 s
% ------------------------------
% ------------------------------
