% Refutation found. Thanks to Tanya!
% SZS status Theorem for query03_query08_tptp4X
% SZS output start Proof for query03_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X5,X6,X7,X8,X9] : (from_4(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
5. ! [X5] : (where_5(X5) <=> (gte(X5,X5) & lte(X5,X5))) [input]
6. ! [X5] : (select_3(X5) <=> ? [X6,X7,X8,X9] : (where_5(X5) & from_4(X5,X6,X7,X8,X9))) [input]
7. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
11. ! [X10] : lte(X10,X10) [input]
17. ! [X10,X11] : (gte(X10,X11) <=> lte(X11,X10)) [input]
41. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
42. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
45. ! [X0] : (where_5(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 5]
46. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
47. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
51. ! [X0] : lte(X0,X0) [rectify 11]
57. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 17]
132. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 42]
133. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 134]
136. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 135,136]
138. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
139. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 44]
140. ! [X0] : ((where_5(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [nnf transformation 45]
141. ! [X0] : ((where_5(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [flattening 140]
142. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 46]
143. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 142]
144. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
145. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 143,144]
146. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 47]
149. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 57]
150. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 132]
151. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
152. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 150,151]
153. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 133]
154. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 133]
155. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 137]
156. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 137]
157. ~main_query_1(X0) | select_1(X0) [cnf transformation 138]
158. ~select_1(X0) | main_query_1(X0) [cnf transformation 138]
159. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 139]
160. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 139]
163. ~gte(X0,X0) | where_5(X0) | ~lte(X0,X0) [cnf transformation 141]
164. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 145]
166. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 145]
167. ~main_query_2(X0) | select_3(X0) [cnf transformation 146]
168. ~select_3(X0) | main_query_2(X0) [cnf transformation 146]
172. lte(X0,X0) [cnf transformation 51]
178. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 149]
202. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 152]
203. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 152]
228. 0 <=> main_query_1(sK8) [avatar definition]
229. main_query_1(sK8) <- {0} [avatar component clause 228]
231. 1 <=> ~main_query_1(sK8) [avatar definition]
232. ~main_query_1(sK8) <- {1} [avatar component clause 231]
234. 2 <=> main_query_2(sK8) [avatar definition]
235. main_query_2(sK8) <- {2} [avatar component clause 234]
237. 3 <=> ~main_query_2(sK8) [avatar definition]
238. ~main_query_2(sK8) <- {3} [avatar component clause 237]
239. ~1 | ~3 [avatar split clause 203,237,231]
240. 0 | 2 [avatar split clause 202,234,228]
241. select_3(sK8) <- {2} [resolution 167,235]
249. gte(X1,X1) [resolution 178,172]
251. where_5(X0) | ~lte(X0,X0) [resolution 163,249]
253. where_5(X0) [subsumption resolution 251,172]
259. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 155,153]
261. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 164,159]
263. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 259,160]
266. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 261,154]
268. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 263,166]
269. ~select_1(X1) | select_3(X1) [subsumption resolution 268,253]
271. ~select_3(X1) | select_1(X1) [resolution 266,156]
272. select_1(sK8) <- {2} [resolution 271,241]
274. main_query_1(sK8) <- {2} [resolution 272,158]
275. $false <- {1, 2} [subsumption resolution 274,232]
276. 1 | ~2 [avatar contradiction clause 275]
277. select_1(sK8) <- {0} [resolution 229,157]
278. select_3(sK8) <- {0} [resolution 277,269]
281. main_query_2(sK8) <- {0} [resolution 278,168]
283. $false <- {0, 3} [subsumption resolution 281,238]
284. ~0 | 3 [avatar contradiction clause 283]
285. $false [avatar sat refutation 239,240,276,284]
% SZS output end Proof for query03_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
