% Refutation found. Thanks to Tanya!
% SZS status Theorem for query03_query07_tptp4X
% SZS output start Proof for query03_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_4(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (where_5(X6) <=> (~gt(X6,X6) & ~lt(X6,X6))) [input]
6. ! [X6] : (select_3(X6) <=> ? [X7,X8,X9,X10] : (where_5(X6) & from_4(X6,X7,X8,X9,X10))) [input]
7. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
8. ! [X11] : eq(X11,X11) [input]
14. ! [X11,X12] : (lt(X11,X12) <=> (~eq(X11,X12) & lte(X11,X12))) [input]
15. ! [X11,X12] : (gt(X11,X12) <=> lt(X12,X11)) [input]
41. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
42. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
45. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
46. ! [X0] : (where_5(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [rectify 5]
47. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
48. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
49. ! [X0] : eq(X0,X0) [rectify 8]
55. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 14]
56. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 15]
82. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 42]
134. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 82]
135. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 136]
138. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
139. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 137,138]
140. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 44]
141. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 45]
142. ! [X0] : ((where_5(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [nnf transformation 46]
143. ! [X0] : ((where_5(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [flattening 142]
144. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 47]
145. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 144]
146. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
147. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 145,146]
148. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 48]
151. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 55]
152. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 151]
153. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 56]
154. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 134]
155. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
156. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 154,155]
157. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 135]
158. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 135]
159. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 139]
160. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 139]
161. ~main_query_1(X0) | select_1(X0) [cnf transformation 140]
162. ~select_1(X0) | main_query_1(X0) [cnf transformation 140]
163. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 141]
164. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 141]
167. gt(X0,X0) | where_5(X0) | lt(X0,X0) [cnf transformation 143]
168. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 147]
170. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 147]
171. ~main_query_2(X0) | select_3(X0) [cnf transformation 148]
172. ~select_3(X0) | main_query_2(X0) [cnf transformation 148]
173. eq(X0,X0) [cnf transformation 49]
182. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 152]
184. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 153]
209. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 156]
210. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 156]
235. 0 <=> main_query_1(sK8) [avatar definition]
236. main_query_1(sK8) <- {0} [avatar component clause 235]
238. 1 <=> ~main_query_1(sK8) [avatar definition]
239. ~main_query_1(sK8) <- {1} [avatar component clause 238]
241. 2 <=> main_query_2(sK8) [avatar definition]
242. main_query_2(sK8) <- {2} [avatar component clause 241]
244. 3 <=> ~main_query_2(sK8) [avatar definition]
245. ~main_query_2(sK8) <- {3} [avatar component clause 244]
246. ~1 | ~3 [avatar split clause 210,244,238]
247. 0 | 2 [avatar split clause 209,241,235]
248. select_3(sK8) <- {2} [resolution 171,242]
254. ~lt(X0,X0) [resolution 182,173]
256. where_5(X1) | lt(X1,X1) | lt(X1,X1) [resolution 167,184]
257. where_5(X1) | lt(X1,X1) [duplicate literal removal 256]
258. where_5(X1) [subsumption resolution 257,254]
263. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 159,157]
265. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 168,163]
267. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 263,164]
270. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 265,158]
272. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 267,170]
273. ~select_1(X1) | select_3(X1) [subsumption resolution 272,258]
275. ~select_3(X1) | select_1(X1) [resolution 270,160]
276. select_1(sK8) <- {2} [resolution 275,248]
278. main_query_1(sK8) <- {2} [resolution 276,162]
279. $false <- {1, 2} [subsumption resolution 278,239]
280. 1 | ~2 [avatar contradiction clause 279]
281. select_1(sK8) <- {0} [resolution 236,161]
282. select_3(sK8) <- {0} [resolution 281,273]
285. main_query_2(sK8) <- {0} [resolution 282,172]
287. $false <- {0, 3} [subsumption resolution 285,245]
288. ~0 | 3 [avatar contradiction clause 287]
289. $false [avatar sat refutation 246,247,280,288]
% SZS output end Proof for query03_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.013 s
% ------------------------------
% ------------------------------
