% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_20_query_29_tptp4X
% SZS output start Proof for query_20_query_29_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X5,X6,X7,X8,X9] : (from_6(X5,X6,X7,X8,X9) <=> tab2(X5,X6,X7,X8,X9)) [input]
3. ! [X5] : (select_5(X5) <=> ? [X6,X7,X8,X9] : ($true & from_6(X5,X6,X7,X8,X9))) [input]
4. ! [X10] : (where_3_subquery_4(X10) <=> select_5(X10)) [input]
5. ! [X0] : (where_3(X0) <=> where_3_subquery_4(X0)) [input]
6. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
7. ! [X11,X12] : (main_query_1(X11,X12) <=> select_1(X11,X12)) [input]
8. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) <=> (tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4))) [input]
9. ! [X5,X6,X7,X8,X9] : (from_12(X5,X6,X7,X8,X9) <=> tab2(X5,X6,X7,X8,X9)) [input]
10. ! [X5,X0] : (where_13(X5,X0) <=> eq(X5,X0)) [input]
11. ! [X5,X0] : (select_11(X5,X0) <=> ? [X6,X7,X8,X9] : (where_13(X5,X0) & from_12(X5,X6,X7,X8,X9))) [input]
12. ! [X10,X0] : (where_9_subquery_10(X10,X0) <=> select_11(X10,X0)) [input]
13. ! [X0] : (where_9(X0) <=> where_9_subquery_10(X0,X0)) [input]
14. ! [X0,X1] : (select_7(X0,X1) <=> ? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_9(X0) & from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [input]
15. ! [X11,X12] : (main_query_2(X11,X12) <=> select_7(X11,X12)) [input]
16. ! [X13] : eq(X13,X13) [input]
75. ! [X11,X12] : (main_query_1(X11,X12) <=> main_query_2(X11,X12)) [input]
76. ~! [X11,X12] : (main_query_1(X11,X12) <=> main_query_2(X11,X12)) [negated conjecture 75]
77. ! [X0,X1,X2,X3,X4] : (from_6(X0,X1,X2,X3,X4) <=> tab2(X0,X1,X2,X3,X4)) [rectify 2]
78. ! [X0] : (select_5(X0) <=> ? [X1,X2,X3,X4] : ($true & from_6(X0,X1,X2,X3,X4))) [rectify 3]
79. ! [X0] : (select_5(X0) <=> ? [X1,X2,X3,X4] : from_6(X0,X1,X2,X3,X4)) [true and false elimination 78]
80. ! [X0] : (where_3_subquery_4(X0) <=> select_5(X0)) [rectify 4]
81. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 7]
82. ! [X0,X1,X2,X3,X4] : (from_12(X0,X1,X2,X3,X4) <=> tab2(X0,X1,X2,X3,X4)) [rectify 9]
83. ! [X0,X1] : (where_13(X0,X1) <=> eq(X0,X1)) [rectify 10]
84. ! [X0,X1] : (select_11(X0,X1) <=> ? [X2,X3,X4,X5] : (where_13(X0,X1) & from_12(X0,X2,X3,X4,X5))) [rectify 11]
85. ! [X0,X1] : (where_9_subquery_10(X0,X1) <=> select_11(X0,X1)) [rectify 12]
86. ! [X0,X1] : (main_query_2(X0,X1) <=> select_7(X0,X1)) [rectify 15]
87. ! [X0] : eq(X0,X0) [rectify 16]
146. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 76]
250. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 146]
251. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
252. ! [X0,X1,X2,X3,X4] : ((from_6(X0,X1,X2,X3,X4) | ~tab2(X0,X1,X2,X3,X4)) & (tab2(X0,X1,X2,X3,X4) | ~from_6(X0,X1,X2,X3,X4))) [nnf transformation 77]
253. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_6(X0,X1,X2,X3,X4) | ~select_5(X0))) [nnf transformation 79]
254. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_6(X0,X5,X6,X7,X8) | ~select_5(X0))) [rectify 253]
255. ! [X0] : (? [X5,X6,X7,X8] : from_6(X0,X5,X6,X7,X8) => from_6(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
256. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (from_6(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_5(X0))) [skolemisation 254,255]
257. ! [X0] : ((where_3_subquery_4(X0) | ~select_5(X0)) & (select_5(X0) | ~where_3_subquery_4(X0))) [nnf transformation 80]
258. ! [X0] : ((where_3(X0) | ~where_3_subquery_4(X0)) & (where_3_subquery_4(X0) | ~where_3(X0))) [nnf transformation 5]
259. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 6]
260. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 259]
261. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0) & from_2(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1)))) [choice axiom]
262. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1))) | ~select_1(X0,X1))) [skolemisation 260,261]
263. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 81]
264. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | (~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4))) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [nnf transformation 8]
265. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4)) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [flattening 264]
266. ! [X0,X1,X2,X3,X4] : ((from_12(X0,X1,X2,X3,X4) | ~tab2(X0,X1,X2,X3,X4)) & (tab2(X0,X1,X2,X3,X4) | ~from_12(X0,X1,X2,X3,X4))) [nnf transformation 82]
267. ! [X0,X1] : ((where_13(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_13(X0,X1))) [nnf transformation 83]
268. ! [X0,X1] : ((select_11(X0,X1) | ! [X2,X3,X4,X5] : (~where_13(X0,X1) | ~from_12(X0,X2,X3,X4,X5))) & (? [X2,X3,X4,X5] : (where_13(X0,X1) & from_12(X0,X2,X3,X4,X5)) | ~select_11(X0,X1))) [nnf transformation 84]
269. ! [X0,X1] : ((select_11(X0,X1) | ! [X2,X3,X4,X5] : (~where_13(X0,X1) | ~from_12(X0,X2,X3,X4,X5))) & (? [X6,X7,X8,X9] : (where_13(X0,X1) & from_12(X0,X6,X7,X8,X9)) | ~select_11(X0,X1))) [rectify 268]
270. ! [X1,X0] : (? [X6,X7,X8,X9] : (where_13(X0,X1) & from_12(X0,X6,X7,X8,X9)) => (where_13(X0,X1) & from_12(X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)))) [choice axiom]
271. ! [X0,X1] : ((select_11(X0,X1) | ! [X2,X3,X4,X5] : (~where_13(X0,X1) | ~from_12(X0,X2,X3,X4,X5))) & ((where_13(X0,X1) & from_12(X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1))) | ~select_11(X0,X1))) [skolemisation 269,270]
272. ! [X0,X1] : ((where_9_subquery_10(X0,X1) | ~select_11(X0,X1)) & (select_11(X0,X1) | ~where_9_subquery_10(X0,X1))) [nnf transformation 85]
273. ! [X0] : ((where_9(X0) | ~where_9_subquery_10(X0,X0)) & (where_9_subquery_10(X0,X0) | ~where_9(X0))) [nnf transformation 13]
274. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_9(X0) | ~from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & (? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_9(X0) & from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9)) | ~select_7(X0,X1))) [nnf transformation 14]
275. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_9(X0) | ~from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_9(X0) & from_8(X0,X1,X10,X11,X12,X13,X14,X15,X16,X17)) | ~select_7(X0,X1))) [rectify 274]
276. ! [X1,X0] : (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_9(X0) & from_8(X0,X1,X10,X11,X12,X13,X14,X15,X16,X17)) => (where_9(X0) & from_8(X0,X1,sK11(X0,X1),sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1),sK16(X0,X1),sK17(X0,X1),sK18(X0,X1)))) [choice axiom]
277. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_9(X0) | ~from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & ((where_9(X0) & from_8(X0,X1,sK11(X0,X1),sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1),sK16(X0,X1),sK17(X0,X1),sK18(X0,X1))) | ~select_7(X0,X1))) [skolemisation 275,276]
278. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_7(X0,X1)) & (select_7(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 86]
281. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 250]
282. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK19,sK20) | ~main_query_1(sK19,sK20)) & (main_query_2(sK19,sK20) | main_query_1(sK19,sK20))) [choice axiom]
283. (~main_query_2(sK19,sK20) | ~main_query_1(sK19,sK20)) & (main_query_2(sK19,sK20) | main_query_1(sK19,sK20)) [skolemisation 281,282]
284. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 251]
285. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 251]
286. ~from_6(X0,X1,X2,X3,X4) | tab2(X0,X1,X2,X3,X4) [cnf transformation 252]
287. ~tab2(X0,X1,X2,X3,X4) | from_6(X0,X1,X2,X3,X4) [cnf transformation 252]
288. from_6(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_5(X0) [cnf transformation 256]
289. ~from_6(X0,X1,X2,X3,X4) | select_5(X0) [cnf transformation 256]
290. ~where_3_subquery_4(X0) | select_5(X0) [cnf transformation 257]
291. ~select_5(X0) | where_3_subquery_4(X0) [cnf transformation 257]
292. ~where_3(X0) | where_3_subquery_4(X0) [cnf transformation 258]
293. ~where_3_subquery_4(X0) | where_3(X0) [cnf transformation 258]
294. from_2(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1)) | ~select_1(X0,X1) [cnf transformation 262]
295. ~select_1(X0,X1) | where_3(X0) [cnf transformation 262]
296. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0,X1) [cnf transformation 262]
297. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 263]
298. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 263]
299. ~from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | tab1(X0,X1,X2,X3,X4) [cnf transformation 265]
301. from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4) [cnf transformation 265]
302. ~from_12(X0,X1,X2,X3,X4) | tab2(X0,X1,X2,X3,X4) [cnf transformation 266]
303. ~tab2(X0,X1,X2,X3,X4) | from_12(X0,X1,X2,X3,X4) [cnf transformation 266]
305. ~eq(X0,X1) | where_13(X0,X1) [cnf transformation 267]
306. from_12(X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) | ~select_11(X0,X1) [cnf transformation 271]
308. ~from_12(X0,X2,X3,X4,X5) | ~where_13(X0,X1) | select_11(X0,X1) [cnf transformation 271]
309. ~where_9_subquery_10(X0,X1) | select_11(X0,X1) [cnf transformation 272]
310. ~select_11(X0,X1) | where_9_subquery_10(X0,X1) [cnf transformation 272]
311. where_9_subquery_10(X0,X0) | ~where_9(X0) [cnf transformation 273]
312. ~where_9_subquery_10(X0,X0) | where_9(X0) [cnf transformation 273]
313. from_8(X0,X1,sK11(X0,X1),sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1),sK16(X0,X1),sK17(X0,X1),sK18(X0,X1)) | ~select_7(X0,X1) [cnf transformation 277]
314. ~select_7(X0,X1) | where_9(X0) [cnf transformation 277]
315. ~from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~where_9(X0) | select_7(X0,X1) [cnf transformation 277]
316. ~main_query_2(X0,X1) | select_7(X0,X1) [cnf transformation 278]
317. ~select_7(X0,X1) | main_query_2(X0,X1) [cnf transformation 278]
318. eq(X0,X0) [cnf transformation 87]
375. main_query_2(sK19,sK20) | main_query_1(sK19,sK20) [cnf transformation 283]
376. ~main_query_2(sK19,sK20) | ~main_query_1(sK19,sK20) [cnf transformation 283]
378. 0 <=> main_query_1(sK19,sK20) [avatar definition]
379. main_query_1(sK19,sK20) <- {0} [avatar component clause 378]
381. 1 <=> ~main_query_1(sK19,sK20) [avatar definition]
384. 2 <=> main_query_2(sK19,sK20) [avatar definition]
385. main_query_2(sK19,sK20) <- {2} [avatar component clause 384]
387. 3 <=> ~main_query_2(sK19,sK20) [avatar definition]
388. ~main_query_2(sK19,sK20) <- {3} [avatar component clause 387]
389. ~1 | ~3 [avatar split clause 376,387,381]
390. 0 | 2 [avatar split clause 375,384,378]
392. where_13(X0,X0) [resolution 305,318]
394. select_11(X0,X0) | ~where_9(X0) [resolution 309,311]
397. select_7(sK19,sK20) <- {2} [resolution 316,385]
398. where_9(sK19) <- {2} [resolution 397,314]
428. tab2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_5(X0) [resolution 288,286]
430. tab1(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1)) | ~select_1(X0,X1) [resolution 294,284]
432. from_12(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_5(X0) [resolution 428,303]
435. ~where_13(X2,X3) | ~select_5(X2) | select_11(X2,X3) [resolution 432,308]
438. select_11(X0,X0) | ~select_5(X0) [resolution 435,392]
442. where_9_subquery_10(X4,X4) | ~select_5(X4) [resolution 438,310]
445. ~select_5(X0) | where_9(X0) [resolution 442,312]
483. tab2(X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) | ~select_11(X0,X1) [resolution 306,302]
494. ~tab2(X20,X21,X22,X23,X24) | ~tab1(X25,X26,X27,X28,X29) | ~where_9(X25) | select_7(X25,X26) [resolution 301,315]
496. 4 <=> ! [X25,X27,X29,X26,X28] : (~tab1(X25,X26,X27,X28,X29) | select_7(X25,X26) | ~where_9(X25)) [avatar definition]
497. ~tab1(X25,X26,X27,X28,X29) | select_7(X25,X26) | ~where_9(X25) <- {4} [avatar component clause 496]
499. 6 <=> ! [X20,X22,X21,X23,X24] : ~tab2(X20,X21,X22,X23,X24) [avatar definition]
500. ~tab2(X20,X21,X22,X23,X24) <- {6} [avatar component clause 499]
501. 4 | 6 [avatar split clause 494,499,496]
504. ~select_1(X0,X1) | ~where_9(X0) | select_7(X0,X1) <- {4} [resolution 497,430]
570. tab1(X0,X1,sK11(X0,X1),sK12(X0,X1),sK13(X0,X1)) | ~select_7(X0,X1) [resolution 313,299]
583. from_2(X0,X1,sK11(X0,X1),sK12(X0,X1),sK13(X0,X1)) | ~select_7(X0,X1) [resolution 570,285]
590. from_6(X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) | ~select_11(X0,X1) [resolution 483,287]
612. ~select_7(X2,X3) | ~where_3(X2) | select_1(X2,X3) [resolution 583,296]
625. ~where_3(sK19) | select_1(sK19,sK20) <- {2} [resolution 612,397]
630. 8 <=> select_1(sK19,sK20) [avatar definition]
631. select_1(sK19,sK20) <- {8} [avatar component clause 630]
633. 10 <=> where_3(sK19) [avatar definition]
634. where_3(sK19) <- {10} [avatar component clause 633]
636. 11 <=> ~where_3(sK19) [avatar definition]
637. ~where_3(sK19) <- {11} [avatar component clause 636]
638. 8 | ~11 | ~2 [avatar split clause 625,384,636,630]
750. ~select_11(X2,X3) | select_5(X2) [resolution 590,289]
765. ~where_9(X1) | select_5(X1) [resolution 750,394]
774. select_5(sK19) <- {2} [resolution 765,398]
784. where_3_subquery_4(sK19) <- {2} [resolution 774,291]
793. where_3(sK19) <- {2} [resolution 784,293]
795. $false <- {2, 11} [subsumption resolution 793,637]
796. ~2 | 11 [avatar contradiction clause 795]
797. select_1(sK19,sK20) <- {0} [resolution 379,297]
798. 8 | ~0 [avatar split clause 797,378,630]
799. where_3_subquery_4(sK19) <- {10} [resolution 634,292]
801. select_5(sK19) <- {10} [resolution 799,290]
805. where_9(sK19) <- {10} [resolution 801,445]
814. ~where_9(sK19) | select_7(sK19,sK20) <- {4, 8} [resolution 631,504]
815. main_query_1(sK19,sK20) <- {8} [resolution 631,298]
816. where_3(sK19) <- {8} [resolution 631,295]
819. 0 | ~8 [avatar split clause 815,630,378]
820. select_7(sK19,sK20) <- {4, 8, 10} [subsumption resolution 814,805]
833. main_query_2(sK19,sK20) <- {4, 8, 10} [resolution 820,317]
835. $false <- {3, 4, 8, 10} [subsumption resolution 833,388]
836. 3 | ~4 | ~8 | ~10 [avatar contradiction clause 835]
837. ~select_5(X0) <- {6} [resolution 500,428]
845. $false <- {6, 10} [resolution 837,801]
846. ~6 | ~10 [avatar contradiction clause 845]
847. 10 | ~8 [avatar split clause 816,630,633]
848. $false [avatar sat refutation 389,390,501,638,796,798,819,836,846,847]
% SZS output end Proof for query_20_query_29_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.054 s
% ------------------------------
% ------------------------------
