% Refutation found. Thanks to Tanya!
% SZS status Theorem for query04_query09_tptp4X
% SZS output start Proof for query04_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X0,X5,X6,X7,X8] : (from_8(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
6. ! [X0] : (select_7(X0) <=> ? [X5,X6,X7,X8] : ($true & from_8(X0,X5,X6,X7,X8))) [input]
7. ! [X0] : (subquery_from_6(X0) <=> select_7(X0)) [input]
8. ! [X0] : (from_5(X0) <=> subquery_from_6(X0)) [input]
9. ! [X0] : (select_4(X0) <=> ($true & from_5(X0))) [input]
10. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [input]
11. ! [X9] : eq(X9,X9) [input]
44. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
45. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 44]
46. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
47. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
48. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 47]
49. ! [X0] : (select_4(X0) <=> from_5(X0)) [true and false elimination 9]
50. ! [X0] : eq(X0,X0) [rectify 11]
134. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 45]
135. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
136. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 137]
139. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
140. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 138,139]
141. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
142. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 46]
143. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 48]
144. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 143]
145. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
146. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0))) [skolemisation 144,145]
147. ! [X0] : ((subquery_from_6(X0) | ~select_7(X0)) & (select_7(X0) | ~subquery_from_6(X0))) [nnf transformation 7]
148. ! [X0] : ((from_5(X0) | ~subquery_from_6(X0)) & (subquery_from_6(X0) | ~from_5(X0))) [nnf transformation 8]
149. ! [X0] : ((select_4(X0) | ~from_5(X0)) & (from_5(X0) | ~select_4(X0))) [nnf transformation 49]
150. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 10]
153. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 134]
154. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
155. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 153,154]
156. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 135]
157. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 135]
159. ~eq(X0,X0) | where_3(X0) [cnf transformation 136]
160. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 140]
162. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 140]
163. ~main_query_1(X0) | select_1(X0) [cnf transformation 141]
164. ~select_1(X0) | main_query_1(X0) [cnf transformation 141]
165. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 142]
166. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 142]
167. from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [cnf transformation 146]
168. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 146]
169. ~subquery_from_6(X0) | select_7(X0) [cnf transformation 147]
170. ~select_7(X0) | subquery_from_6(X0) [cnf transformation 147]
171. ~from_5(X0) | subquery_from_6(X0) [cnf transformation 148]
172. ~subquery_from_6(X0) | from_5(X0) [cnf transformation 148]
173. ~select_4(X0) | from_5(X0) [cnf transformation 149]
174. ~from_5(X0) | select_4(X0) [cnf transformation 149]
175. ~main_query_2(X0) | select_4(X0) [cnf transformation 150]
176. ~select_4(X0) | main_query_2(X0) [cnf transformation 150]
177. eq(X0,X0) [cnf transformation 50]
208. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 155]
209. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 155]
234. 0 <=> main_query_1(sK8) [avatar definition]
235. main_query_1(sK8) <- {0} [avatar component clause 234]
237. 1 <=> ~main_query_1(sK8) [avatar definition]
238. ~main_query_1(sK8) <- {1} [avatar component clause 237]
240. 2 <=> main_query_2(sK8) [avatar definition]
241. main_query_2(sK8) <- {2} [avatar component clause 240]
243. 3 <=> ~main_query_2(sK8) [avatar definition]
244. ~main_query_2(sK8) <- {3} [avatar component clause 243]
245. ~1 | ~3 [avatar split clause 209,243,237]
246. 0 | 2 [avatar split clause 208,240,234]
247. select_4(sK8) <- {2} [resolution 175,241]
248. from_5(sK8) <- {2} [resolution 247,173]
251. subquery_from_6(sK8) <- {2} [resolution 248,171]
253. select_7(sK8) <- {2} [resolution 251,169]
255. where_3(X1) [resolution 159,177]
267. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 160,156]
269. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [resolution 167,165]
271. from_8(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 267,166]
274. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_7(X1) [resolution 269,157]
276. ~select_1(X1) | select_7(X1) [resolution 271,168]
278. ~select_7(X1) | ~where_3(X1) | select_1(X1) [resolution 274,162]
279. ~select_7(X1) | select_1(X1) [subsumption resolution 278,255]
280. select_1(sK8) <- {2} [resolution 279,253]
282. main_query_1(sK8) <- {2} [resolution 280,164]
284. $false <- {1, 2} [subsumption resolution 282,238]
285. 1 | ~2 [avatar contradiction clause 284]
286. select_1(sK8) <- {0} [resolution 235,163]
287. select_7(sK8) <- {0} [resolution 286,276]
291. subquery_from_6(sK8) <- {0} [resolution 287,170]
292. from_5(sK8) <- {0} [resolution 291,172]
294. select_4(sK8) <- {0} [resolution 292,174]
296. main_query_2(sK8) <- {0} [resolution 294,176]
298. $false <- {0, 3} [subsumption resolution 296,244]
299. ~0 | 3 [avatar contradiction clause 298]
300. $false [avatar sat refutation 245,246,285,299]
% SZS output end Proof for query04_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
