% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query07_tptp4X
% SZS output start Proof for query00_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X1,X2,X3,X4] : (from_4(X6,X1,X2,X3,X4) <=> tab1(X6,X1,X2,X3,X4)) [input]
5. ! [X6] : (where_5(X6) <=> (~gt(X6,X6) & ~lt(X6,X6))) [input]
6. ! [X6] : (select_3(X6) <=> ? [X1,X2,X3,X4] : (where_5(X6) & from_4(X6,X1,X2,X3,X4))) [input]
7. ! [X6] : (main_query_2(X6) <=> select_3(X6)) [input]
8. ! [X7] : eq(X7,X7) [input]
15. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
16. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
39. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
40. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 39]
41. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
42. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
43. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
44. ! [X0] : (where_5(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [rectify 5]
45. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
46. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
47. ! [X0] : eq(X0,X0) [rectify 8]
54. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 15]
55. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 16]
78. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 40]
124. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 78]
125. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
126. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 41]
127. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 126]
128. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
129. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 127,128]
130. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 42]
131. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 43]
132. ! [X0] : ((where_5(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [nnf transformation 44]
133. ! [X0] : ((where_5(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [flattening 132]
134. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 45]
135. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 134]
136. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
137. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 135,136]
138. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 46]
141. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 54]
142. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 141]
143. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 55]
144. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 124]
145. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
146. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 144,145]
147. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 125]
148. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 125]
149. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 129]
150. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 129]
151. ~main_query_1(X0) | select_1(X0) [cnf transformation 130]
152. ~select_1(X0) | main_query_1(X0) [cnf transformation 130]
153. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 131]
154. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 131]
157. gt(X0,X0) | where_5(X0) | lt(X0,X0) [cnf transformation 133]
158. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 137]
160. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 137]
161. ~main_query_2(X0) | select_3(X0) [cnf transformation 138]
162. ~select_3(X0) | main_query_2(X0) [cnf transformation 138]
163. eq(X0,X0) [cnf transformation 47]
173. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 142]
175. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 143]
197. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 146]
198. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 146]
200. 0 <=> main_query_1(sK8) [avatar definition]
201. main_query_1(sK8) <- {0} [avatar component clause 200]
203. 1 <=> ~main_query_1(sK8) [avatar definition]
204. ~main_query_1(sK8) <- {1} [avatar component clause 203]
206. 2 <=> main_query_2(sK8) [avatar definition]
207. main_query_2(sK8) <- {2} [avatar component clause 206]
209. 3 <=> ~main_query_2(sK8) [avatar definition]
210. ~main_query_2(sK8) <- {3} [avatar component clause 209]
211. ~1 | ~3 [avatar split clause 198,209,203]
212. 0 | 2 [avatar split clause 197,206,200]
213. select_3(sK8) <- {2} [resolution 161,207]
220. ~lt(X0,X0) [resolution 173,163]
222. where_5(X1) | lt(X1,X1) | lt(X1,X1) [resolution 157,175]
223. where_5(X1) | lt(X1,X1) [duplicate literal removal 222]
224. where_5(X1) [subsumption resolution 223,220]
248. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 149,147]
290. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 158,153]
414. from_4(X2,sK0(X2),sK1(X2),sK2(X2),sK3(X2)) | ~select_1(X2) [resolution 248,154]
420. from_2(X5,sK4(X5),sK5(X5),sK6(X5),sK7(X5)) | ~select_3(X5) [resolution 290,148]
424. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 414,160]
430. ~select_1(X1) | select_3(X1) [subsumption resolution 424,224]
436. ~select_3(X1) | select_1(X1) [resolution 420,150]
442. select_1(sK8) <- {2} [resolution 436,213]
444. main_query_1(sK8) <- {2} [resolution 442,152]
445. $false <- {1, 2} [subsumption resolution 444,204]
446. 1 | ~2 [avatar contradiction clause 445]
447. select_1(sK8) <- {0} [resolution 201,151]
448. select_3(sK8) <- {0} [resolution 447,430]
451. main_query_2(sK8) <- {0} [resolution 448,162]
453. $false <- {0, 3} [subsumption resolution 451,210]
454. ~0 | 3 [avatar contradiction clause 453]
455. $false [avatar sat refutation 211,212,446,454]
% SZS output end Proof for query00_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.012 s
% ------------------------------
% ------------------------------
