% Refutation found. Thanks to Tanya!
% SZS status Theorem for query11_query13_tptp4X
% SZS output start Proof for query11_query13_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8] : (where_6(X7,X8) <=> (~gt(X7,X8) & ~lt(X7,X8))) [input]
7. ! [X7,X8] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X7,X8) & from_5(X7,X8,X9,X10,X11))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
9. ! [X12] : eq(X12,X12) [input]
13. ! [X12,X13] : ((lte(X13,X12) & lte(X12,X13)) <=> eq(X12,X13)) [input]
15. ! [X12,X13] : (lte(X13,X12) | lte(X12,X13)) [input]
16. ! [X12,X13] : (lt(X12,X13) <=> (~eq(X12,X13) & lte(X12,X13))) [input]
17. ! [X12,X13] : (gt(X12,X13) <=> lt(X13,X12)) [input]
46. ! [X12,X7,X8] : ((where_6(X7,X8) & eq(X7,X12)) => where_6(X12,X8)) [input]
55. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
56. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 55]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
58. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
59. ! [X0,X1] : (where_6(X0,X1) <=> (~gt(X0,X1) & ~lt(X0,X1))) [rectify 6]
60. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
61. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
62. ! [X0] : eq(X0,X0) [rectify 9]
66. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
68. ! [X0,X1] : (lte(X1,X0) | lte(X0,X1)) [rectify 15]
69. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 16]
70. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 17]
99. ! [X0,X1,X2] : ((where_6(X1,X2) & eq(X1,X0)) => where_6(X0,X2)) [rectify 46]
108. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 56]
166. ! [X0,X1,X2] : (where_6(X0,X2) | (~where_6(X1,X2) | ~eq(X1,X0))) [ennf transformation 99]
167. ! [X0,X1,X2] : (where_6(X0,X2) | ~where_6(X1,X2) | ~eq(X1,X0)) [flattening 166]
184. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 108]
185. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
186. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
187. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
188. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 187]
189. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
190. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 188,189]
191. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
192. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 58]
193. ! [X0,X1] : ((where_6(X0,X1) | (gt(X0,X1) | lt(X0,X1))) & ((~gt(X0,X1) & ~lt(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 59]
194. ! [X0,X1] : ((where_6(X0,X1) | gt(X0,X1) | lt(X0,X1)) & ((~gt(X0,X1) & ~lt(X0,X1)) | ~where_6(X0,X1))) [flattening 193]
195. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 60]
196. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 195]
197. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
198. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 196,197]
199. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 61]
200. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 66]
201. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 200]
202. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 69]
203. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 202]
204. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 70]
205. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 184]
206. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
207. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 205,206]
208. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 185]
209. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 185]
210. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 186]
211. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 186]
212. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 190]
213. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 190]
214. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 190]
215. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 191]
216. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 191]
217. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 192]
218. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 192]
219. ~lt(X0,X1) | ~where_6(X0,X1) [cnf transformation 194]
220. ~gt(X0,X1) | ~where_6(X0,X1) [cnf transformation 194]
221. gt(X0,X1) | where_6(X0,X1) | lt(X0,X1) [cnf transformation 194]
222. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 198]
223. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 198]
224. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 198]
225. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 199]
226. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 199]
227. eq(X0,X0) [cnf transformation 62]
233. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 201]
235. lte(X1,X0) | lte(X0,X1) [cnf transformation 68]
237. ~lt(X0,X1) | ~eq(X0,X1) [cnf transformation 203]
238. ~lte(X0,X1) | eq(X0,X1) | lt(X0,X1) [cnf transformation 203]
239. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 204]
240. ~lt(X1,X0) | gt(X0,X1) [cnf transformation 204]
267. ~where_6(X1,X2) | where_6(X0,X2) | ~eq(X1,X0) [cnf transformation 167]
276. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 207]
277. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 207]
279. 0 <=> main_query_1(sK6,sK7) [avatar definition]
280. main_query_1(sK6,sK7) <- {0} [avatar component clause 279]
282. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
283. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 282]
285. 2 <=> main_query_2(sK6,sK7) [avatar definition]
286. main_query_2(sK6,sK7) <- {2} [avatar component clause 285]
288. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
289. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 288]
290. ~1 | ~3 [avatar split clause 277,288,282]
291. 0 | 2 [avatar split clause 276,285,279]
294. select_4(sK6,sK7) <- {2} [resolution 225,286]
295. where_6(sK6,sK7) <- {2} [resolution 294,223]
301. lt(X1,X0) | lt(X0,X1) | where_6(X0,X1) [resolution 221,239]
311. eq(X1,X2) | lt(X1,X2) | lte(X2,X1) [resolution 238,235]
313. lte(X2,X1) | lt(X1,X2) [subsumption resolution 311,233]
314. lt(X1,X0) | lt(X0,X1) | eq(X1,X0) [resolution 313,238]
336. ~eq(X3,X2) | where_6(X3,X2) | lt(X2,X3) [resolution 301,237]
339. lt(X0,X0) | where_6(X0,X0) [factoring 301]
341. where_6(X1,X1) | ~eq(X1,X1) [resolution 339,237]
344. where_6(X1,X1) [subsumption resolution 341,227]
346. ~eq(X3,X2) | where_6(X2,X3) [resolution 344,267]
369. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 212,208]
372. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 222,217]
376. 5 <=> ~lt(sK7,sK6) [avatar definition]
377. ~lt(sK7,sK6) <- {5} [avatar component clause 376]
379. 4 <=> lt(sK7,sK6) [avatar definition]
380. lt(sK7,sK6) <- {4} [avatar component clause 379]
382. 7 <=> ~where_6(sK7,sK6) [avatar definition]
385. 6 <=> where_6(sK7,sK6) [avatar definition]
389. gt(sK6,sK7) <- {4} [resolution 380,240]
392. ~where_6(sK7,sK6) <- {4} [resolution 380,219]
393. ~7 | ~4 [avatar split clause 392,379,382]
397. ~where_6(sK6,sK7) <- {4} [resolution 389,220]
398. $false <- {2, 4} [subsumption resolution 397,295]
399. ~2 | ~4 [avatar contradiction clause 398]
401. lt(sK6,sK7) | eq(sK6,sK7) <- {5} [resolution 377,314]
406. 9 <=> ~eq(sK6,sK7) [avatar definition]
407. ~eq(sK6,sK7) <- {9} [avatar component clause 406]
409. 8 <=> eq(sK6,sK7) [avatar definition]
410. eq(sK6,sK7) <- {8} [avatar component clause 409]
415. 10 <=> lt(sK6,sK7) [avatar definition]
416. lt(sK6,sK7) <- {10} [avatar component clause 415]
417. 8 | 10 | 5 [avatar split clause 401,376,415,409]
432. ~where_6(sK6,sK7) <- {10} [resolution 416,219]
434. $false <- {2, 10} [subsumption resolution 432,295]
435. ~2 | ~10 [avatar contradiction clause 434]
439. where_3(sK6,sK7) <- {8} [resolution 410,211]
448. where_6(sK7,sK6) <- {8} [resolution 410,346]
504. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 369,218]
512. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 372,209]
624. where_6(sK6,sK7) | lt(sK7,sK6) <- {8} [resolution 336,410]
694. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 504,224]
713. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 512,214]
962. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 713,294]
966. select_1(sK6,sK7) <- {2, 8} [subsumption resolution 962,439]
981. main_query_1(sK6,sK7) <- {2, 8} [resolution 966,216]
983. $false <- {1, 2, 8} [subsumption resolution 981,283]
984. 1 | ~2 | ~8 [avatar contradiction clause 983]
998. where_6(sK6,sK7) <- {5, 8} [subsumption resolution 624,377]
1001. select_1(sK6,sK7) <- {0} [resolution 280,215]
1007. where_3(sK6,sK7) <- {0} [resolution 1001,213]
1017. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {5, 8} [resolution 998,694]
1020. select_4(sK6,sK7) <- {0, 5, 8} [subsumption resolution 1017,1001]
1036. main_query_2(sK6,sK7) <- {0, 5, 8} [resolution 1020,226]
1038. $false <- {0, 3, 5, 8} [subsumption resolution 1036,289]
1039. ~0 | 3 | 5 | ~8 [avatar contradiction clause 1038]
1044. 6 | ~8 [avatar split clause 448,409,385]
1046. eq(sK6,sK7) <- {0} [resolution 1007,210]
1047. $false <- {0, 9} [subsumption resolution 1046,407]
1048. ~0 | 9 [avatar contradiction clause 1047]
1049. $false [avatar sat refutation 290,291,393,399,417,435,984,1039,1044,1048]
% SZS output end Proof for query11_query13_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.174 s
% ------------------------------
% ------------------------------
