% Refutation found. Thanks to Tanya!
% SZS status Theorem for query10_query14_tptp4X
% SZS output start Proof for query10_query14_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (where_6_subquery_7(X5) <=> select_8(X5)) [input]
9. ! [X0,X1] : (where_6(X0,X1) <=> (where_6_subquery_7(X0) & eq(X0,X1))) [input]
10. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
63. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
64. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 63]
65. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
66. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 7]
67. ! [X0] : (where_6_subquery_7(X0) <=> select_8(X0)) [rectify 8]
68. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
120. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 64]
206. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 120]
207. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
208. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
209. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
210. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 209]
211. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
212. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 210,211]
213. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 65]
214. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
215. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
216. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 66]
217. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 216]
218. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
219. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 217,218]
220. ! [X0] : ((where_6_subquery_7(X0) | ~select_8(X0)) & (select_8(X0) | ~where_6_subquery_7(X0))) [nnf transformation 67]
221. ! [X0,X1] : ((where_6(X0,X1) | (~where_6_subquery_7(X0) | ~eq(X0,X1))) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 9]
222. ! [X0,X1] : ((where_6(X0,X1) | ~where_6_subquery_7(X0) | ~eq(X0,X1)) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 221]
223. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 10]
224. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 223]
225. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
226. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 224,225]
227. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 68]
230. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 206]
231. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
232. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 230,231]
233. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 207]
234. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 207]
235. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 208]
236. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 208]
237. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 212]
238. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 212]
239. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 212]
240. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 213]
241. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 213]
242. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 214]
243. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 214]
245. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 215]
247. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 219]
249. ~select_8(X0) | where_6_subquery_7(X0) [cnf transformation 220]
250. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 222]
252. ~eq(X0,X1) | ~where_6_subquery_7(X0) | where_6(X0,X1) [cnf transformation 222]
253. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 226]
254. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 226]
255. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 226]
256. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 227]
257. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 227]
307. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 232]
308. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 232]
310. 0 <=> main_query_1(sK10,sK11) [avatar definition]
311. main_query_1(sK10,sK11) <- {0} [avatar component clause 310]
313. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
314. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 313]
316. 2 <=> main_query_2(sK10,sK11) [avatar definition]
317. main_query_2(sK10,sK11) <- {2} [avatar component clause 316]
319. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
320. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 319]
321. ~1 | ~3 [avatar split clause 308,319,313]
322. 0 | 2 [avatar split clause 307,316,310]
406. select_1(sK10,sK11) <- {0} [resolution 311,240]
410. where_3(sK10,sK11) <- {0} [resolution 406,238]
411. eq(sK10,sK11) <- {0} [resolution 410,235]
414. ~where_6_subquery_7(sK10) | where_6(sK10,sK11) <- {0} [resolution 411,252]
424. 8 <=> where_6(sK10,sK11) [avatar definition]
425. where_6(sK10,sK11) <- {8} [avatar component clause 424]
427. 10 <=> where_6_subquery_7(sK10) [avatar definition]
430. 11 <=> ~where_6_subquery_7(sK10) [avatar definition]
432. 8 | ~11 | ~0 [avatar split clause 414,310,430,424]
472. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 237,233]
476. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 253,242]
601. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 472,243]
602. from_9(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 472,245]
621. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 476,234]
657. 14 <=> select_8(sK10) [avatar definition]
658. select_8(sK10) <- {14} [avatar component clause 657]
1189. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 601,255]
1197. ~select_1(X2,X3) | select_8(X2) [resolution 602,247]
1203. select_8(sK10) <- {0} [resolution 1197,406]
1218. 14 | ~0 [avatar split clause 1203,310,657]
1220. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 621,239]
1238. where_6_subquery_7(sK10) <- {14} [resolution 658,249]
1241. 10 | ~14 [avatar split clause 1238,657,427]
1268. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {8} [resolution 425,1189]
1273. select_4(sK10,sK11) <- {0, 8} [subsumption resolution 1268,406]
1298. main_query_2(sK10,sK11) <- {0, 8} [resolution 1273,257]
1300. $false <- {0, 3, 8} [subsumption resolution 1298,320]
1301. ~0 | 3 | ~8 [avatar contradiction clause 1300]
1317. select_4(sK10,sK11) <- {2} [resolution 317,256]
1321. where_6(sK10,sK11) <- {2} [resolution 1317,254]
1322. 8 | ~2 [avatar split clause 1321,316,424]
1326. eq(sK10,sK11) <- {8} [resolution 425,250]
1329. where_3(sK10,sK11) <- {8} [resolution 1326,236]
1747. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 1220,1317]
1751. select_1(sK10,sK11) <- {2, 8} [subsumption resolution 1747,1329]
1760. main_query_1(sK10,sK11) <- {2, 8} [resolution 1751,241]
1766. $false <- {1, 2, 8} [subsumption resolution 1760,314]
1767. 1 | ~2 | ~8 [avatar contradiction clause 1766]
1768. $false [avatar sat refutation 321,322,432,1218,1241,1301,1322,1767]
% SZS output end Proof for query10_query14_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5500
% Time elapsed: 0.082 s
% ------------------------------
% ------------------------------
