% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_16_query_28_tptp4X
% SZS output start Proof for query_16_query_28_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_6(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (subquery_from_3(X5,X6) <=> select_4(X5,X6)) [input]
5. ! [X7,X8] : (from_2(X7,X8) <=> subquery_from_3(X7,X8)) [input]
6. ! [X7,X8] : (select_1(X7,X8) <=> ($true & from_2(X7,X8))) [input]
7. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
8. ! [X0,X9,X2,X3,X4] : (from_8(X0,X9,X2,X3,X4) <=> tab1(X0,X9,X2,X3,X4)) [input]
9. ! [X0,X1,X2,X3,X4] : (from_12(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
10. ! [X0,X9] : (where_13(X0,X9) <=> eq(X0,X9)) [input]
11. ! [X0,X9] : (select_11(X0,X9) <=> ? [X1,X2,X3,X4] : (where_13(X0,X9) & from_12(X0,X1,X2,X3,X4))) [input]
12. ! [X5,X9] : (where_9_subquery_10(X5,X9) <=> select_11(X5,X9)) [input]
13. ! [X0,X9] : (where_9(X0,X9) <=> where_9_subquery_10(X0,X9)) [input]
14. ! [X0,X9] : (select_7(X0,X9) <=> ? [X2,X3,X4] : (where_9(X0,X9) & from_8(X0,X9,X2,X3,X4))) [input]
15. ! [X5,X9] : (main_query_2(X5,X9) <=> select_7(X5,X9)) [input]
75. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
76. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 75]
77. ! [X0,X1] : (subquery_from_3(X0,X1) <=> select_4(X0,X1)) [rectify 4]
78. ! [X0,X1] : (from_2(X0,X1) <=> subquery_from_3(X0,X1)) [rectify 5]
79. ! [X0,X1] : (select_1(X0,X1) <=> ($true & from_2(X0,X1))) [rectify 6]
80. ! [X0,X1] : (select_1(X0,X1) <=> from_2(X0,X1)) [true and false elimination 79]
81. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 7]
82. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 8]
83. ! [X0,X1] : (where_13(X0,X1) <=> eq(X0,X1)) [rectify 10]
84. ! [X0,X1] : (select_11(X0,X1) <=> ? [X2,X3,X4,X5] : (where_13(X0,X1) & from_12(X0,X2,X3,X4,X5))) [rectify 11]
85. ! [X0,X1] : (where_9_subquery_10(X0,X1) <=> select_11(X0,X1)) [rectify 12]
86. ! [X0,X1] : (where_9(X0,X1) <=> where_9_subquery_10(X0,X1)) [rectify 13]
87. ! [X0,X1] : (select_7(X0,X1) <=> ? [X2,X3,X4] : (where_9(X0,X1) & from_8(X0,X1,X2,X3,X4))) [rectify 14]
88. ! [X0,X1] : (main_query_2(X0,X1) <=> select_7(X0,X1)) [rectify 15]
148. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 76]
252. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 148]
253. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 1]
254. ! [X0,X1] : ((where_6(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_6(X0,X1))) [nnf transformation 2]
255. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 3]
256. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 255]
257. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
258. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_4(X0,X1))) [skolemisation 256,257]
259. ! [X0,X1] : ((subquery_from_3(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~subquery_from_3(X0,X1))) [nnf transformation 77]
260. ! [X0,X1] : ((from_2(X0,X1) | ~subquery_from_3(X0,X1)) & (subquery_from_3(X0,X1) | ~from_2(X0,X1))) [nnf transformation 78]
261. ! [X0,X1] : ((select_1(X0,X1) | ~from_2(X0,X1)) & (from_2(X0,X1) | ~select_1(X0,X1))) [nnf transformation 80]
262. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 81]
263. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 82]
264. ! [X0,X1,X2,X3,X4] : ((from_12(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_12(X0,X1,X2,X3,X4))) [nnf transformation 9]
265. ! [X0,X1] : ((where_13(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_13(X0,X1))) [nnf transformation 83]
266. ! [X0,X1] : ((select_11(X0,X1) | ! [X2,X3,X4,X5] : (~where_13(X0,X1) | ~from_12(X0,X2,X3,X4,X5))) & (? [X2,X3,X4,X5] : (where_13(X0,X1) & from_12(X0,X2,X3,X4,X5)) | ~select_11(X0,X1))) [nnf transformation 84]
267. ! [X0,X1] : ((select_11(X0,X1) | ! [X2,X3,X4,X5] : (~where_13(X0,X1) | ~from_12(X0,X2,X3,X4,X5))) & (? [X6,X7,X8,X9] : (where_13(X0,X1) & from_12(X0,X6,X7,X8,X9)) | ~select_11(X0,X1))) [rectify 266]
268. ! [X1,X0] : (? [X6,X7,X8,X9] : (where_13(X0,X1) & from_12(X0,X6,X7,X8,X9)) => (where_13(X0,X1) & from_12(X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1),sK6(X0,X1)))) [choice axiom]
269. ! [X0,X1] : ((select_11(X0,X1) | ! [X2,X3,X4,X5] : (~where_13(X0,X1) | ~from_12(X0,X2,X3,X4,X5))) & ((where_13(X0,X1) & from_12(X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1),sK6(X0,X1))) | ~select_11(X0,X1))) [skolemisation 267,268]
270. ! [X0,X1] : ((where_9_subquery_10(X0,X1) | ~select_11(X0,X1)) & (select_11(X0,X1) | ~where_9_subquery_10(X0,X1))) [nnf transformation 85]
271. ! [X0,X1] : ((where_9(X0,X1) | ~where_9_subquery_10(X0,X1)) & (where_9_subquery_10(X0,X1) | ~where_9(X0,X1))) [nnf transformation 86]
272. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : (~where_9(X0,X1) | ~from_8(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_9(X0,X1) & from_8(X0,X1,X2,X3,X4)) | ~select_7(X0,X1))) [nnf transformation 87]
273. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : (~where_9(X0,X1) | ~from_8(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_9(X0,X1) & from_8(X0,X1,X5,X6,X7)) | ~select_7(X0,X1))) [rectify 272]
274. ! [X1,X0] : (? [X5,X6,X7] : (where_9(X0,X1) & from_8(X0,X1,X5,X6,X7)) => (where_9(X0,X1) & from_8(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
275. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : (~where_9(X0,X1) | ~from_8(X0,X1,X2,X3,X4))) & ((where_9(X0,X1) & from_8(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_7(X0,X1))) [skolemisation 273,274]
276. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_7(X0,X1)) & (select_7(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 88]
279. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 252]
280. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
281. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 279,280]
282. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 253]
283. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 253]
284. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 254]
285. ~eq(X0,X1) | where_6(X0,X1) [cnf transformation 254]
286. from_5(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_4(X0,X1) [cnf transformation 258]
287. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 258]
288. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 258]
289. ~subquery_from_3(X0,X1) | select_4(X0,X1) [cnf transformation 259]
290. ~select_4(X0,X1) | subquery_from_3(X0,X1) [cnf transformation 259]
291. ~from_2(X0,X1) | subquery_from_3(X0,X1) [cnf transformation 260]
292. ~subquery_from_3(X0,X1) | from_2(X0,X1) [cnf transformation 260]
293. ~select_1(X0,X1) | from_2(X0,X1) [cnf transformation 261]
294. ~from_2(X0,X1) | select_1(X0,X1) [cnf transformation 261]
295. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 262]
296. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 262]
297. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 263]
298. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 263]
300. ~tab1(X0,X1,X2,X3,X4) | from_12(X0,X1,X2,X3,X4) [cnf transformation 264]
301. ~where_13(X0,X1) | eq(X0,X1) [cnf transformation 265]
302. ~eq(X0,X1) | where_13(X0,X1) [cnf transformation 265]
304. ~select_11(X0,X1) | where_13(X0,X1) [cnf transformation 269]
305. ~from_12(X0,X2,X3,X4,X5) | ~where_13(X0,X1) | select_11(X0,X1) [cnf transformation 269]
306. ~where_9_subquery_10(X0,X1) | select_11(X0,X1) [cnf transformation 270]
307. ~select_11(X0,X1) | where_9_subquery_10(X0,X1) [cnf transformation 270]
308. ~where_9(X0,X1) | where_9_subquery_10(X0,X1) [cnf transformation 271]
309. ~where_9_subquery_10(X0,X1) | where_9(X0,X1) [cnf transformation 271]
310. from_8(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_7(X0,X1) [cnf transformation 275]
311. ~select_7(X0,X1) | where_9(X0,X1) [cnf transformation 275]
312. ~from_8(X0,X1,X2,X3,X4) | ~where_9(X0,X1) | select_7(X0,X1) [cnf transformation 275]
313. ~main_query_2(X0,X1) | select_7(X0,X1) [cnf transformation 276]
314. ~select_7(X0,X1) | main_query_2(X0,X1) [cnf transformation 276]
372. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 281]
373. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 281]
375. 0 <=> main_query_1(sK10,sK11) [avatar definition]
376. main_query_1(sK10,sK11) <- {0} [avatar component clause 375]
378. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
379. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 378]
381. 2 <=> main_query_2(sK10,sK11) [avatar definition]
382. main_query_2(sK10,sK11) <- {2} [avatar component clause 381]
384. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
385. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 384]
386. ~1 | ~3 [avatar split clause 373,384,378]
387. 0 | 2 [avatar split clause 372,381,375]
392. select_7(sK10,sK11) <- {2} [resolution 313,382]
393. where_9(sK10,sK11) <- {2} [resolution 392,311]
395. where_9_subquery_10(sK10,sK11) <- {2} [resolution 393,308]
398. select_11(sK10,sK11) <- {2} [resolution 395,306]
401. where_13(sK10,sK11) <- {2} [resolution 398,304]
403. eq(sK10,sK11) <- {2} [resolution 401,301]
410. where_6(sK10,sK11) <- {2} [resolution 403,285]
524. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_4(X2,X3) [resolution 286,282]
528. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_7(X2,X3) [resolution 310,297]
615. from_8(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_4(X2,X3) [resolution 524,298]
616. from_12(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_4(X4,X5) [resolution 524,300]
624. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_7(X0,X1) [resolution 528,283]
733. ~where_9(X2,X3) | ~select_4(X2,X3) | select_7(X2,X3) [resolution 615,312]
791. ~where_13(X2,X4) | ~select_4(X2,X3) | select_11(X2,X4) [resolution 616,305]
800. ~select_7(X2,X3) | ~where_6(X2,X3) | select_4(X2,X3) [resolution 624,288]
823. ~where_6(sK10,sK11) | select_4(sK10,sK11) <- {2} [resolution 800,392]
827. select_4(sK10,sK11) <- {2} [subsumption resolution 823,410]
841. subquery_from_3(sK10,sK11) <- {2} [resolution 827,290]
853. from_2(sK10,sK11) <- {2} [resolution 841,292]
865. select_1(sK10,sK11) <- {2} [resolution 853,294]
877. main_query_1(sK10,sK11) <- {2} [resolution 865,296]
879. $false <- {1, 2} [subsumption resolution 877,379]
880. 1 | ~2 [avatar contradiction clause 879]
883. select_1(sK10,sK11) <- {0} [resolution 376,295]
887. from_2(sK10,sK11) <- {0} [resolution 883,293]
891. subquery_from_3(sK10,sK11) <- {0} [resolution 887,291]
896. select_4(sK10,sK11) <- {0} [resolution 891,289]
902. where_6(sK10,sK11) <- {0} [resolution 896,287]
911. eq(sK10,sK11) <- {0} [resolution 902,284]
935. where_13(sK10,sK11) <- {0} [resolution 911,302]
999. ~select_4(sK10,X0) | select_11(sK10,sK11) <- {0} [resolution 935,791]
1007. 8 <=> select_11(sK10,sK11) [avatar definition]
1008. select_11(sK10,sK11) <- {8} [avatar component clause 1007]
1010. 10 <=> ! [X0] : ~select_4(sK10,X0) [avatar definition]
1011. ~select_4(sK10,X0) <- {10} [avatar component clause 1010]
1012. 8 | 10 | ~0 [avatar split clause 999,375,1010,1007]
1098. where_9_subquery_10(sK10,sK11) <- {8} [resolution 1008,307]
1155. where_9(sK10,sK11) <- {8} [resolution 1098,309]
1192. ~select_4(sK10,sK11) | select_7(sK10,sK11) <- {8} [resolution 1155,733]
1196. select_7(sK10,sK11) <- {0, 8} [subsumption resolution 1192,896]
1276. main_query_2(sK10,sK11) <- {0, 8} [resolution 1196,314]
1282. $false <- {0, 3, 8} [subsumption resolution 1276,385]
1283. ~0 | 3 | ~8 [avatar contradiction clause 1282]
1298. $false <- {0, 10} [resolution 1011,896]
1301. ~0 | ~10 [avatar contradiction clause 1298]
1302. $false [avatar sat refutation 386,387,880,1012,1283,1301]
% SZS output end Proof for query_16_query_28_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.037 s
% ------------------------------
% ------------------------------
