% Refutation found. Thanks to Tanya!
% SZS status Theorem for query01_query04_tptp4X
% SZS output start Proof for query01_query04_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [input]
6. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
8. ! [X5] : eq(X5,X5) [input]
41. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
42. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0] : eq(X0,X0) [rectify 8]
128. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 42]
129. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
130. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
131. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 130]
132. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
133. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 131,132]
134. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
135. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
136. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 5]
137. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 6]
138. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 137]
139. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
140. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 138,139]
141. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 7]
144. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 128]
145. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
146. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 144,145]
147. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 129]
148. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 129]
149. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 133]
150. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 133]
151. ~main_query_1(X0) | select_1(X0) [cnf transformation 134]
152. ~select_1(X0) | main_query_1(X0) [cnf transformation 134]
153. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 135]
154. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 135]
156. ~eq(X0,X0) | where_5(X0) [cnf transformation 136]
157. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 140]
159. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 140]
160. ~main_query_2(X0) | select_3(X0) [cnf transformation 141]
161. ~select_3(X0) | main_query_2(X0) [cnf transformation 141]
162. eq(X0,X0) [cnf transformation 44]
193. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 146]
194. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 146]
219. 0 <=> main_query_1(sK8) [avatar definition]
220. main_query_1(sK8) <- {0} [avatar component clause 219]
222. 1 <=> ~main_query_1(sK8) [avatar definition]
223. ~main_query_1(sK8) <- {1} [avatar component clause 222]
225. 2 <=> main_query_2(sK8) [avatar definition]
226. main_query_2(sK8) <- {2} [avatar component clause 225]
228. 3 <=> ~main_query_2(sK8) [avatar definition]
229. ~main_query_2(sK8) <- {3} [avatar component clause 228]
230. ~1 | ~3 [avatar split clause 194,228,222]
231. 0 | 2 [avatar split clause 193,225,219]
232. select_3(sK8) <- {2} [resolution 160,226]
236. where_5(X1) [resolution 156,162]
247. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 149,147]
249. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 157,153]
251. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 247,154]
254. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 249,148]
256. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 251,159]
257. ~select_1(X1) | select_3(X1) [subsumption resolution 256,236]
259. ~select_3(X1) | select_1(X1) [resolution 254,150]
260. select_1(sK8) <- {2} [resolution 259,232]
262. main_query_1(sK8) <- {2} [resolution 260,152]
263. $false <- {1, 2} [subsumption resolution 262,223]
264. 1 | ~2 [avatar contradiction clause 263]
265. select_1(sK8) <- {0} [resolution 220,151]
266. select_3(sK8) <- {0} [resolution 265,257]
269. main_query_2(sK8) <- {0} [resolution 266,161]
271. $false <- {0, 3} [subsumption resolution 269,229]
272. ~0 | 3 [avatar contradiction clause 271]
273. $false [avatar sat refutation 230,231,264,272]
% SZS output end Proof for query01_query04_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
