% Refutation found. Thanks to Tanya!
% SZS status Theorem for query10_query14_tptp4X
% SZS output start Proof for query10_query14_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X1,X0] : (where_3(X1,X0) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X1,X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (query_7(X5) <=> select_8(X5)) [input]
9. ! [X1,X0] : (where_6(X1,X0) <=> (query_7(X0) & eq(X0,X1))) [input]
10. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X1,X0) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
64. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
65. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 64]
66. ! [X0,X1] : (where_3(X0,X1) <=> eq(X1,X0)) [rectify 2]
67. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
68. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 7]
69. ! [X0] : (query_7(X0) <=> select_8(X0)) [rectify 8]
70. ! [X0,X1] : (where_6(X0,X1) <=> (query_7(X1) & eq(X1,X0))) [rectify 9]
71. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
124. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 65]
214. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 124]
215. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
216. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X1,X0)) & (eq(X1,X0) | ~where_3(X0,X1))) [nnf transformation 66]
217. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X1,X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X1,X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
218. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X1,X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X1,X0) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 217]
219. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X1,X0) & from_2(X0,X1,X5,X6,X7)) => (where_3(X1,X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
220. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X1,X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X1,X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 218,219]
221. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 67]
222. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
223. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
224. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 68]
225. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 224]
226. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
227. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 225,226]
228. ! [X0] : ((query_7(X0) | ~select_8(X0)) & (select_8(X0) | ~query_7(X0))) [nnf transformation 69]
229. ! [X0,X1] : ((where_6(X0,X1) | (~query_7(X1) | ~eq(X1,X0))) & ((query_7(X1) & eq(X1,X0)) | ~where_6(X0,X1))) [nnf transformation 70]
230. ! [X0,X1] : ((where_6(X0,X1) | ~query_7(X1) | ~eq(X1,X0)) & ((query_7(X1) & eq(X1,X0)) | ~where_6(X0,X1))) [flattening 229]
231. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X1,X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X1,X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 10]
232. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X1,X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X1,X0) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 231]
233. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X1,X0) & from_5(X0,X1,X5,X6,X7)) => (where_6(X1,X0) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
234. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X1,X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X1,X0) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 232,233]
235. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 71]
238. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 214]
239. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
240. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 238,239]
241. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 215]
242. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 215]
243. ~where_3(X0,X1) | eq(X1,X0) [cnf transformation 216]
244. ~eq(X1,X0) | where_3(X0,X1) [cnf transformation 216]
245. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 220]
246. ~select_1(X0,X1) | where_3(X1,X0) [cnf transformation 220]
247. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X1,X0) | select_1(X0,X1) [cnf transformation 220]
248. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 221]
249. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 221]
250. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 222]
251. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 222]
253. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 223]
255. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 227]
257. ~select_8(X0) | query_7(X0) [cnf transformation 228]
258. ~where_6(X0,X1) | eq(X1,X0) [cnf transformation 230]
260. ~eq(X1,X0) | ~query_7(X1) | where_6(X0,X1) [cnf transformation 230]
261. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 234]
262. ~select_4(X0,X1) | where_6(X1,X0) [cnf transformation 234]
263. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X1,X0) | select_4(X0,X1) [cnf transformation 234]
264. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 235]
265. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 235]
316. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 240]
317. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 240]
361. 0 <=> main_query_1(sK10,sK11) [avatar definition]
362. main_query_1(sK10,sK11) <- {0} [avatar component clause 361]
364. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
365. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 364]
367. 2 <=> main_query_2(sK10,sK11) [avatar definition]
368. main_query_2(sK10,sK11) <- {2} [avatar component clause 367]
370. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
371. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 370]
372. ~1 | ~3 [avatar split clause 317,370,364]
373. 0 | 2 [avatar split clause 316,367,361]
376. select_4(sK10,sK11) <- {2} [resolution 264,368]
377. where_6(sK11,sK10) <- {2} [resolution 376,262]
379. eq(sK10,sK11) <- {2} [resolution 377,258]
389. where_3(sK11,sK10) <- {2} [resolution 379,244]
437. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 245,241]
441. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 261,250]
442. from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 437,253]
443. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 437,251]
449. from_2(X4,X5,sK7(X4,X5),sK8(X4,X5),sK9(X4,X5)) | ~select_4(X4,X5) [resolution 441,242]
451. ~select_1(X2,X3) | select_8(X2) [resolution 442,255]
452. ~where_6(X1,X0) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 443,263]
459. ~select_4(X0,X1) | ~where_3(X1,X0) | select_1(X0,X1) [resolution 449,247]
461. ~where_3(sK11,sK10) | select_1(sK10,sK11) <- {2} [resolution 459,376]
462. select_1(sK10,sK11) <- {2} [subsumption resolution 461,389]
464. main_query_1(sK10,sK11) <- {2} [resolution 462,249]
466. $false <- {1, 2} [subsumption resolution 464,365]
467. 1 | ~2 [avatar contradiction clause 466]
468. select_1(sK10,sK11) <- {0} [resolution 362,248]
469. select_8(sK10) <- {0} [resolution 468,451]
471. where_3(sK11,sK10) <- {0} [resolution 468,246]
472. query_7(sK10) <- {0} [resolution 469,257]
473. eq(sK10,sK11) <- {0} [resolution 471,243]
476. ~query_7(sK10) | where_6(sK11,sK10) <- {0} [resolution 473,260]
481. where_6(sK11,sK10) <- {0} [subsumption resolution 476,472]
492. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {0} [resolution 481,452]
495. select_4(sK10,sK11) <- {0} [subsumption resolution 492,468]
501. main_query_2(sK10,sK11) <- {0} [resolution 495,265]
503. $false <- {0, 3} [subsumption resolution 501,371]
504. ~0 | 3 [avatar contradiction clause 503]
505. $false [avatar sat refutation 372,373,467,504]
% SZS output end Proof for query10_query14_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.019 s
% ------------------------------
% ------------------------------
