% Refutation found. Thanks to Tanya!
% SZS status Theorem for query05_query09_tptp4X
% SZS output start Proof for query05_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X5,X6,X7,X8,X9] : (from_8(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
6. ! [X5] : (select_7(X5) <=> ? [X6,X7,X8,X9] : ($true & from_8(X5,X6,X7,X8,X9))) [input]
7. ! [X5] : (subquery_from_6(X5) <=> select_7(X5)) [input]
8. ! [X5] : (from_5(X5) <=> subquery_from_6(X5)) [input]
9. ! [X5] : (select_4(X5) <=> ($true & from_5(X5))) [input]
10. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
11. ! [X10] : eq(X10,X10) [input]
44. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
45. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 44]
46. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
47. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
48. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 47]
49. ! [X0] : (subquery_from_6(X0) <=> select_7(X0)) [rectify 7]
50. ! [X0] : (from_5(X0) <=> subquery_from_6(X0)) [rectify 8]
51. ! [X0] : (select_4(X0) <=> ($true & from_5(X0))) [rectify 9]
52. ! [X0] : (select_4(X0) <=> from_5(X0)) [true and false elimination 51]
53. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 10]
54. ! [X0] : eq(X0,X0) [rectify 11]
138. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 45]
139. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
140. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
141. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
142. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 141]
143. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
144. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 142,143]
145. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
146. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 46]
147. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 48]
148. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 147]
149. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
150. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0))) [skolemisation 148,149]
151. ! [X0] : ((subquery_from_6(X0) | ~select_7(X0)) & (select_7(X0) | ~subquery_from_6(X0))) [nnf transformation 49]
152. ! [X0] : ((from_5(X0) | ~subquery_from_6(X0)) & (subquery_from_6(X0) | ~from_5(X0))) [nnf transformation 50]
153. ! [X0] : ((select_4(X0) | ~from_5(X0)) & (from_5(X0) | ~select_4(X0))) [nnf transformation 52]
154. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 53]
157. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 138]
158. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
159. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 157,158]
160. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 139]
161. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 139]
163. ~eq(X0,X0) | where_3(X0) [cnf transformation 140]
164. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 144]
166. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 144]
167. ~main_query_1(X0) | select_1(X0) [cnf transformation 145]
168. ~select_1(X0) | main_query_1(X0) [cnf transformation 145]
169. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 146]
170. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 146]
171. from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [cnf transformation 150]
172. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 150]
173. ~subquery_from_6(X0) | select_7(X0) [cnf transformation 151]
174. ~select_7(X0) | subquery_from_6(X0) [cnf transformation 151]
175. ~from_5(X0) | subquery_from_6(X0) [cnf transformation 152]
176. ~subquery_from_6(X0) | from_5(X0) [cnf transformation 152]
177. ~select_4(X0) | from_5(X0) [cnf transformation 153]
178. ~from_5(X0) | select_4(X0) [cnf transformation 153]
179. ~main_query_2(X0) | select_4(X0) [cnf transformation 154]
180. ~select_4(X0) | main_query_2(X0) [cnf transformation 154]
181. eq(X0,X0) [cnf transformation 54]
212. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 159]
213. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 159]
238. 0 <=> main_query_1(sK8) [avatar definition]
239. main_query_1(sK8) <- {0} [avatar component clause 238]
241. 1 <=> ~main_query_1(sK8) [avatar definition]
242. ~main_query_1(sK8) <- {1} [avatar component clause 241]
244. 2 <=> main_query_2(sK8) [avatar definition]
245. main_query_2(sK8) <- {2} [avatar component clause 244]
247. 3 <=> ~main_query_2(sK8) [avatar definition]
248. ~main_query_2(sK8) <- {3} [avatar component clause 247]
249. ~1 | ~3 [avatar split clause 213,247,241]
250. 0 | 2 [avatar split clause 212,244,238]
251. select_4(sK8) <- {2} [resolution 179,245]
252. from_5(sK8) <- {2} [resolution 251,177]
255. subquery_from_6(sK8) <- {2} [resolution 252,175]
257. select_7(sK8) <- {2} [resolution 255,173]
259. where_3(X1) [resolution 163,181]
271. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 164,160]
273. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [resolution 171,169]
275. from_8(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 271,170]
278. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_7(X1) [resolution 273,161]
280. ~select_1(X1) | select_7(X1) [resolution 275,172]
282. ~select_7(X1) | ~where_3(X1) | select_1(X1) [resolution 278,166]
283. ~select_7(X1) | select_1(X1) [subsumption resolution 282,259]
284. select_1(sK8) <- {2} [resolution 283,257]
286. main_query_1(sK8) <- {2} [resolution 284,168]
288. $false <- {1, 2} [subsumption resolution 286,242]
289. 1 | ~2 [avatar contradiction clause 288]
290. select_1(sK8) <- {0} [resolution 239,167]
291. select_7(sK8) <- {0} [resolution 290,280]
295. subquery_from_6(sK8) <- {0} [resolution 291,174]
296. from_5(sK8) <- {0} [resolution 295,176]
298. select_4(sK8) <- {0} [resolution 296,178]
300. main_query_2(sK8) <- {0} [resolution 298,180]
302. $false <- {0, 3} [subsumption resolution 300,248]
303. ~0 | 3 [avatar contradiction clause 302]
304. $false [avatar sat refutation 249,250,289,303]
% SZS output end Proof for query05_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.010 s
% ------------------------------
% ------------------------------
