% Refutation found. Thanks to Tanya!
% SZS status Theorem for query03_query06_tptp4X
% SZS output start Proof for query03_query06_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_6(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (select_5(X0) <=> ? [X1,X2,X3,X4] : ($true & from_6(X0,X1,X2,X3,X4))) [input]
6. ! [X6] : (query_4(X6) <=> select_5(X6)) [input]
7. ! [X7] : (from_7(X7) <=> query_4(X7)) [input]
8. ! [X7] : (select_3(X7) <=> ($true & from_7(X7))) [input]
9. ! [X6] : (main_query_2(X6) <=> select_3(X6)) [input]
50. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
51. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 50]
52. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
53. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
54. ! [X0] : (select_5(X0) <=> ? [X1,X2,X3,X4] : from_6(X0,X1,X2,X3,X4)) [true and false elimination 5]
55. ! [X0] : (query_4(X0) <=> select_5(X0)) [rectify 6]
56. ! [X0] : (from_7(X0) <=> query_4(X0)) [rectify 7]
57. ! [X0] : (select_3(X0) <=> ($true & from_7(X0))) [rectify 8]
58. ! [X0] : (select_3(X0) <=> from_7(X0)) [true and false elimination 57]
59. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 9]
100. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 51]
166. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 100]
167. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
168. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 52]
169. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 168]
170. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
171. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 169,170]
172. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 53]
173. ! [X0,X1,X2,X3,X4] : ((from_6(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_6(X0,X1,X2,X3,X4))) [nnf transformation 4]
174. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_6(X0,X1,X2,X3,X4) | ~select_5(X0))) [nnf transformation 54]
175. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_6(X0,X5,X6,X7,X8) | ~select_5(X0))) [rectify 174]
176. ! [X0] : (? [X5,X6,X7,X8] : from_6(X0,X5,X6,X7,X8) => from_6(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
177. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (from_6(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_5(X0))) [skolemisation 175,176]
178. ! [X0] : ((query_4(X0) | ~select_5(X0)) & (select_5(X0) | ~query_4(X0))) [nnf transformation 55]
179. ! [X0] : ((from_7(X0) | ~query_4(X0)) & (query_4(X0) | ~from_7(X0))) [nnf transformation 56]
180. ! [X0] : ((select_3(X0) | ~from_7(X0)) & (from_7(X0) | ~select_3(X0))) [nnf transformation 58]
181. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 59]
184. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 166]
185. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
186. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 184,185]
187. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 167]
188. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 167]
189. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 171]
190. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 171]
191. ~main_query_1(X0) | select_1(X0) [cnf transformation 172]
192. ~select_1(X0) | main_query_1(X0) [cnf transformation 172]
193. ~from_6(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 173]
194. ~tab1(X0,X1,X2,X3,X4) | from_6(X0,X1,X2,X3,X4) [cnf transformation 173]
195. from_6(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_5(X0) [cnf transformation 177]
196. ~from_6(X0,X1,X2,X3,X4) | select_5(X0) [cnf transformation 177]
197. ~query_4(X0) | select_5(X0) [cnf transformation 178]
198. ~select_5(X0) | query_4(X0) [cnf transformation 178]
199. ~from_7(X0) | query_4(X0) [cnf transformation 179]
200. ~query_4(X0) | from_7(X0) [cnf transformation 179]
201. ~select_3(X0) | from_7(X0) [cnf transformation 180]
202. ~from_7(X0) | select_3(X0) [cnf transformation 180]
203. ~main_query_2(X0) | select_3(X0) [cnf transformation 181]
204. ~select_3(X0) | main_query_2(X0) [cnf transformation 181]
243. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 186]
244. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 186]
276. 0 <=> main_query_1(sK8) [avatar definition]
277. main_query_1(sK8) <- {0} [avatar component clause 276]
279. 1 <=> ~main_query_1(sK8) [avatar definition]
280. ~main_query_1(sK8) <- {1} [avatar component clause 279]
282. 2 <=> main_query_2(sK8) [avatar definition]
283. main_query_2(sK8) <- {2} [avatar component clause 282]
285. 3 <=> ~main_query_2(sK8) [avatar definition]
286. ~main_query_2(sK8) <- {3} [avatar component clause 285]
287. ~1 | ~3 [avatar split clause 244,285,279]
288. 0 | 2 [avatar split clause 243,282,276]
289. select_3(sK8) <- {2} [resolution 203,283]
291. from_7(sK8) <- {2} [resolution 289,201]
294. query_4(sK8) <- {2} [resolution 291,199]
297. select_5(sK8) <- {2} [resolution 294,197]
303. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 189,187]
305. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_5(X0) [resolution 195,193]
307. from_6(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 303,194]
310. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_5(X1) [resolution 305,188]
312. ~select_1(X1) | select_5(X1) [resolution 307,196]
314. ~select_5(X1) | select_1(X1) [resolution 310,190]
315. select_1(sK8) <- {2} [resolution 314,297]
317. main_query_1(sK8) <- {2} [resolution 315,192]
318. $false <- {1, 2} [subsumption resolution 317,280]
319. 1 | ~2 [avatar contradiction clause 318]
320. select_1(sK8) <- {0} [resolution 277,191]
321. select_5(sK8) <- {0} [resolution 320,312]
324. query_4(sK8) <- {0} [resolution 321,198]
325. from_7(sK8) <- {0} [resolution 324,200]
327. select_3(sK8) <- {0} [resolution 325,202]
329. main_query_2(sK8) <- {0} [resolution 327,204]
331. $false <- {0, 3} [subsumption resolution 329,286]
332. ~0 | 3 [avatar contradiction clause 331]
333. $false [avatar sat refutation 287,288,319,332]
% SZS output end Proof for query03_query06_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
