% Refutation found. Thanks to Tanya!
% SZS status Theorem for query10_query11_tptp4X
% SZS output start Proof for query10_query11_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8] : (where_6(X7,X8) <=> eq(X7,X8)) [input]
7. ! [X7,X8] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X7,X8) & from_5(X7,X8,X9,X10,X11))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
55. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
56. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 55]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
58. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
59. ! [X0,X1] : (where_6(X0,X1) <=> eq(X0,X1)) [rectify 6]
60. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
61. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
108. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 56]
184. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 108]
185. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
186. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
187. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
188. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 187]
189. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
190. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 188,189]
191. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
192. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 58]
193. ! [X0,X1] : ((where_6(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_6(X0,X1))) [nnf transformation 59]
194. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 60]
195. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 194]
196. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
197. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 195,196]
198. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 61]
201. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 184]
202. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
203. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 201,202]
204. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 185]
205. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 185]
206. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 186]
207. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 186]
208. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 190]
209. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 190]
210. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 190]
211. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 191]
212. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 191]
213. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 192]
214. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 192]
215. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 193]
216. ~eq(X0,X1) | where_6(X0,X1) [cnf transformation 193]
217. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 197]
218. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 197]
219. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 197]
220. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 198]
221. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 198]
266. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 203]
267. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 203]
269. 0 <=> main_query_1(sK6,sK7) [avatar definition]
270. main_query_1(sK6,sK7) <- {0} [avatar component clause 269]
272. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
273. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 272]
275. 2 <=> main_query_2(sK6,sK7) [avatar definition]
276. main_query_2(sK6,sK7) <- {2} [avatar component clause 275]
278. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
279. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 278]
280. ~1 | ~3 [avatar split clause 267,278,272]
281. 0 | 2 [avatar split clause 266,275,269]
286. select_4(sK6,sK7) <- {2} [resolution 220,276]
288. where_6(sK6,sK7) <- {2} [resolution 286,218]
290. eq(sK6,sK7) <- {2} [resolution 288,215]
295. where_3(sK6,sK7) <- {2} [resolution 290,207]
391. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 208,204]
395. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 217,213]
515. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 391,214]
534. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 395,205]
961. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 515,219]
971. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 534,210]
984. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 971,286]
988. select_1(sK6,sK7) <- {2} [subsumption resolution 984,295]
1001. main_query_1(sK6,sK7) <- {2} [resolution 988,212]
1003. $false <- {1, 2} [subsumption resolution 1001,273]
1004. 1 | ~2 [avatar contradiction clause 1003]
1007. select_1(sK6,sK7) <- {0} [resolution 270,211]
1011. where_3(sK6,sK7) <- {0} [resolution 1007,209]
1012. eq(sK6,sK7) <- {0} [resolution 1011,206]
1015. where_6(sK6,sK7) <- {0} [resolution 1012,216]
1025. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 1015,961]
1029. select_4(sK6,sK7) <- {0} [subsumption resolution 1025,1007]
1115. main_query_2(sK6,sK7) <- {0} [resolution 1029,221]
1117. $false <- {0, 3} [subsumption resolution 1115,279]
1118. ~0 | 3 [avatar contradiction clause 1117]
1119. $false [avatar sat refutation 280,281,1004,1118]
% SZS output end Proof for query10_query11_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.025 s
% ------------------------------
% ------------------------------
