% Refutation found. Thanks to Tanya!
% SZS status Theorem for query07_query09_tptp4X
% SZS output start Proof for query07_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X6,X7,X8,X9,X10] : (from_7(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
6. ! [X6] : (select_6(X6) <=> ? [X7,X8,X9,X10] : ($true & from_7(X6,X7,X8,X9,X10))) [input]
7. ! [X5] : (query_5(X5) <=> select_6(X5)) [input]
8. ! [X11] : (from_8(X11) <=> query_5(X11)) [input]
9. ! [X11] : (select_4(X11) <=> ($true & from_8(X11))) [input]
10. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
11. ! [X12] : eq(X12,X12) [input]
17. ! [X12,X13] : (lt(X12,X13) <=> (~eq(X12,X13) & lte(X12,X13))) [input]
18. ! [X12,X13] : (gt(X12,X13) <=> lt(X13,X12)) [input]
52. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
53. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 52]
54. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
55. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
56. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [rectify 6]
57. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 56]
58. ! [X0] : (query_5(X0) <=> select_6(X0)) [rectify 7]
59. ! [X0] : (from_8(X0) <=> query_5(X0)) [rectify 8]
60. ! [X0] : (select_4(X0) <=> ($true & from_8(X0))) [rectify 9]
61. ! [X0] : (select_4(X0) <=> from_8(X0)) [true and false elimination 60]
62. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 10]
63. ! [X0] : eq(X0,X0) [rectify 11]
69. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 17]
70. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 18]
104. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 53]
172. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 104]
173. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
174. ! [X0] : ((where_3(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
175. ! [X0] : ((where_3(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [flattening 174]
176. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
177. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 176]
178. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
179. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 177,178]
180. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 54]
181. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 55]
182. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 57]
183. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 182]
184. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
185. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 183,184]
186. ! [X0] : ((query_5(X0) | ~select_6(X0)) & (select_6(X0) | ~query_5(X0))) [nnf transformation 58]
187. ! [X0] : ((from_8(X0) | ~query_5(X0)) & (query_5(X0) | ~from_8(X0))) [nnf transformation 59]
188. ! [X0] : ((select_4(X0) | ~from_8(X0)) & (from_8(X0) | ~select_4(X0))) [nnf transformation 61]
189. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 62]
192. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 69]
193. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 192]
194. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 70]
195. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 172]
196. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
197. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 195,196]
198. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 173]
199. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 173]
202. gt(X0,X0) | where_3(X0) | lt(X0,X0) [cnf transformation 175]
203. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 179]
205. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 179]
206. ~main_query_1(X0) | select_1(X0) [cnf transformation 180]
207. ~select_1(X0) | main_query_1(X0) [cnf transformation 180]
208. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 181]
209. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 181]
210. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 185]
211. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 185]
212. ~query_5(X0) | select_6(X0) [cnf transformation 186]
213. ~select_6(X0) | query_5(X0) [cnf transformation 186]
214. ~from_8(X0) | query_5(X0) [cnf transformation 187]
215. ~query_5(X0) | from_8(X0) [cnf transformation 187]
216. ~select_4(X0) | from_8(X0) [cnf transformation 188]
217. ~from_8(X0) | select_4(X0) [cnf transformation 188]
218. ~main_query_2(X0) | select_4(X0) [cnf transformation 189]
219. ~select_4(X0) | main_query_2(X0) [cnf transformation 189]
220. eq(X0,X0) [cnf transformation 63]
229. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 193]
231. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 194]
264. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 197]
265. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 197]
298. 0 <=> main_query_1(sK8) [avatar definition]
299. main_query_1(sK8) <- {0} [avatar component clause 298]
301. 1 <=> ~main_query_1(sK8) [avatar definition]
302. ~main_query_1(sK8) <- {1} [avatar component clause 301]
304. 2 <=> main_query_2(sK8) [avatar definition]
305. main_query_2(sK8) <- {2} [avatar component clause 304]
307. 3 <=> ~main_query_2(sK8) [avatar definition]
308. ~main_query_2(sK8) <- {3} [avatar component clause 307]
309. ~1 | ~3 [avatar split clause 265,307,301]
310. 0 | 2 [avatar split clause 264,304,298]
311. select_4(sK8) <- {2} [resolution 218,305]
312. from_8(sK8) <- {2} [resolution 311,216]
315. query_5(sK8) <- {2} [resolution 312,214]
317. select_6(sK8) <- {2} [resolution 315,212]
322. ~lt(X0,X0) [resolution 229,220]
324. where_3(X1) | lt(X1,X1) | lt(X1,X1) [resolution 202,231]
325. where_3(X1) | lt(X1,X1) [duplicate literal removal 324]
326. where_3(X1) [subsumption resolution 325,322]
331. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 203,198]
333. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 210,208]
335. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 331,209]
338. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 333,199]
340. ~select_1(X1) | select_6(X1) [resolution 335,211]
342. ~select_6(X1) | ~where_3(X1) | select_1(X1) [resolution 338,205]
343. ~select_6(X1) | select_1(X1) [subsumption resolution 342,326]
344. select_1(sK8) <- {2} [resolution 343,317]
346. main_query_1(sK8) <- {2} [resolution 344,207]
348. $false <- {1, 2} [subsumption resolution 346,302]
349. 1 | ~2 [avatar contradiction clause 348]
350. select_1(sK8) <- {0} [resolution 299,206]
351. select_6(sK8) <- {0} [resolution 350,340]
355. query_5(sK8) <- {0} [resolution 351,213]
356. from_8(sK8) <- {0} [resolution 355,215]
358. select_4(sK8) <- {0} [resolution 356,217]
360. main_query_2(sK8) <- {0} [resolution 358,219]
362. $false <- {0, 3} [subsumption resolution 360,308]
363. ~0 | 3 [avatar contradiction clause 362]
364. $false [avatar sat refutation 309,310,349,363]
% SZS output end Proof for query07_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
