% Refutation found. Thanks to Tanya!
% SZS status Theorem for query10_query11_tptp4X
% SZS output start Proof for query10_query11_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [input]
5. ! [X0,X1,X5,X6,X7] : (from_5(X0,X1,X5,X6,X7) <=> tab1(X0,X1,X5,X6,X7)) [input]
6. ! [X0,X1] : (where_6(X0,X1) <=> eq(X0,X1)) [input]
7. ! [X0,X1] : (select_4(X0,X1) <=> ? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7))) [input]
8. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [input]
54. ! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [input]
55. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [negated conjecture 54]
56. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
57. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
178. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 55]
179. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
180. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
181. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
182. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 181]
183. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
184. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 182,183]
185. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 4]
186. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 56]
187. ! [X0,X1] : ((where_6(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_6(X0,X1))) [nnf transformation 6]
188. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 57]
189. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 188]
190. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
191. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 189,190]
192. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 8]
195. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 178]
196. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
197. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 195,196]
198. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 179]
199. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 179]
200. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 180]
201. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 180]
202. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 184]
203. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 184]
204. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 184]
205. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 185]
206. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 185]
207. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 186]
208. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 186]
209. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 187]
210. ~eq(X0,X1) | where_6(X0,X1) [cnf transformation 187]
211. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 191]
212. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 191]
213. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 191]
214. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 192]
215. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 192]
259. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 197]
260. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 197]
297. 0 <=> main_query_1(sK6,sK7) [avatar definition]
298. main_query_1(sK6,sK7) <- {0} [avatar component clause 297]
300. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
301. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 300]
303. 2 <=> main_query_2(sK6,sK7) [avatar definition]
304. main_query_2(sK6,sK7) <- {2} [avatar component clause 303]
306. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
307. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 306]
308. ~1 | ~3 [avatar split clause 260,306,300]
309. 0 | 2 [avatar split clause 259,303,297]
314. select_4(sK6,sK7) <- {2} [resolution 214,304]
316. where_6(sK6,sK7) <- {2} [resolution 314,212]
318. eq(sK6,sK7) <- {2} [resolution 316,209]
323. where_3(sK6,sK7) <- {2} [resolution 318,201]
352. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 202,198]
354. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 211,207]
355. from_5(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 352,208]
358. from_2(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 354,199]
359. ~where_6(X0,X1) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 355,213]
377. ~select_4(X0,X1) | ~where_3(X0,X1) | select_1(X0,X1) [resolution 358,204]
379. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 377,314]
380. select_1(sK6,sK7) <- {2} [subsumption resolution 379,323]
381. main_query_1(sK6,sK7) <- {2} [resolution 380,206]
383. $false <- {1, 2} [subsumption resolution 381,301]
384. 1 | ~2 [avatar contradiction clause 383]
385. select_1(sK6,sK7) <- {0} [resolution 298,205]
387. where_3(sK6,sK7) <- {0} [resolution 385,203]
388. eq(sK6,sK7) <- {0} [resolution 387,200]
393. where_6(sK6,sK7) <- {0} [resolution 388,210]
405. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 393,359]
407. select_4(sK6,sK7) <- {0} [subsumption resolution 405,385]
414. main_query_2(sK6,sK7) <- {0} [resolution 407,215]
416. $false <- {0, 3} [subsumption resolution 414,307]
417. ~0 | 3 [avatar contradiction clause 416]
418. $false [avatar sat refutation 308,309,384,417]
% SZS output end Proof for query10_query11_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.019 s
% ------------------------------
% ------------------------------
