% Refutation found. Thanks to Tanya!
% SZS status Theorem for query10_query13_tptp4X
% SZS output start Proof for query10_query13_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1] : (where_6(X0,X1) <=> (~gt(X0,X1) & ~lt(X0,X1))) [input]
7. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
13. ! [X7,X8] : ((lte(X8,X7) & lte(X7,X8)) <=> eq(X7,X8)) [input]
15. ! [X7,X8] : (lte(X8,X7) | lte(X7,X8)) [input]
16. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
17. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
29. ! [X7,X0,X1] : ((where_6(X0,X1) & eq(X0,X7)) => where_6(X7,X1)) [input]
55. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
56. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 55]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
58. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
59. ! [X0] : eq(X0,X0) [rectify 9]
63. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
65. ! [X0,X1] : (lte(X1,X0) | lte(X0,X1)) [rectify 15]
66. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 16]
67. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 17]
79. ! [X0,X1,X2] : ((where_6(X1,X2) & eq(X1,X0)) => where_6(X0,X2)) [rectify 29]
105. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 56]
129. ! [X0,X1,X2] : (where_6(X0,X2) | (~where_6(X1,X2) | ~eq(X1,X0))) [ennf transformation 79]
130. ! [X0,X1,X2] : (where_6(X0,X2) | ~where_6(X1,X2) | ~eq(X1,X0)) [flattening 129]
181. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 105]
182. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
183. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
184. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
185. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 184]
186. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
187. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 185,186]
188. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
189. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
190. ! [X0,X1] : ((where_6(X0,X1) | (gt(X0,X1) | lt(X0,X1))) & ((~gt(X0,X1) & ~lt(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 6]
191. ! [X0,X1] : ((where_6(X0,X1) | gt(X0,X1) | lt(X0,X1)) & ((~gt(X0,X1) & ~lt(X0,X1)) | ~where_6(X0,X1))) [flattening 190]
192. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 7]
193. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 192]
194. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
195. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 193,194]
196. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 58]
197. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 63]
198. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 197]
199. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 66]
200. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 199]
201. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 67]
202. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 181]
203. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
204. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 202,203]
205. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 182]
206. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 182]
207. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 183]
208. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 183]
209. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 187]
210. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 187]
211. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 187]
212. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 188]
213. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 188]
214. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 189]
215. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 189]
216. ~lt(X0,X1) | ~where_6(X0,X1) [cnf transformation 191]
217. ~gt(X0,X1) | ~where_6(X0,X1) [cnf transformation 191]
218. gt(X0,X1) | where_6(X0,X1) | lt(X0,X1) [cnf transformation 191]
219. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 195]
220. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 195]
221. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 195]
222. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 196]
223. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 196]
224. eq(X0,X0) [cnf transformation 59]
230. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 198]
232. lte(X1,X0) | lte(X0,X1) [cnf transformation 65]
234. ~lt(X0,X1) | ~eq(X0,X1) [cnf transformation 200]
235. ~lte(X0,X1) | eq(X0,X1) | lt(X0,X1) [cnf transformation 200]
236. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 201]
237. ~lt(X1,X0) | gt(X0,X1) [cnf transformation 201]
247. ~where_6(X1,X2) | where_6(X0,X2) | ~eq(X1,X0) [cnf transformation 130]
273. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 204]
274. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 204]
276. 0 <=> main_query_1(sK6,sK7) [avatar definition]
277. main_query_1(sK6,sK7) <- {0} [avatar component clause 276]
279. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
280. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 279]
282. 2 <=> main_query_2(sK6,sK7) [avatar definition]
283. main_query_2(sK6,sK7) <- {2} [avatar component clause 282]
285. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
286. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 285]
287. ~1 | ~3 [avatar split clause 274,285,279]
288. 0 | 2 [avatar split clause 273,282,276]
291. select_4(sK6,sK7) <- {2} [resolution 222,283]
292. where_6(sK6,sK7) <- {2} [resolution 291,220]
298. lt(X1,X0) | lt(X0,X1) | where_6(X0,X1) [resolution 218,236]
308. eq(X1,X2) | lt(X1,X2) | lte(X2,X1) [resolution 235,232]
310. lte(X2,X1) | lt(X1,X2) [subsumption resolution 308,230]
311. lt(X1,X0) | lt(X0,X1) | eq(X1,X0) [resolution 310,235]
333. ~eq(X3,X2) | where_6(X3,X2) | lt(X2,X3) [resolution 298,234]
336. lt(X0,X0) | where_6(X0,X0) [factoring 298]
338. where_6(X1,X1) | ~eq(X1,X1) [resolution 336,234]
341. where_6(X1,X1) [subsumption resolution 338,224]
343. ~eq(X3,X2) | where_6(X2,X3) [resolution 341,247]
366. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 209,205]
369. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 219,214]
373. 5 <=> ~lt(sK7,sK6) [avatar definition]
374. ~lt(sK7,sK6) <- {5} [avatar component clause 373]
376. 4 <=> lt(sK7,sK6) [avatar definition]
377. lt(sK7,sK6) <- {4} [avatar component clause 376]
379. 7 <=> ~where_6(sK7,sK6) [avatar definition]
382. 6 <=> where_6(sK7,sK6) [avatar definition]
386. gt(sK6,sK7) <- {4} [resolution 377,237]
389. ~where_6(sK7,sK6) <- {4} [resolution 377,216]
390. ~7 | ~4 [avatar split clause 389,376,379]
394. ~where_6(sK6,sK7) <- {4} [resolution 386,217]
395. $false <- {2, 4} [subsumption resolution 394,292]
396. ~2 | ~4 [avatar contradiction clause 395]
398. lt(sK6,sK7) | eq(sK6,sK7) <- {5} [resolution 374,311]
403. 9 <=> ~eq(sK6,sK7) [avatar definition]
404. ~eq(sK6,sK7) <- {9} [avatar component clause 403]
406. 8 <=> eq(sK6,sK7) [avatar definition]
407. eq(sK6,sK7) <- {8} [avatar component clause 406]
412. 10 <=> lt(sK6,sK7) [avatar definition]
413. lt(sK6,sK7) <- {10} [avatar component clause 412]
414. 8 | 10 | 5 [avatar split clause 398,373,412,406]
429. ~where_6(sK6,sK7) <- {10} [resolution 413,216]
431. $false <- {2, 10} [subsumption resolution 429,292]
432. ~2 | ~10 [avatar contradiction clause 431]
436. where_3(sK6,sK7) <- {8} [resolution 407,208]
445. where_6(sK7,sK6) <- {8} [resolution 407,343]
501. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 366,215]
509. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 369,206]
621. where_6(sK6,sK7) | lt(sK7,sK6) <- {8} [resolution 333,407]
691. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 501,221]
710. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 509,211]
959. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 710,291]
963. select_1(sK6,sK7) <- {2, 8} [subsumption resolution 959,436]
978. main_query_1(sK6,sK7) <- {2, 8} [resolution 963,213]
980. $false <- {1, 2, 8} [subsumption resolution 978,280]
981. 1 | ~2 | ~8 [avatar contradiction clause 980]
995. where_6(sK6,sK7) <- {5, 8} [subsumption resolution 621,374]
998. select_1(sK6,sK7) <- {0} [resolution 277,212]
1004. where_3(sK6,sK7) <- {0} [resolution 998,210]
1014. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {5, 8} [resolution 995,691]
1017. select_4(sK6,sK7) <- {0, 5, 8} [subsumption resolution 1014,998]
1033. main_query_2(sK6,sK7) <- {0, 5, 8} [resolution 1017,223]
1035. $false <- {0, 3, 5, 8} [subsumption resolution 1033,286]
1036. ~0 | 3 | 5 | ~8 [avatar contradiction clause 1035]
1041. 6 | ~8 [avatar split clause 445,406,382]
1043. eq(sK6,sK7) <- {0} [resolution 1004,207]
1044. $false <- {0, 9} [subsumption resolution 1043,404]
1045. ~0 | 9 [avatar contradiction clause 1044]
1046. $false [avatar sat refutation 287,288,390,396,414,432,981,1036,1041,1045]
% SZS output end Proof for query10_query13_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.026 s
% ------------------------------
% ------------------------------
