% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_08_query_09_tptp4X
% SZS output start Proof for query_08_query_09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : ($true & from_5(X0,X1,X2,X3,X4))) [input]
3. ! [X5,X6] : (subquery_from_3(X5,X6) <=> select_4(X5,X6)) [input]
4. ! [X7,X8] : (from_2(X7,X8) <=> subquery_from_3(X7,X8)) [input]
5. ! [X7,X8] : (select_1(X7,X8) <=> ($true & from_2(X7,X8))) [input]
6. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
7. ! [X9,X10,X11,X12,X13] : (from_7(X9,X10,X11,X12,X13) <=> tab1(X9,X10,X11,X12,X13)) [input]
8. ! [X9,X10,X11,X12,X13] : (from_11(X9,X10,X11,X12,X13) <=> tab1(X9,X10,X11,X12,X13)) [input]
9. ! [X9] : (select_10(X9) <=> ? [X10,X11,X12,X13] : ($true & from_11(X9,X10,X11,X12,X13))) [input]
10. ! [X5] : (where_8_subquery_9(X5) <=> select_10(X5)) [input]
11. ! [X9] : (where_8(X9) <=> where_8_subquery_9(X9)) [input]
12. ! [X9,X10] : (select_6(X9,X10) <=> ? [X11,X12,X13] : (where_8(X9) & from_7(X9,X10,X11,X12,X13))) [input]
13. ! [X5,X6] : (main_query_2(X5,X6) <=> select_6(X5,X6)) [input]
63. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
64. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 63]
65. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : from_5(X0,X1,X2,X3,X4)) [true and false elimination 2]
66. ! [X0,X1] : (subquery_from_3(X0,X1) <=> select_4(X0,X1)) [rectify 3]
67. ! [X0,X1] : (from_2(X0,X1) <=> subquery_from_3(X0,X1)) [rectify 4]
68. ! [X0,X1] : (select_1(X0,X1) <=> ($true & from_2(X0,X1))) [rectify 5]
69. ! [X0,X1] : (select_1(X0,X1) <=> from_2(X0,X1)) [true and false elimination 68]
70. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 6]
71. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 7]
72. ! [X0,X1,X2,X3,X4] : (from_11(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 8]
73. ! [X0] : (select_10(X0) <=> ? [X1,X2,X3,X4] : ($true & from_11(X0,X1,X2,X3,X4))) [rectify 9]
74. ! [X0] : (select_10(X0) <=> ? [X1,X2,X3,X4] : from_11(X0,X1,X2,X3,X4)) [true and false elimination 73]
75. ! [X0] : (where_8_subquery_9(X0) <=> select_10(X0)) [rectify 10]
76. ! [X0] : (where_8(X0) <=> where_8_subquery_9(X0)) [rectify 11]
77. ! [X0,X1] : (select_6(X0,X1) <=> ? [X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4))) [rectify 12]
78. ! [X0,X1] : (main_query_2(X0,X1) <=> select_6(X0,X1)) [rectify 13]
128. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 64]
212. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 128]
213. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 1]
214. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_5(X0,X1,X2,X3,X4) | ~select_4(X0,X1))) [nnf transformation 65]
215. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_5(X0,X1,X5,X6,X7) | ~select_4(X0,X1))) [rectify 214]
216. ! [X1,X0] : (? [X5,X6,X7] : from_5(X0,X1,X5,X6,X7) => from_5(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) [choice axiom]
217. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (from_5(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_4(X0,X1))) [skolemisation 215,216]
218. ! [X0,X1] : ((subquery_from_3(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~subquery_from_3(X0,X1))) [nnf transformation 66]
219. ! [X0,X1] : ((from_2(X0,X1) | ~subquery_from_3(X0,X1)) & (subquery_from_3(X0,X1) | ~from_2(X0,X1))) [nnf transformation 67]
220. ! [X0,X1] : ((select_1(X0,X1) | ~from_2(X0,X1)) & (from_2(X0,X1) | ~select_1(X0,X1))) [nnf transformation 69]
221. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 70]
222. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 71]
223. ! [X0,X1,X2,X3,X4] : ((from_11(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_11(X0,X1,X2,X3,X4))) [nnf transformation 72]
224. ! [X0] : ((select_10(X0) | ! [X1,X2,X3,X4] : ~from_11(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_11(X0,X1,X2,X3,X4) | ~select_10(X0))) [nnf transformation 74]
225. ! [X0] : ((select_10(X0) | ! [X1,X2,X3,X4] : ~from_11(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_11(X0,X5,X6,X7,X8) | ~select_10(X0))) [rectify 224]
226. ! [X0] : (? [X5,X6,X7,X8] : from_11(X0,X5,X6,X7,X8) => from_11(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
227. ! [X0] : ((select_10(X0) | ! [X1,X2,X3,X4] : ~from_11(X0,X1,X2,X3,X4)) & (from_11(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_10(X0))) [skolemisation 225,226]
228. ! [X0] : ((where_8_subquery_9(X0) | ~select_10(X0)) & (select_10(X0) | ~where_8_subquery_9(X0))) [nnf transformation 75]
229. ! [X0] : ((where_8(X0) | ~where_8_subquery_9(X0)) & (where_8_subquery_9(X0) | ~where_8(X0))) [nnf transformation 76]
230. ! [X0,X1] : ((select_6(X0,X1) | ! [X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4)) | ~select_6(X0,X1))) [nnf transformation 77]
231. ! [X0,X1] : ((select_6(X0,X1) | ! [X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_8(X0) & from_7(X0,X1,X5,X6,X7)) | ~select_6(X0,X1))) [rectify 230]
232. ! [X1,X0] : (? [X5,X6,X7] : (where_8(X0) & from_7(X0,X1,X5,X6,X7)) => (where_8(X0) & from_7(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
233. ! [X0,X1] : ((select_6(X0,X1) | ! [X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & ((where_8(X0) & from_7(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_6(X0,X1))) [skolemisation 231,232]
234. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_6(X0,X1)) & (select_6(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 78]
237. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 212]
238. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
239. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 237,238]
240. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 213]
241. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 213]
242. from_5(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_4(X0,X1) [cnf transformation 217]
243. ~from_5(X0,X1,X2,X3,X4) | select_4(X0,X1) [cnf transformation 217]
244. ~subquery_from_3(X0,X1) | select_4(X0,X1) [cnf transformation 218]
245. ~select_4(X0,X1) | subquery_from_3(X0,X1) [cnf transformation 218]
246. ~from_2(X0,X1) | subquery_from_3(X0,X1) [cnf transformation 219]
247. ~subquery_from_3(X0,X1) | from_2(X0,X1) [cnf transformation 219]
248. ~select_1(X0,X1) | from_2(X0,X1) [cnf transformation 220]
249. ~from_2(X0,X1) | select_1(X0,X1) [cnf transformation 220]
250. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 221]
251. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 221]
252. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 222]
253. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 222]
255. ~tab1(X0,X1,X2,X3,X4) | from_11(X0,X1,X2,X3,X4) [cnf transformation 223]
257. ~from_11(X0,X1,X2,X3,X4) | select_10(X0) [cnf transformation 227]
259. ~select_10(X0) | where_8_subquery_9(X0) [cnf transformation 228]
261. ~where_8_subquery_9(X0) | where_8(X0) [cnf transformation 229]
262. from_7(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_6(X0,X1) [cnf transformation 233]
264. ~from_7(X0,X1,X2,X3,X4) | ~where_8(X0) | select_6(X0,X1) [cnf transformation 233]
265. ~main_query_2(X0,X1) | select_6(X0,X1) [cnf transformation 234]
266. ~select_6(X0,X1) | main_query_2(X0,X1) [cnf transformation 234]
314. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 239]
315. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 239]
317. 0 <=> main_query_1(sK10,sK11) [avatar definition]
318. main_query_1(sK10,sK11) <- {0} [avatar component clause 317]
320. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
321. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 320]
323. 2 <=> main_query_2(sK10,sK11) [avatar definition]
324. main_query_2(sK10,sK11) <- {2} [avatar component clause 323]
326. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
327. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 326]
328. ~1 | ~3 [avatar split clause 315,326,320]
329. 0 | 2 [avatar split clause 314,323,317]
330. select_6(sK10,sK11) <- {2} [resolution 265,324]
359. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_4(X0,X1) [resolution 242,240]
364. tab1(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_6(X0,X1) [resolution 262,252]
412. from_7(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_4(X2,X3) [resolution 359,253]
413. from_11(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_4(X4,X5) [resolution 359,255]
419. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_6(X0,X1) [resolution 364,241]
539. ~select_4(X2,X3) | ~where_8(X2) | select_6(X2,X3) [resolution 412,264]
561. ~select_4(X2,X3) | select_10(X2) [resolution 413,257]
583. ~select_6(X2,X3) | select_4(X2,X3) [resolution 419,243]
596. select_4(sK10,sK11) <- {2} [resolution 583,330]
607. subquery_from_3(sK10,sK11) <- {2} [resolution 596,245]
616. from_2(sK10,sK11) <- {2} [resolution 607,247]
626. select_1(sK10,sK11) <- {2} [resolution 616,249]
636. main_query_1(sK10,sK11) <- {2} [resolution 626,251]
638. $false <- {1, 2} [subsumption resolution 636,321]
639. 1 | ~2 [avatar contradiction clause 638]
640. select_1(sK10,sK11) <- {0} [resolution 318,250]
642. from_2(sK10,sK11) <- {0} [resolution 640,248]
644. subquery_from_3(sK10,sK11) <- {0} [resolution 642,246]
654. select_4(sK10,sK11) <- {0} [resolution 644,244]
663. select_10(sK10) <- {0} [resolution 654,561]
664. ~where_8(sK10) | select_6(sK10,sK11) <- {0} [resolution 654,539]
670. 4 <=> select_6(sK10,sK11) [avatar definition]
671. select_6(sK10,sK11) <- {4} [avatar component clause 670]
673. 6 <=> where_8(sK10) [avatar definition]
676. 7 <=> ~where_8(sK10) [avatar definition]
678. 4 | ~7 | ~0 [avatar split clause 664,317,676,670]
687. where_8_subquery_9(sK10) <- {0} [resolution 663,259]
704. where_8(sK10) <- {0} [resolution 687,261]
708. 6 | ~0 [avatar split clause 704,317,673]
727. main_query_2(sK10,sK11) <- {4} [resolution 671,266]
729. $false <- {3, 4} [subsumption resolution 727,327]
730. 3 | ~4 [avatar contradiction clause 729]
731. $false [avatar sat refutation 328,329,639,678,708,730]
% SZS output end Proof for query_08_query_09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.016 s
% ------------------------------
% ------------------------------
