% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_06_query_08_tptp4X
% SZS output start Proof for query_06_query_08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_8(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8] : (select_7(X7,X8) <=> ? [X9,X10,X11] : ($true & from_8(X7,X8,X9,X10,X11))) [input]
7. ! [X5,X6] : (subquery_from_6(X5,X6) <=> select_7(X5,X6)) [input]
8. ! [X12,X13] : (from_5(X12,X13) <=> subquery_from_6(X12,X13)) [input]
9. ! [X12,X13] : (select_4(X12,X13) <=> ($true & from_5(X12,X13))) [input]
10. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
14. ! [X14] : lte(X14,X14) [input]
20. ! [X14,X15] : (gte(X14,X15) <=> lte(X15,X14)) [input]
55. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
56. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 55]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
58. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
59. ! [X0,X1] : (select_7(X0,X1) <=> ? [X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
60. ! [X0,X1] : (select_7(X0,X1) <=> ? [X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 59]
61. ! [X0,X1] : (subquery_from_6(X0,X1) <=> select_7(X0,X1)) [rectify 7]
62. ! [X0,X1] : (from_5(X0,X1) <=> subquery_from_6(X0,X1)) [rectify 8]
63. ! [X0,X1] : (select_4(X0,X1) <=> ($true & from_5(X0,X1))) [rectify 9]
64. ! [X0,X1] : (select_4(X0,X1) <=> from_5(X0,X1)) [true and false elimination 63]
65. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 10]
69. ! [X0] : lte(X0,X0) [rectify 14]
75. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 20]
110. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 56]
184. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 110]
185. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
186. ! [X0] : ((where_3(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
187. ! [X0] : ((where_3(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_3(X0))) [flattening 186]
188. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
189. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 188]
190. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
191. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 189,190]
192. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
193. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 58]
194. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0,X1))) [nnf transformation 60]
195. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_8(X0,X1,X5,X6,X7) | ~select_7(X0,X1))) [rectify 194]
196. ! [X1,X0] : (? [X5,X6,X7] : from_8(X0,X1,X5,X6,X7) => from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) [choice axiom]
197. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1))) [skolemisation 195,196]
198. ! [X0,X1] : ((subquery_from_6(X0,X1) | ~select_7(X0,X1)) & (select_7(X0,X1) | ~subquery_from_6(X0,X1))) [nnf transformation 61]
199. ! [X0,X1] : ((from_5(X0,X1) | ~subquery_from_6(X0,X1)) & (subquery_from_6(X0,X1) | ~from_5(X0,X1))) [nnf transformation 62]
200. ! [X0,X1] : ((select_4(X0,X1) | ~from_5(X0,X1)) & (from_5(X0,X1) | ~select_4(X0,X1))) [nnf transformation 64]
201. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 65]
204. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 75]
205. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 184]
206. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
207. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 205,206]
208. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 185]
209. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 185]
212. ~gte(X0,X0) | where_3(X0) | ~lte(X0,X0) [cnf transformation 187]
213. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 191]
215. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0,X1) [cnf transformation 191]
216. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 192]
217. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 192]
218. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 193]
219. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 193]
220. from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1) [cnf transformation 197]
221. ~from_8(X0,X1,X2,X3,X4) | select_7(X0,X1) [cnf transformation 197]
222. ~subquery_from_6(X0,X1) | select_7(X0,X1) [cnf transformation 198]
223. ~select_7(X0,X1) | subquery_from_6(X0,X1) [cnf transformation 198]
224. ~from_5(X0,X1) | subquery_from_6(X0,X1) [cnf transformation 199]
225. ~subquery_from_6(X0,X1) | from_5(X0,X1) [cnf transformation 199]
226. ~select_4(X0,X1) | from_5(X0,X1) [cnf transformation 200]
227. ~from_5(X0,X1) | select_4(X0,X1) [cnf transformation 200]
228. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 201]
229. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 201]
233. lte(X0,X0) [cnf transformation 69]
239. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 204]
274. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 207]
275. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 207]
277. 0 <=> main_query_1(sK6,sK7) [avatar definition]
278. main_query_1(sK6,sK7) <- {0} [avatar component clause 277]
280. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
281. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 280]
283. 2 <=> main_query_2(sK6,sK7) [avatar definition]
284. main_query_2(sK6,sK7) <- {2} [avatar component clause 283]
286. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
287. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 286]
288. ~1 | ~3 [avatar split clause 275,286,280]
289. 0 | 2 [avatar split clause 274,283,277]
290. select_4(sK6,sK7) <- {2} [resolution 228,284]
291. from_5(sK6,sK7) <- {2} [resolution 290,226]
294. subquery_from_6(sK6,sK7) <- {2} [resolution 291,224]
297. select_7(sK6,sK7) <- {2} [resolution 294,222]
303. gte(X1,X1) [resolution 239,233]
305. where_3(X0) | ~lte(X0,X0) [resolution 212,303]
307. where_3(X0) [subsumption resolution 305,233]
327. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 213,208]
329. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1) [resolution 220,218]
342. from_8(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 327,219]
348. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1) [resolution 329,209]
356. ~select_1(X2,X3) | select_7(X2,X3) [resolution 342,221]
377. ~select_7(X2,X3) | ~where_3(X2) | select_1(X2,X3) [resolution 348,215]
383. ~select_7(X2,X3) | select_1(X2,X3) [subsumption resolution 377,307]
391. select_1(sK6,sK7) <- {2} [resolution 383,297]
400. main_query_1(sK6,sK7) <- {2} [resolution 391,217]
402. $false <- {1, 2} [subsumption resolution 400,281]
403. 1 | ~2 [avatar contradiction clause 402]
404. select_1(sK6,sK7) <- {0} [resolution 278,216]
405. select_7(sK6,sK7) <- {0} [resolution 404,356]
409. subquery_from_6(sK6,sK7) <- {0} [resolution 405,223]
417. from_5(sK6,sK7) <- {0} [resolution 409,225]
426. select_4(sK6,sK7) <- {0} [resolution 417,227]
435. main_query_2(sK6,sK7) <- {0} [resolution 426,229]
437. $false <- {0, 3} [subsumption resolution 435,287]
438. ~0 | 3 [avatar contradiction clause 437]
439. $false [avatar sat refutation 288,289,403,438]
% SZS output end Proof for query_06_query_08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
