% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_21_query_29_tptp4X
% SZS output start Proof for query_21_query_29_tptp4X
1. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) <=> (tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4))) [input]
2. ! [X0,X5] : (where_3(X0,X5) <=> eq(X0,X5)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_3(X0,X5) & from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [input]
4. ! [X10,X11] : (main_query_1(X10,X11) <=> select_1(X10,X11)) [input]
5. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) <=> (tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4))) [input]
6. ! [X5,X6,X7,X8,X9] : (from_9(X5,X6,X7,X8,X9) <=> tab2(X5,X6,X7,X8,X9)) [input]
7. ! [X5,X0] : (where_10(X5,X0) <=> eq(X5,X0)) [input]
8. ! [X5,X0] : (select_8(X5,X0) <=> ? [X6,X7,X8,X9] : (where_10(X5,X0) & from_9(X5,X6,X7,X8,X9))) [input]
9. ! [X12,X0] : (where_6_subquery_7(X12,X0) <=> select_8(X12,X0)) [input]
10. ! [X0] : (where_6(X0) <=> where_6_subquery_7(X0,X0)) [input]
11. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_6(X0) & from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [input]
12. ! [X10,X11] : (main_query_2(X10,X11) <=> select_4(X10,X11)) [input]
13. ! [X13] : eq(X13,X13) [input]
14. ! [X13,X14] : (eq(X13,X14) => eq(X14,X13)) [input]
27. ! [X13,X10,X11,X16,X17,X18] : ((tab1(X10,X11,X16,X17,X18) & eq(X10,X13)) => tab1(X13,X11,X16,X17,X18)) [input]
32. ! [X13,X5,X6,X7,X8,X9] : ((from_9(X5,X6,X7,X8,X9) & eq(X5,X13)) => from_9(X13,X6,X7,X8,X9)) [input]
74. ! [X10,X11] : (main_query_1(X10,X11) <=> main_query_2(X10,X11)) [input]
75. ~! [X10,X11] : (main_query_1(X10,X11) <=> main_query_2(X10,X11)) [negated conjecture 74]
76. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [rectify 2]
77. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
78. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab2(X0,X1,X2,X3,X4)) [rectify 6]
79. ! [X0,X1] : (where_10(X0,X1) <=> eq(X0,X1)) [rectify 7]
80. ! [X0,X1] : (select_8(X0,X1) <=> ? [X2,X3,X4,X5] : (where_10(X0,X1) & from_9(X0,X2,X3,X4,X5))) [rectify 8]
81. ! [X0,X1] : (where_6_subquery_7(X0,X1) <=> select_8(X0,X1)) [rectify 9]
82. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 12]
83. ! [X0] : eq(X0,X0) [rectify 13]
84. ! [X0,X1] : (eq(X0,X1) => eq(X1,X0)) [rectify 14]
97. ! [X0,X1,X2,X3,X4,X5] : ((tab1(X1,X2,X3,X4,X5) & eq(X1,X0)) => tab1(X0,X2,X3,X4,X5)) [rectify 27]
102. ! [X0,X1,X2,X3,X4,X5] : ((from_9(X1,X2,X3,X4,X5) & eq(X1,X0)) => from_9(X0,X2,X3,X4,X5)) [rectify 32]
144. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 75]
145. ! [X0,X1] : (eq(X1,X0) | ~eq(X0,X1)) [ennf transformation 84]
158. ! [X0,X1,X2,X3,X4,X5] : (tab1(X0,X2,X3,X4,X5) | (~tab1(X1,X2,X3,X4,X5) | ~eq(X1,X0))) [ennf transformation 97]
159. ! [X0,X1,X2,X3,X4,X5] : (tab1(X0,X2,X3,X4,X5) | ~tab1(X1,X2,X3,X4,X5) | ~eq(X1,X0)) [flattening 158]
168. ! [X0,X1,X2,X3,X4,X5] : (from_9(X0,X2,X3,X4,X5) | (~from_9(X1,X2,X3,X4,X5) | ~eq(X1,X0))) [ennf transformation 102]
169. ! [X0,X1,X2,X3,X4,X5] : (from_9(X0,X2,X3,X4,X5) | ~from_9(X1,X2,X3,X4,X5) | ~eq(X1,X0)) [flattening 168]
252. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 144]
253. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | (~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4))) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [nnf transformation 1]
254. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4)) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [flattening 253]
255. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 76]
256. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_3(X0,X5) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & (? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_3(X0,X5) & from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9)) | ~select_1(X0,X1))) [nnf transformation 3]
257. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_3(X0,X5) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_3(X0,X13) & from_2(X0,X1,X10,X11,X12,X13,X14,X15,X16,X17)) | ~select_1(X0,X1))) [rectify 256]
258. ! [X1,X0] : (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_3(X0,X13) & from_2(X0,X1,X10,X11,X12,X13,X14,X15,X16,X17)) => (where_3(X0,sK3(X0,X1)) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1),sK5(X0,X1),sK6(X0,X1),sK7(X0,X1)))) [choice axiom]
259. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_3(X0,X5) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & ((where_3(X0,sK3(X0,X1)) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1),sK5(X0,X1),sK6(X0,X1),sK7(X0,X1))) | ~select_1(X0,X1))) [skolemisation 257,258]
260. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 77]
261. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | (~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4))) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [nnf transformation 5]
262. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4)) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [flattening 261]
263. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab2(X0,X1,X2,X3,X4)) & (tab2(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 78]
264. ! [X0,X1] : ((where_10(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_10(X0,X1))) [nnf transformation 79]
265. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5] : (~where_10(X0,X1) | ~from_9(X0,X2,X3,X4,X5))) & (? [X2,X3,X4,X5] : (where_10(X0,X1) & from_9(X0,X2,X3,X4,X5)) | ~select_8(X0,X1))) [nnf transformation 80]
266. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5] : (~where_10(X0,X1) | ~from_9(X0,X2,X3,X4,X5))) & (? [X6,X7,X8,X9] : (where_10(X0,X1) & from_9(X0,X6,X7,X8,X9)) | ~select_8(X0,X1))) [rectify 265]
267. ! [X1,X0] : (? [X6,X7,X8,X9] : (where_10(X0,X1) & from_9(X0,X6,X7,X8,X9)) => (where_10(X0,X1) & from_9(X0,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1),sK11(X0,X1)))) [choice axiom]
268. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5] : (~where_10(X0,X1) | ~from_9(X0,X2,X3,X4,X5))) & ((where_10(X0,X1) & from_9(X0,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1),sK11(X0,X1))) | ~select_8(X0,X1))) [skolemisation 266,267]
269. ! [X0,X1] : ((where_6_subquery_7(X0,X1) | ~select_8(X0,X1)) & (select_8(X0,X1) | ~where_6_subquery_7(X0,X1))) [nnf transformation 81]
270. ! [X0] : ((where_6(X0) | ~where_6_subquery_7(X0,X0)) & (where_6_subquery_7(X0,X0) | ~where_6(X0))) [nnf transformation 10]
271. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & (? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_6(X0) & from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9)) | ~select_4(X0,X1))) [nnf transformation 11]
272. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_6(X0) & from_5(X0,X1,X10,X11,X12,X13,X14,X15,X16,X17)) | ~select_4(X0,X1))) [rectify 271]
273. ! [X1,X0] : (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_6(X0) & from_5(X0,X1,X10,X11,X12,X13,X14,X15,X16,X17)) => (where_6(X0) & from_5(X0,X1,sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1),sK16(X0,X1),sK17(X0,X1),sK18(X0,X1),sK19(X0,X1)))) [choice axiom]
274. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & ((where_6(X0) & from_5(X0,X1,sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1),sK16(X0,X1),sK17(X0,X1),sK18(X0,X1),sK19(X0,X1))) | ~select_4(X0,X1))) [skolemisation 272,273]
275. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 82]
278. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 252]
279. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK20,sK21) | ~main_query_1(sK20,sK21)) & (main_query_2(sK20,sK21) | main_query_1(sK20,sK21))) [choice axiom]
280. (~main_query_2(sK20,sK21) | ~main_query_1(sK20,sK21)) & (main_query_2(sK20,sK21) | main_query_1(sK20,sK21)) [skolemisation 278,279]
281. ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | tab1(X0,X1,X2,X3,X4) [cnf transformation 254]
282. ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | tab2(X5,X6,X7,X8,X9) [cnf transformation 254]
283. from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4) [cnf transformation 254]
284. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 255]
285. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 255]
286. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1),sK5(X0,X1),sK6(X0,X1),sK7(X0,X1)) | ~select_1(X0,X1) [cnf transformation 259]
287. where_3(X0,sK3(X0,X1)) | ~select_1(X0,X1) [cnf transformation 259]
288. ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~where_3(X0,X5) | select_1(X0,X1) [cnf transformation 259]
289. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 260]
290. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 260]
291. ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | tab1(X0,X1,X2,X3,X4) [cnf transformation 262]
293. from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4) [cnf transformation 262]
294. ~from_9(X0,X1,X2,X3,X4) | tab2(X0,X1,X2,X3,X4) [cnf transformation 263]
295. ~tab2(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 263]
297. ~eq(X0,X1) | where_10(X0,X1) [cnf transformation 264]
298. from_9(X0,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1),sK11(X0,X1)) | ~select_8(X0,X1) [cnf transformation 268]
300. ~from_9(X0,X2,X3,X4,X5) | ~where_10(X0,X1) | select_8(X0,X1) [cnf transformation 268]
301. ~where_6_subquery_7(X0,X1) | select_8(X0,X1) [cnf transformation 269]
302. ~select_8(X0,X1) | where_6_subquery_7(X0,X1) [cnf transformation 269]
303. where_6_subquery_7(X0,X0) | ~where_6(X0) [cnf transformation 270]
304. ~where_6_subquery_7(X0,X0) | where_6(X0) [cnf transformation 270]
305. from_5(X0,X1,sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1),sK16(X0,X1),sK17(X0,X1),sK18(X0,X1),sK19(X0,X1)) | ~select_4(X0,X1) [cnf transformation 274]
306. ~select_4(X0,X1) | where_6(X0) [cnf transformation 274]
307. ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~where_6(X0) | select_4(X0,X1) [cnf transformation 274]
308. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 275]
309. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 275]
310. eq(X0,X0) [cnf transformation 83]
311. ~eq(X0,X1) | eq(X1,X0) [cnf transformation 145]
322. ~tab1(X1,X2,X3,X4,X5) | tab1(X0,X2,X3,X4,X5) | ~eq(X1,X0) [cnf transformation 159]
327. ~from_9(X1,X2,X3,X4,X5) | from_9(X0,X2,X3,X4,X5) | ~eq(X1,X0) [cnf transformation 169]
369. main_query_2(sK20,sK21) | main_query_1(sK20,sK21) [cnf transformation 280]
370. ~main_query_2(sK20,sK21) | ~main_query_1(sK20,sK21) [cnf transformation 280]
372. 0 <=> main_query_1(sK20,sK21) [avatar definition]
373. main_query_1(sK20,sK21) <- {0} [avatar component clause 372]
375. 1 <=> ~main_query_1(sK20,sK21) [avatar definition]
376. ~main_query_1(sK20,sK21) <- {1} [avatar component clause 375]
378. 2 <=> main_query_2(sK20,sK21) [avatar definition]
379. main_query_2(sK20,sK21) <- {2} [avatar component clause 378]
381. 3 <=> ~main_query_2(sK20,sK21) [avatar definition]
382. ~main_query_2(sK20,sK21) <- {3} [avatar component clause 381]
383. ~1 | ~3 [avatar split clause 370,381,375]
384. 0 | 2 [avatar split clause 369,378,372]
386. where_3(X0,X0) [resolution 285,310]
388. where_10(X0,X0) [resolution 297,310]
390. select_8(X0,X0) | ~where_6(X0) [resolution 301,303]
399. eq(X0,sK3(X0,X1)) | ~select_1(X0,X1) [resolution 287,284]
404. eq(sK3(X7,X8),X7) | ~select_1(X7,X8) [resolution 399,311]
575. tab2(X0,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1),sK11(X0,X1)) | ~select_8(X0,X1) [resolution 298,294]
577. from_9(X7,sK8(X5,X6),sK9(X5,X6),sK10(X5,X6),sK11(X5,X6)) | ~select_8(X5,X6) | ~eq(X5,X7) [resolution 298,327]
585. ~tab2(X0,X1,X2,X3,X4) | ~tab1(X5,X6,X7,X8,X9) | ~where_3(X5,X0) | select_1(X5,X6) [resolution 283,288]
593. ~tab2(X20,X21,X22,X23,X24) | ~tab1(X25,X26,X27,X28,X29) | ~where_6(X25) | select_4(X25,X26) [resolution 293,307]
595. 4 <=> ! [X25,X27,X29,X26,X28] : (~tab1(X25,X26,X27,X28,X29) | select_4(X25,X26) | ~where_6(X25)) [avatar definition]
596. ~tab1(X25,X26,X27,X28,X29) | select_4(X25,X26) | ~where_6(X25) <- {4} [avatar component clause 595]
598. 6 <=> ! [X20,X22,X21,X23,X24] : ~tab2(X20,X21,X22,X23,X24) [avatar definition]
599. ~tab2(X20,X21,X22,X23,X24) <- {6} [avatar component clause 598]
600. 4 | 6 [avatar split clause 593,598,595]
720. select_4(sK20,sK21) <- {2} [resolution 379,308]
725. where_6(sK20) <- {2} [resolution 720,306]
727. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 286,281]
728. tab2(sK3(X2,X3),sK4(X2,X3),sK5(X2,X3),sK6(X2,X3),sK7(X2,X3)) | ~select_1(X2,X3) [resolution 286,282]
740. tab1(X0,X1,sK12(X0,X1),sK13(X0,X1),sK14(X0,X1)) | ~select_4(X0,X1) [resolution 305,291]
764. tab1(X16,X15,sK0(X14,X15),sK1(X14,X15),sK2(X14,X15)) | ~select_1(X14,X15) | ~eq(X14,X16) [resolution 727,322]
783. ~tab1(X0,X1,X2,X3,X4) | ~where_3(X0,X5) | select_1(X0,X1) | ~select_8(X5,X6) [resolution 585,575]
787. from_9(sK3(X0,X1),sK4(X0,X1),sK5(X0,X1),sK6(X0,X1),sK7(X0,X1)) | ~select_1(X0,X1) [resolution 728,295]
809. ~select_8(X3,X4) | ~eq(X3,X5) | ~where_10(X5,X6) | select_8(X5,X6) [resolution 577,300]
1382. ~where_10(X1,X2) | ~eq(X0,X1) | select_8(X1,X2) | ~where_6(X0) [resolution 809,390]
1385. ~eq(X0,X1) | select_8(X1,X1) | ~where_6(X0) [resolution 1382,388]
1397. ~select_4(X4,X6) | select_1(X4,X6) | ~select_8(X5,X7) | ~where_3(X4,X5) [resolution 783,740]
1404. ~where_6(sK3(X15,X16)) | select_8(X15,X15) | ~select_1(X15,X16) [resolution 1385,404]
1462. ~select_1(X20,X21) | ~eq(X20,X22) | select_4(X22,X21) | ~where_6(X22) <- {4} [resolution 764,596]
1527. ~where_10(sK3(X2,X3),X4) | ~select_1(X2,X3) | select_8(sK3(X2,X3),X4) [resolution 787,300]
1755. select_1(sK20,sK21) | ~select_8(X0,X1) | ~where_3(sK20,X0) <- {2} [resolution 1397,720]
1837. 18 <=> ! [X16,X17] : (~select_8(X16,X17) | ~where_3(sK20,X16)) [avatar definition]
1838. ~select_8(X16,X17) | ~where_3(sK20,X16) <- {18} [avatar component clause 1837]
1859. 28 <=> select_1(sK20,sK21) [avatar definition]
1860. select_1(sK20,sK21) <- {28} [avatar component clause 1859]
1861. 18 | 28 | ~2 [avatar split clause 1755,378,1859,1837]
1898. ~where_3(sK20,X0) | ~where_6(X0) <- {18} [resolution 1838,390]
1931. ~where_6(sK20) <- {18} [resolution 1898,386]
1937. $false <- {2, 18} [subsumption resolution 1931,725]
1938. ~2 | ~18 [avatar contradiction clause 1937]
1969. main_query_1(sK20,sK21) <- {28} [resolution 1860,290]
1990. $false <- {1, 28} [subsumption resolution 1969,376]
1991. 1 | ~28 [avatar contradiction clause 1990]
1996. select_1(sK20,sK21) <- {0} [resolution 373,289]
1997. 28 | ~0 [avatar split clause 1996,372,1859]
2017. ~eq(sK20,X11) | select_4(X11,sK21) | ~where_6(X11) <- {4, 28} [resolution 1860,1462]
2022. 30 <=> where_6(sK20) [avatar definition]
2025. 31 <=> ~where_6(sK20) [avatar definition]
2031. 32 <=> select_4(sK20,sK21) [avatar definition]
2032. select_4(sK20,sK21) <- {32} [avatar component clause 2031]
2296. select_4(sK20,sK21) | ~where_6(sK20) <- {4, 28} [resolution 2017,310]
3292. select_8(sK3(X14,X15),sK3(X14,X15)) | ~select_1(X14,X15) [resolution 1527,388]
3350. where_6_subquery_7(sK3(X2,X3),sK3(X2,X3)) | ~select_1(X2,X3) [resolution 3292,302]
3368. where_6(sK3(X0,X1)) | ~select_1(X0,X1) [resolution 3350,304]
3384. ~select_1(X0,X1) | select_8(X0,X0) | ~select_1(X0,X1) [resolution 3368,1404]
3385. ~select_1(X0,X1) | select_8(X0,X0) [duplicate literal removal 3384]
3398. select_8(sK20,sK20) <- {28} [resolution 3385,1860]
3422. where_6_subquery_7(sK20,sK20) <- {28} [resolution 3398,302]
3441. where_6(sK20) <- {28} [resolution 3422,304]
3447. 30 | ~28 [avatar split clause 3441,1859,2022]
3540. ~31 | 32 | ~4 | ~28 [avatar split clause 2296,1859,595,2031,2025]
3645. main_query_2(sK20,sK21) <- {32} [resolution 2032,309]
3647. $false <- {3, 32} [subsumption resolution 3645,382]
3648. 3 | ~32 [avatar contradiction clause 3647]
3651. ~select_8(X6,X7) <- {6} [resolution 599,575]
3677. $false <- {6, 28} [resolution 3651,3398]
3678. ~6 | ~28 [avatar contradiction clause 3677]
3679. $false [avatar sat refutation 383,384,600,1861,1938,1991,1997,3447,3540,3648,3678]
% SZS output end Proof for query_21_query_29_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 6908
% Time elapsed: 0.230 s
% ------------------------------
% ------------------------------
