% Refutation found. Thanks to Tanya!
% SZS status Theorem for query10_query14_tptp4X
% SZS output start Proof for query10_query14_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X5,X2,X3,X4] : (from_9(X0,X5,X2,X3,X4) <=> tab1(X0,X5,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X5,X2,X3,X4] : ($true & from_9(X0,X5,X2,X3,X4))) [input]
8. ! [X0] : (where_6_subquery_7(X0) <=> select_8(X0)) [input]
9. ! [X0,X1] : (where_6(X0,X1) <=> (where_6_subquery_7(X0) & eq(X0,X1))) [input]
10. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [input]
62. ! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [input]
63. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [negated conjecture 62]
64. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 6]
65. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [rectify 7]
66. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 65]
202. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 63]
203. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
204. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
205. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
206. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 205]
207. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
208. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 206,207]
209. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 4]
210. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
211. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 64]
212. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 66]
213. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 212]
214. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
215. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 213,214]
216. ! [X0] : ((where_6_subquery_7(X0) | ~select_8(X0)) & (select_8(X0) | ~where_6_subquery_7(X0))) [nnf transformation 8]
217. ! [X0,X1] : ((where_6(X0,X1) | (~where_6_subquery_7(X0) | ~eq(X0,X1))) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 9]
218. ! [X0,X1] : ((where_6(X0,X1) | ~where_6_subquery_7(X0) | ~eq(X0,X1)) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 217]
219. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 10]
220. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 219]
221. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
222. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 220,221]
223. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 11]
226. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 202]
227. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
228. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 226,227]
229. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 203]
230. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 203]
231. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 204]
232. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 204]
233. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 208]
234. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 208]
235. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 208]
236. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 209]
237. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 209]
238. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 210]
239. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 210]
241. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 211]
243. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 215]
245. ~select_8(X0) | where_6_subquery_7(X0) [cnf transformation 216]
246. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 218]
248. ~eq(X0,X1) | ~where_6_subquery_7(X0) | where_6(X0,X1) [cnf transformation 218]
249. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 222]
250. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 222]
251. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 222]
252. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 223]
253. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 223]
302. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 228]
303. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 228]
345. 0 <=> main_query_1(sK10,sK11) [avatar definition]
346. main_query_1(sK10,sK11) <- {0} [avatar component clause 345]
348. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
349. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 348]
351. 2 <=> main_query_2(sK10,sK11) [avatar definition]
352. main_query_2(sK10,sK11) <- {2} [avatar component clause 351]
354. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
355. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 354]
356. ~1 | ~3 [avatar split clause 303,354,348]
357. 0 | 2 [avatar split clause 302,351,345]
360. select_4(sK10,sK11) <- {2} [resolution 252,352]
361. where_6(sK10,sK11) <- {2} [resolution 360,250]
363. eq(sK10,sK11) <- {2} [resolution 361,246]
373. where_3(sK10,sK11) <- {2} [resolution 363,232]
421. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 233,229]
425. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 249,238]
426. from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 421,241]
427. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 421,239]
433. from_2(X4,X5,sK7(X4,X5),sK8(X4,X5),sK9(X4,X5)) | ~select_4(X4,X5) [resolution 425,230]
435. ~select_1(X2,X3) | select_8(X2) [resolution 426,243]
436. ~where_6(X0,X1) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 427,251]
443. ~select_4(X0,X1) | ~where_3(X0,X1) | select_1(X0,X1) [resolution 433,235]
445. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 443,360]
446. select_1(sK10,sK11) <- {2} [subsumption resolution 445,373]
448. main_query_1(sK10,sK11) <- {2} [resolution 446,237]
450. $false <- {1, 2} [subsumption resolution 448,349]
451. 1 | ~2 [avatar contradiction clause 450]
452. select_1(sK10,sK11) <- {0} [resolution 346,236]
453. select_8(sK10) <- {0} [resolution 452,435]
455. where_3(sK10,sK11) <- {0} [resolution 452,234]
456. where_6_subquery_7(sK10) <- {0} [resolution 453,245]
457. eq(sK10,sK11) <- {0} [resolution 455,231]
460. ~where_6_subquery_7(sK10) | where_6(sK10,sK11) <- {0} [resolution 457,248]
465. where_6(sK10,sK11) <- {0} [subsumption resolution 460,456]
476. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {0} [resolution 465,436]
479. select_4(sK10,sK11) <- {0} [subsumption resolution 476,452]
485. main_query_2(sK10,sK11) <- {0} [resolution 479,253]
487. $false <- {0, 3} [subsumption resolution 485,355]
488. ~0 | 3 [avatar contradiction clause 487]
489. $false [avatar sat refutation 356,357,451,488]
% SZS output end Proof for query10_query14_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
