% Refutation found. Thanks to Tanya!
% SZS status Theorem for query05_query07_tptp4X
% SZS output start Proof for query05_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X6,X1,X2,X3,X4] : (from_5(X6,X1,X2,X3,X4) <=> tab1(X6,X1,X2,X3,X4)) [input]
6. ! [X6] : (where_6(X6) <=> (~gt(X6,X6) & ~lt(X6,X6))) [input]
7. ! [X6] : (select_4(X6) <=> ? [X1,X2,X3,X4] : (where_6(X6) & from_5(X6,X1,X2,X3,X4))) [input]
8. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
16. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
17. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
40. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
41. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 40]
42. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
43. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
44. ! [X0] : (where_6(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [rectify 6]
45. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
46. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
47. ! [X0] : eq(X0,X0) [rectify 9]
54. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 16]
55. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 17]
78. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 41]
124. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 78]
125. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
126. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
127. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
128. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 127]
129. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
130. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 128,129]
131. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 42]
132. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 43]
133. ! [X0] : ((where_6(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [nnf transformation 44]
134. ! [X0] : ((where_6(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [flattening 133]
135. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 45]
136. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 135]
137. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
138. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 136,137]
139. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 46]
142. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 54]
143. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 142]
144. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 55]
145. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 124]
146. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
147. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 145,146]
148. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 125]
149. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 125]
151. ~eq(X0,X0) | where_3(X0) [cnf transformation 126]
152. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 130]
154. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 130]
155. ~main_query_1(X0) | select_1(X0) [cnf transformation 131]
156. ~select_1(X0) | main_query_1(X0) [cnf transformation 131]
157. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 132]
158. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 132]
161. gt(X0,X0) | where_6(X0) | lt(X0,X0) [cnf transformation 134]
162. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 138]
164. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 138]
165. ~main_query_2(X0) | select_4(X0) [cnf transformation 139]
166. ~select_4(X0) | main_query_2(X0) [cnf transformation 139]
167. eq(X0,X0) [cnf transformation 47]
177. ~lt(X0,X1) | ~eq(X0,X1) [cnf transformation 143]
179. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 144]
201. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 147]
202. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 147]
204. 0 <=> main_query_1(sK8) [avatar definition]
205. main_query_1(sK8) <- {0} [avatar component clause 204]
207. 1 <=> ~main_query_1(sK8) [avatar definition]
208. ~main_query_1(sK8) <- {1} [avatar component clause 207]
210. 2 <=> main_query_2(sK8) [avatar definition]
211. main_query_2(sK8) <- {2} [avatar component clause 210]
213. 3 <=> ~main_query_2(sK8) [avatar definition]
214. ~main_query_2(sK8) <- {3} [avatar component clause 213]
215. ~1 | ~3 [avatar split clause 202,213,207]
216. 0 | 2 [avatar split clause 201,210,204]
217. select_4(sK8) <- {2} [resolution 165,211]
221. where_3(X1) [resolution 151,167]
230. where_6(X1) | lt(X1,X1) | lt(X1,X1) [resolution 161,179]
231. lt(X1,X1) | where_6(X1) [duplicate literal removal 230]
234. where_6(X2) | ~eq(X2,X2) [resolution 231,177]
236. where_6(X2) [subsumption resolution 234,167]
271. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 152,148]
311. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 162,157]
433. from_5(X1,sK0(X1),sK1(X1),sK2(X1),sK3(X1)) | ~select_1(X1) [resolution 271,158]
440. from_2(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 311,149]
451. ~select_1(X5) | ~where_6(X5) | select_4(X5) [resolution 433,164]
452. ~select_1(X5) | select_4(X5) [subsumption resolution 451,236]
458. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 440,154]
464. ~select_4(X1) | select_1(X1) [subsumption resolution 458,221]
467. select_1(sK8) <- {2} [resolution 464,217]
477. main_query_1(sK8) <- {2} [resolution 467,156]
479. $false <- {1, 2} [subsumption resolution 477,208]
480. 1 | ~2 [avatar contradiction clause 479]
481. select_1(sK8) <- {0} [resolution 205,155]
482. select_4(sK8) <- {0} [resolution 481,452]
486. main_query_2(sK8) <- {0} [resolution 482,166]
488. $false <- {0, 3} [subsumption resolution 486,214]
489. ~0 | 3 [avatar contradiction clause 488]
490. $false [avatar sat refutation 215,216,480,489]
% SZS output end Proof for query05_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.016 s
% ------------------------------
% ------------------------------
