% Refutation found. Thanks to Tanya!
% SZS status Theorem for query07_query08_tptp4X
% SZS output start Proof for query07_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
7. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
12. ! [X7] : lte(X7,X7) [input]
15. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
16. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
18. ! [X7,X8] : (gte(X7,X8) <=> lte(X8,X7)) [input]
42. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
43. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 42]
44. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
45. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
46. ! [X0] : eq(X0,X0) [rectify 9]
49. ! [X0] : lte(X0,X0) [rectify 12]
52. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 15]
53. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 16]
55. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 18]
79. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 43]
131. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 79]
132. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
133. ! [X0] : ((where_3(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
134. ! [X0] : ((where_3(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [flattening 133]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 135]
137. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 136,137]
139. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 44]
140. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
141. ! [X0] : ((where_6(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [nnf transformation 6]
142. ! [X0] : ((where_6(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [flattening 141]
143. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 7]
144. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 143]
145. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
146. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 144,145]
147. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 45]
150. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 52]
151. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 150]
152. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 53]
153. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 55]
154. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 131]
155. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
156. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 154,155]
157. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 132]
158. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 132]
161. gt(X0,X0) | where_3(X0) | lt(X0,X0) [cnf transformation 134]
162. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 138]
164. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 138]
165. ~main_query_1(X0) | select_1(X0) [cnf transformation 139]
166. ~select_1(X0) | main_query_1(X0) [cnf transformation 139]
167. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 140]
168. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 140]
171. ~gte(X0,X0) | where_6(X0) | ~lte(X0,X0) [cnf transformation 142]
172. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 146]
174. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 146]
175. ~main_query_2(X0) | select_4(X0) [cnf transformation 147]
176. ~select_4(X0) | main_query_2(X0) [cnf transformation 147]
177. eq(X0,X0) [cnf transformation 46]
180. lte(X0,X0) [cnf transformation 49]
186. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 151]
188. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 152]
191. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 153]
215. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 156]
216. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 156]
241. 0 <=> main_query_1(sK8) [avatar definition]
242. main_query_1(sK8) <- {0} [avatar component clause 241]
244. 1 <=> ~main_query_1(sK8) [avatar definition]
245. ~main_query_1(sK8) <- {1} [avatar component clause 244]
247. 2 <=> main_query_2(sK8) [avatar definition]
248. main_query_2(sK8) <- {2} [avatar component clause 247]
250. 3 <=> ~main_query_2(sK8) [avatar definition]
251. ~main_query_2(sK8) <- {3} [avatar component clause 250]
252. ~1 | ~3 [avatar split clause 216,250,244]
253. 0 | 2 [avatar split clause 215,247,241]
254. select_4(sK8) <- {2} [resolution 175,248]
260. ~lt(X0,X0) [resolution 186,177]
263. gte(X1,X1) [resolution 191,180]
266. where_3(X1) | lt(X1,X1) | lt(X1,X1) [resolution 161,188]
267. where_3(X1) | lt(X1,X1) [duplicate literal removal 266]
268. where_3(X1) [subsumption resolution 267,260]
269. where_6(X0) | ~lte(X0,X0) [resolution 171,263]
271. where_6(X0) [subsumption resolution 269,180]
279. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 162,157]
281. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 172,167]
283. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 279,168]
286. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 281,158]
288. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 283,174]
289. ~select_1(X1) | select_4(X1) [subsumption resolution 288,271]
291. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 286,164]
292. ~select_4(X1) | select_1(X1) [subsumption resolution 291,268]
293. select_1(sK8) <- {2} [resolution 292,254]
295. main_query_1(sK8) <- {2} [resolution 293,166]
297. $false <- {1, 2} [subsumption resolution 295,245]
298. 1 | ~2 [avatar contradiction clause 297]
299. select_1(sK8) <- {0} [resolution 242,165]
300. select_4(sK8) <- {0} [resolution 299,289]
304. main_query_2(sK8) <- {0} [resolution 300,176]
306. $false <- {0, 3} [subsumption resolution 304,251]
307. ~0 | 3 [avatar contradiction clause 306]
308. $false [avatar sat refutation 252,253,298,307]
% SZS output end Proof for query07_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
