% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_04_query_08_tptp4X
% SZS output start Proof for query_04_query_08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X5,X6,X7,X8] : (from_4(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
5. ! [X0] : (where_5(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
6. ! [X0] : (select_3(X0) <=> ? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8))) [input]
7. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
8. ! [X9] : eq(X9,X9) [input]
15. ! [X9,X10] : (lt(X9,X10) <=> (~eq(X9,X10) & lte(X9,X10))) [input]
16. ! [X9,X10] : (gt(X9,X10) <=> lt(X10,X9)) [input]
39. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
40. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 39]
41. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
42. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
43. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
44. ! [X0] : eq(X0,X0) [rectify 8]
51. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 15]
52. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 16]
120. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 40]
121. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
122. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 41]
123. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 122]
124. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
125. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 123,124]
126. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
127. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 42]
128. ! [X0] : ((where_5(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [nnf transformation 5]
129. ! [X0] : ((where_5(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [flattening 128]
130. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 43]
131. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 130]
132. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
133. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 131,132]
134. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 7]
137. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 51]
138. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 137]
139. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 52]
140. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 120]
141. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
142. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 140,141]
143. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 121]
144. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 121]
145. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 125]
146. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 125]
147. ~main_query_1(X0) | select_1(X0) [cnf transformation 126]
148. ~select_1(X0) | main_query_1(X0) [cnf transformation 126]
149. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 127]
150. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 127]
153. gt(X0,X0) | where_5(X0) | lt(X0,X0) [cnf transformation 129]
154. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 133]
156. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 133]
157. ~main_query_2(X0) | select_3(X0) [cnf transformation 134]
158. ~select_3(X0) | main_query_2(X0) [cnf transformation 134]
159. eq(X0,X0) [cnf transformation 44]
169. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 138]
171. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 139]
193. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 142]
194. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 142]
196. 0 <=> main_query_1(sK8) [avatar definition]
197. main_query_1(sK8) <- {0} [avatar component clause 196]
199. 1 <=> ~main_query_1(sK8) [avatar definition]
200. ~main_query_1(sK8) <- {1} [avatar component clause 199]
202. 2 <=> main_query_2(sK8) [avatar definition]
203. main_query_2(sK8) <- {2} [avatar component clause 202]
205. 3 <=> ~main_query_2(sK8) [avatar definition]
206. ~main_query_2(sK8) <- {3} [avatar component clause 205]
207. ~1 | ~3 [avatar split clause 194,205,199]
208. 0 | 2 [avatar split clause 193,202,196]
209. select_3(sK8) <- {2} [resolution 157,203]
216. ~lt(X0,X0) [resolution 169,159]
218. where_5(X1) | lt(X1,X1) | lt(X1,X1) [resolution 153,171]
219. where_5(X1) | lt(X1,X1) [duplicate literal removal 218]
220. where_5(X1) [subsumption resolution 219,216]
244. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 145,143]
286. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 154,149]
405. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 244,150]
409. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 286,144]
412. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 405,156]
418. ~select_1(X1) | select_3(X1) [subsumption resolution 412,220]
422. ~select_3(X1) | select_1(X1) [resolution 409,146]
430. select_1(sK8) <- {2} [resolution 422,209]
434. main_query_1(sK8) <- {2} [resolution 430,148]
435. $false <- {1, 2} [subsumption resolution 434,200]
436. 1 | ~2 [avatar contradiction clause 435]
437. select_1(sK8) <- {0} [resolution 197,147]
440. select_3(sK8) <- {0} [resolution 437,418]
445. main_query_2(sK8) <- {0} [resolution 440,158]
447. $false <- {0, 3} [subsumption resolution 445,206]
448. ~0 | 3 [avatar contradiction clause 447]
449. $false [avatar sat refutation 207,208,436,448]
% SZS output end Proof for query_04_query_08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.035 s
% ------------------------------
% ------------------------------
