% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query03_tptp4X
% SZS output start Proof for query00_query03_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_4(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (select_3(X6) <=> ? [X7,X8,X9,X10] : ($true & from_4(X6,X7,X8,X9,X10))) [input]
6. ! [X6] : (main_query_2(X6) <=> select_3(X6)) [input]
38. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
39. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 38]
40. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
41. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
42. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
43. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [rectify 5]
44. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 43]
45. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 6]
77. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 39]
123. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 77]
124. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
125. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 40]
126. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 125]
127. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
128. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 126,127]
129. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 41]
130. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 42]
131. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 44]
132. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 131]
133. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
134. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0))) [skolemisation 132,133]
135. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 45]
138. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 123]
139. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
140. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 138,139]
141. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 124]
142. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 124]
143. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 128]
144. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 128]
145. ~main_query_1(X0) | select_1(X0) [cnf transformation 129]
146. ~select_1(X0) | main_query_1(X0) [cnf transformation 129]
147. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 130]
148. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 130]
149. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 134]
150. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 134]
151. ~main_query_2(X0) | select_3(X0) [cnf transformation 135]
152. ~select_3(X0) | main_query_2(X0) [cnf transformation 135]
182. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 140]
183. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 140]
185. 0 <=> main_query_1(sK8) [avatar definition]
186. main_query_1(sK8) <- {0} [avatar component clause 185]
188. 1 <=> ~main_query_1(sK8) [avatar definition]
189. ~main_query_1(sK8) <- {1} [avatar component clause 188]
191. 2 <=> main_query_2(sK8) [avatar definition]
192. main_query_2(sK8) <- {2} [avatar component clause 191]
194. 3 <=> ~main_query_2(sK8) [avatar definition]
195. ~main_query_2(sK8) <- {3} [avatar component clause 194]
196. ~1 | ~3 [avatar split clause 183,194,188]
197. 0 | 2 [avatar split clause 182,191,185]
198. select_3(sK8) <- {2} [resolution 151,192]
229. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 143,141]
231. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 149,147]
233. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 229,148]
236. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 231,142]
238. ~select_1(X1) | select_3(X1) [resolution 233,150]
240. ~select_3(X1) | select_1(X1) [resolution 236,144]
241. select_1(sK8) <- {2} [resolution 240,198]
245. main_query_1(sK8) <- {2} [resolution 241,146]
246. $false <- {1, 2} [subsumption resolution 245,189]
247. 1 | ~2 [avatar contradiction clause 246]
248. select_1(sK8) <- {0} [resolution 186,145]
251. select_3(sK8) <- {0} [resolution 248,238]
256. main_query_2(sK8) <- {0} [resolution 251,152]
257. $false <- {0, 3} [subsumption resolution 256,195]
258. ~0 | 3 [avatar contradiction clause 257]
259. $false [avatar sat refutation 196,197,247,258]
% SZS output end Proof for query00_query03_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
