% Refutation found. Thanks to Tanya!
% SZS status Theorem for query10_query12_tptp4X
% SZS output start Proof for query10_query12_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1] : (where_6(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [input]
7. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
12. ! [X7] : lte(X7,X7) [input]
13. ! [X7,X8] : ((lte(X8,X7) & lte(X7,X8)) <=> eq(X7,X8)) [input]
23. ! [X7,X0,X1] : ((where_6(X0,X1) & eq(X1,X7)) => where_6(X0,X7)) [input]
55. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
56. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 55]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
58. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
62. ! [X0] : lte(X0,X0) [rectify 12]
63. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
73. ! [X0,X1,X2] : ((where_6(X1,X2) & eq(X2,X0)) => where_6(X1,X0)) [rectify 23]
105. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 56]
117. ! [X0,X1,X2] : (where_6(X1,X0) | (~where_6(X1,X2) | ~eq(X2,X0))) [ennf transformation 73]
118. ! [X0,X1,X2] : (where_6(X1,X0) | ~where_6(X1,X2) | ~eq(X2,X0)) [flattening 117]
181. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 105]
182. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
183. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
184. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
185. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 184]
186. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
187. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 185,186]
188. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
189. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
190. ! [X0,X1] : ((where_6(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 6]
191. ! [X0,X1] : ((where_6(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [flattening 190]
192. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 7]
193. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 192]
194. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
195. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 193,194]
196. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 58]
197. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 63]
198. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 197]
199. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 181]
200. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
201. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 199,200]
202. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 182]
203. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 182]
204. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 183]
205. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 183]
206. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 187]
207. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 187]
208. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 187]
209. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 188]
210. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 188]
211. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 189]
212. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 189]
213. ~where_6(X0,X1) | lte(X0,X1) [cnf transformation 191]
214. ~where_6(X0,X1) | lte(X1,X0) [cnf transformation 191]
215. ~lte(X1,X0) | where_6(X0,X1) | ~lte(X0,X1) [cnf transformation 191]
216. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 195]
217. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 195]
218. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 195]
219. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 196]
220. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 196]
224. lte(X0,X0) [cnf transformation 62]
225. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 198]
233. ~where_6(X1,X2) | where_6(X1,X0) | ~eq(X2,X0) [cnf transformation 118]
265. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 201]
266. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 201]
268. 0 <=> main_query_1(sK6,sK7) [avatar definition]
269. main_query_1(sK6,sK7) <- {0} [avatar component clause 268]
271. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
272. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 271]
274. 2 <=> main_query_2(sK6,sK7) [avatar definition]
275. main_query_2(sK6,sK7) <- {2} [avatar component clause 274]
277. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
278. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 277]
279. ~1 | ~3 [avatar split clause 266,277,271]
280. 0 | 2 [avatar split clause 265,274,268]
283. select_4(sK6,sK7) <- {2} [resolution 219,275]
284. where_6(sK6,sK7) <- {2} [resolution 283,217]
286. lte(sK7,sK6) <- {2} [resolution 284,214]
287. lte(sK6,sK7) <- {2} [resolution 284,213]
292. where_6(X0,X0) | ~lte(X0,X0) [resolution 215,224]
297. where_6(X0,X0) [subsumption resolution 292,224]
307. eq(sK6,sK7) | ~lte(sK6,sK7) <- {2} [resolution 225,286]
309. eq(sK6,sK7) <- {2} [subsumption resolution 307,287]
315. where_3(sK6,sK7) <- {2} [resolution 309,205]
335. ~eq(X1,X2) | where_6(X1,X2) [resolution 233,297]
395. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 206,202]
399. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 216,211]
512. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 395,212]
529. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 399,203]
903. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 512,218]
913. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 529,208]
926. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 913,283]
930. select_1(sK6,sK7) <- {2} [subsumption resolution 926,315]
943. main_query_1(sK6,sK7) <- {2} [resolution 930,210]
945. $false <- {1, 2} [subsumption resolution 943,272]
946. 1 | ~2 [avatar contradiction clause 945]
949. select_1(sK6,sK7) <- {0} [resolution 269,209]
953. where_3(sK6,sK7) <- {0} [resolution 949,207]
954. eq(sK6,sK7) <- {0} [resolution 953,204]
964. where_6(sK6,sK7) <- {0} [resolution 954,335]
1023. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 964,903]
1028. select_4(sK6,sK7) <- {0} [subsumption resolution 1023,949]
1092. main_query_2(sK6,sK7) <- {0} [resolution 1028,220]
1094. $false <- {0, 3} [subsumption resolution 1092,278]
1095. ~0 | 3 [avatar contradiction clause 1094]
1096. $false [avatar sat refutation 279,280,946,1095]
% SZS output end Proof for query10_query12_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.024 s
% ------------------------------
% ------------------------------
