% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_11_query_13_tptp4X
% SZS output start Proof for query_11_query_13_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1] : (where_6(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [input]
7. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
12. ! [X7] : lte(X7,X7) [input]
13. ! [X7,X8] : ((lte(X8,X7) & lte(X7,X8)) <=> eq(X7,X8)) [input]
51. ! [X7,X0,X1] : ((where_6(X0,X1) & eq(X1,X7)) => where_6(X0,X7)) [input]
52. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
53. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 52]
54. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
55. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
59. ! [X0] : lte(X0,X0) [rectify 12]
60. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
98. ! [X0,X1,X2] : ((where_6(X1,X2) & eq(X2,X0)) => where_6(X1,X0)) [rectify 51]
99. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 53]
167. ! [X0,X1,X2] : (where_6(X1,X0) | (~where_6(X1,X2) | ~eq(X2,X0))) [ennf transformation 98]
168. ! [X0,X1,X2] : (where_6(X1,X0) | ~where_6(X1,X2) | ~eq(X2,X0)) [flattening 167]
169. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 99]
170. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
171. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
172. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
173. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 172]
174. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
175. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 173,174]
176. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 54]
177. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
178. ! [X0,X1] : ((where_6(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 6]
179. ! [X0,X1] : ((where_6(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [flattening 178]
180. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 7]
181. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 180]
182. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
183. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 181,182]
184. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 55]
185. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 60]
186. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 185]
187. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 169]
188. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
189. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 187,188]
190. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 170]
191. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 170]
192. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 171]
193. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 171]
194. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 175]
195. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 175]
196. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 175]
197. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 176]
198. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 176]
199. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 177]
200. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 177]
201. ~where_6(X0,X1) | lte(X0,X1) [cnf transformation 179]
202. ~where_6(X0,X1) | lte(X1,X0) [cnf transformation 179]
203. ~lte(X1,X0) | where_6(X0,X1) | ~lte(X0,X1) [cnf transformation 179]
204. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 183]
205. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 183]
206. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 183]
207. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 184]
208. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 184]
212. lte(X0,X0) [cnf transformation 59]
213. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 186]
249. ~where_6(X1,X2) | where_6(X1,X0) | ~eq(X2,X0) [cnf transformation 168]
250. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 189]
251. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 189]
253. 0 <=> main_query_1(sK6,sK7) [avatar definition]
254. main_query_1(sK6,sK7) <- {0} [avatar component clause 253]
256. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
257. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 256]
259. 2 <=> main_query_2(sK6,sK7) [avatar definition]
260. main_query_2(sK6,sK7) <- {2} [avatar component clause 259]
262. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
263. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 262]
264. ~1 | ~3 [avatar split clause 251,262,256]
265. 0 | 2 [avatar split clause 250,259,253]
268. select_4(sK6,sK7) <- {2} [resolution 207,260]
269. where_6(sK6,sK7) <- {2} [resolution 268,205]
271. lte(sK7,sK6) <- {2} [resolution 269,202]
272. lte(sK6,sK7) <- {2} [resolution 269,201]
277. where_6(X0,X0) | ~lte(X0,X0) [resolution 203,212]
282. where_6(X0,X0) [subsumption resolution 277,212]
292. eq(sK6,sK7) | ~lte(sK6,sK7) <- {2} [resolution 213,271]
294. eq(sK6,sK7) <- {2} [subsumption resolution 292,272]
300. where_3(sK6,sK7) <- {2} [resolution 294,193]
352. ~eq(X1,X2) | where_6(X1,X2) [resolution 249,282]
374. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 194,190]
378. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 204,199]
484. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 374,200]
503. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 378,191]
826. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 484,206]
836. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 503,196]
849. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 836,268]
853. select_1(sK6,sK7) <- {2} [subsumption resolution 849,300]
866. main_query_1(sK6,sK7) <- {2} [resolution 853,198]
868. $false <- {1, 2} [subsumption resolution 866,257]
869. 1 | ~2 [avatar contradiction clause 868]
872. select_1(sK6,sK7) <- {0} [resolution 254,197]
876. where_3(sK6,sK7) <- {0} [resolution 872,195]
877. eq(sK6,sK7) <- {0} [resolution 876,192]
887. where_6(sK6,sK7) <- {0} [resolution 877,352]
946. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 887,826]
951. select_4(sK6,sK7) <- {0} [subsumption resolution 946,872]
1017. main_query_2(sK6,sK7) <- {0} [resolution 951,208]
1019. $false <- {0, 3} [subsumption resolution 1017,263]
1020. ~0 | 3 [avatar contradiction clause 1019]
1021. $false [avatar sat refutation 264,265,869,1020]
% SZS output end Proof for query_11_query_13_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.028 s
% ------------------------------
% ------------------------------
