% Refutation found. Thanks to Tanya!
% SZS status Theorem for query06_query08_tptp4X
% SZS output start Proof for query06_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : ($true & from_5(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (subquery_from_3(X0) <=> select_4(X0)) [input]
4. ! [X0] : (from_2(X0) <=> subquery_from_3(X0)) [input]
5. ! [X0] : (select_1(X0) <=> ($true & from_2(X0))) [input]
6. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
7. ! [X0,X5,X6,X7,X8] : (from_7(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
8. ! [X0] : (where_8(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
9. ! [X0] : (select_6(X0) <=> ? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8))) [input]
10. ! [X0] : (main_query_2(X0) <=> select_6(X0)) [input]
14. ! [X9] : lte(X9,X9) [input]
20. ! [X9,X10] : (gte(X9,X10) <=> lte(X10,X9)) [input]
44. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
45. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 44]
46. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4)) [true and false elimination 2]
47. ! [X0] : (select_1(X0) <=> from_2(X0)) [true and false elimination 5]
48. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 7]
49. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4))) [rectify 9]
53. ! [X0] : lte(X0,X0) [rectify 14]
59. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 20]
134. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 45]
135. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 1]
136. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4) | ~select_4(X0))) [nnf transformation 46]
137. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) | ~select_4(X0))) [rectify 136]
138. ! [X0] : (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) => from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
139. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0))) [skolemisation 137,138]
140. ! [X0] : ((subquery_from_3(X0) | ~select_4(X0)) & (select_4(X0) | ~subquery_from_3(X0))) [nnf transformation 3]
141. ! [X0] : ((from_2(X0) | ~subquery_from_3(X0)) & (subquery_from_3(X0) | ~from_2(X0))) [nnf transformation 4]
142. ! [X0] : ((select_1(X0) | ~from_2(X0)) & (from_2(X0) | ~select_1(X0))) [nnf transformation 47]
143. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 6]
144. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 48]
145. ! [X0] : ((where_8(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_8(X0))) [nnf transformation 8]
146. ! [X0] : ((where_8(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_8(X0))) [flattening 145]
147. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4)) | ~select_6(X0))) [nnf transformation 49]
148. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) | ~select_6(X0))) [rectify 147]
149. ! [X0] : (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) => (where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
150. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & ((where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_6(X0))) [skolemisation 148,149]
151. ! [X0] : ((main_query_2(X0) | ~select_6(X0)) & (select_6(X0) | ~main_query_2(X0))) [nnf transformation 10]
154. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 59]
155. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 134]
156. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
157. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 155,156]
158. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 135]
159. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 135]
160. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [cnf transformation 139]
161. ~from_5(X0,X1,X2,X3,X4) | select_4(X0) [cnf transformation 139]
162. ~subquery_from_3(X0) | select_4(X0) [cnf transformation 140]
163. ~select_4(X0) | subquery_from_3(X0) [cnf transformation 140]
164. ~from_2(X0) | subquery_from_3(X0) [cnf transformation 141]
165. ~subquery_from_3(X0) | from_2(X0) [cnf transformation 141]
166. ~select_1(X0) | from_2(X0) [cnf transformation 142]
167. ~from_2(X0) | select_1(X0) [cnf transformation 142]
168. ~main_query_1(X0) | select_1(X0) [cnf transformation 143]
169. ~select_1(X0) | main_query_1(X0) [cnf transformation 143]
170. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 144]
171. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 144]
174. ~gte(X0,X0) | where_8(X0) | ~lte(X0,X0) [cnf transformation 146]
175. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 150]
177. ~from_7(X0,X1,X2,X3,X4) | ~where_8(X0) | select_6(X0) [cnf transformation 150]
178. ~main_query_2(X0) | select_6(X0) [cnf transformation 151]
179. ~select_6(X0) | main_query_2(X0) [cnf transformation 151]
183. lte(X0,X0) [cnf transformation 53]
189. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 154]
213. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 157]
214. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 157]
239. 0 <=> main_query_1(sK8) [avatar definition]
240. main_query_1(sK8) <- {0} [avatar component clause 239]
242. 1 <=> ~main_query_1(sK8) [avatar definition]
243. ~main_query_1(sK8) <- {1} [avatar component clause 242]
245. 2 <=> main_query_2(sK8) [avatar definition]
246. main_query_2(sK8) <- {2} [avatar component clause 245]
248. 3 <=> ~main_query_2(sK8) [avatar definition]
249. ~main_query_2(sK8) <- {3} [avatar component clause 248]
250. ~1 | ~3 [avatar split clause 214,248,242]
251. 0 | 2 [avatar split clause 213,245,239]
252. select_6(sK8) <- {2} [resolution 178,246]
260. gte(X1,X1) [resolution 189,183]
262. where_8(X0) | ~lte(X0,X0) [resolution 174,260]
264. where_8(X0) [subsumption resolution 262,183]
270. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [resolution 160,158]
272. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 175,170]
274. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [resolution 270,171]
277. from_5(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 272,159]
279. ~select_4(X1) | ~where_8(X1) | select_6(X1) [resolution 274,177]
280. ~select_4(X1) | select_6(X1) [subsumption resolution 279,264]
282. ~select_6(X1) | select_4(X1) [resolution 277,161]
283. select_4(sK8) <- {2} [resolution 282,252]
285. subquery_from_3(sK8) <- {2} [resolution 283,163]
286. from_2(sK8) <- {2} [resolution 285,165]
288. select_1(sK8) <- {2} [resolution 286,167]
290. main_query_1(sK8) <- {2} [resolution 288,169]
292. $false <- {1, 2} [subsumption resolution 290,243]
293. 1 | ~2 [avatar contradiction clause 292]
294. select_1(sK8) <- {0} [resolution 240,168]
296. from_2(sK8) <- {0} [resolution 294,166]
298. subquery_from_3(sK8) <- {0} [resolution 296,164]
300. select_4(sK8) <- {0} [resolution 298,162]
301. select_6(sK8) <- {0} [resolution 300,280]
304. main_query_2(sK8) <- {0} [resolution 301,179]
306. $false <- {0, 3} [subsumption resolution 304,249]
307. ~0 | 3 [avatar contradiction clause 306]
308. $false [avatar sat refutation 250,251,293,307]
% SZS output end Proof for query06_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.020 s
% ------------------------------
% ------------------------------
