% Refutation found. Thanks to Tanya!
% SZS status Theorem for query08_query09_tptp4X
% SZS output start Proof for query08_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X6,X7,X8,X9,X10] : (from_8(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
6. ! [X6] : (select_7(X6) <=> ? [X7,X8,X9,X10] : ($true & from_8(X6,X7,X8,X9,X10))) [input]
7. ! [X6] : (subquery_from_6(X6) <=> select_7(X6)) [input]
8. ! [X11] : (from_5(X11) <=> subquery_from_6(X11)) [input]
9. ! [X11] : (select_4(X11) <=> ($true & from_5(X11))) [input]
10. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
14. ! [X12] : lte(X12,X12) [input]
21. ! [X12,X13] : (gte(X12,X13) <=> lte(X13,X12)) [input]
42. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
43. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 42]
44. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
45. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
46. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
47. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 46]
48. ! [X0] : (subquery_from_6(X0) <=> select_7(X0)) [rectify 7]
49. ! [X0] : (from_5(X0) <=> subquery_from_6(X0)) [rectify 8]
50. ! [X0] : (select_4(X0) <=> ($true & from_5(X0))) [rectify 9]
51. ! [X0] : (select_4(X0) <=> from_5(X0)) [true and false elimination 50]
52. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 10]
56. ! [X0] : lte(X0,X0) [rectify 14]
63. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 21]
84. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 43]
130. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 84]
131. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
132. ! [X0] : ((where_3(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
133. ! [X0] : ((where_3(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_3(X0))) [flattening 132]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 134]
136. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 135,136]
138. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 44]
139. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 45]
140. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 47]
141. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 140]
142. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
143. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0))) [skolemisation 141,142]
144. ! [X0] : ((subquery_from_6(X0) | ~select_7(X0)) & (select_7(X0) | ~subquery_from_6(X0))) [nnf transformation 48]
145. ! [X0] : ((from_5(X0) | ~subquery_from_6(X0)) & (subquery_from_6(X0) | ~from_5(X0))) [nnf transformation 49]
146. ! [X0] : ((select_4(X0) | ~from_5(X0)) & (from_5(X0) | ~select_4(X0))) [nnf transformation 51]
147. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 52]
150. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 63]
151. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 130]
152. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
153. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 151,152]
154. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 131]
155. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 131]
158. ~gte(X0,X0) | where_3(X0) | ~lte(X0,X0) [cnf transformation 133]
159. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 137]
161. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 137]
162. ~main_query_1(X0) | select_1(X0) [cnf transformation 138]
163. ~select_1(X0) | main_query_1(X0) [cnf transformation 138]
164. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 139]
165. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 139]
166. from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [cnf transformation 143]
167. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 143]
168. ~subquery_from_6(X0) | select_7(X0) [cnf transformation 144]
169. ~select_7(X0) | subquery_from_6(X0) [cnf transformation 144]
170. ~from_5(X0) | subquery_from_6(X0) [cnf transformation 145]
171. ~subquery_from_6(X0) | from_5(X0) [cnf transformation 145]
172. ~select_4(X0) | from_5(X0) [cnf transformation 146]
173. ~from_5(X0) | select_4(X0) [cnf transformation 146]
174. ~main_query_2(X0) | select_4(X0) [cnf transformation 147]
175. ~select_4(X0) | main_query_2(X0) [cnf transformation 147]
179. lte(X0,X0) [cnf transformation 56]
186. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 150]
207. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 153]
208. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 153]
210. 0 <=> main_query_1(sK8) [avatar definition]
211. main_query_1(sK8) <- {0} [avatar component clause 210]
213. 1 <=> ~main_query_1(sK8) [avatar definition]
214. ~main_query_1(sK8) <- {1} [avatar component clause 213]
216. 2 <=> main_query_2(sK8) [avatar definition]
217. main_query_2(sK8) <- {2} [avatar component clause 216]
219. 3 <=> ~main_query_2(sK8) [avatar definition]
220. ~main_query_2(sK8) <- {3} [avatar component clause 219]
221. ~1 | ~3 [avatar split clause 208,219,213]
222. 0 | 2 [avatar split clause 207,216,210]
223. select_4(sK8) <- {2} [resolution 174,217]
224. from_5(sK8) <- {2} [resolution 223,172]
227. subquery_from_6(sK8) <- {2} [resolution 224,170]
229. select_7(sK8) <- {2} [resolution 227,168]
237. gte(X1,X1) [resolution 186,179]
242. where_3(X1) | ~lte(X1,X1) [resolution 237,158]
243. where_3(X1) [subsumption resolution 242,179]
288. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 159,154]
311. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [resolution 166,164]
313. from_8(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 288,165]
316. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_7(X1) [resolution 311,155]
318. ~select_1(X1) | select_7(X1) [resolution 313,167]
320. ~select_7(X1) | ~where_3(X1) | select_1(X1) [resolution 316,161]
321. ~select_7(X1) | select_1(X1) [subsumption resolution 320,243]
324. select_1(sK8) <- {2} [resolution 321,229]
328. main_query_1(sK8) <- {2} [resolution 324,163]
330. $false <- {1, 2} [subsumption resolution 328,214]
331. 1 | ~2 [avatar contradiction clause 330]
332. select_1(sK8) <- {0} [resolution 211,162]
335. select_7(sK8) <- {0} [resolution 332,318]
341. subquery_from_6(sK8) <- {0} [resolution 335,169]
344. from_5(sK8) <- {0} [resolution 341,171]
348. select_4(sK8) <- {0} [resolution 344,173]
352. main_query_2(sK8) <- {0} [resolution 348,175]
354. $false <- {0, 3} [subsumption resolution 352,220]
355. ~0 | 3 [avatar contradiction clause 354]
356. $false [avatar sat refutation 221,222,331,355]
% SZS output end Proof for query08_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.013 s
% ------------------------------
% ------------------------------
