% Refutation found. Thanks to Tanya!
% SZS status Theorem for query01_query04_tptp4X
% SZS output start Proof for query01_query04_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [input]
6. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
8. ! [X5] : eq(X5,X5) [input]
42. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
43. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 42]
44. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
45. ! [X0] : eq(X0,X0) [rectify 8]
130. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 43]
131. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
132. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 44]
133. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 132]
134. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 133,134]
136. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
137. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
138. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 5]
139. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 6]
140. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 139]
141. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
142. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 140,141]
143. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 7]
146. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 130]
147. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
148. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 146,147]
149. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 131]
150. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 131]
151. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 135]
152. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 135]
153. ~main_query_1(X0) | select_1(X0) [cnf transformation 136]
154. ~select_1(X0) | main_query_1(X0) [cnf transformation 136]
155. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 137]
156. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 137]
158. ~eq(X0,X0) | where_5(X0) [cnf transformation 138]
159. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 142]
161. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 142]
162. ~main_query_2(X0) | select_3(X0) [cnf transformation 143]
163. ~select_3(X0) | main_query_2(X0) [cnf transformation 143]
164. eq(X0,X0) [cnf transformation 45]
196. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 148]
197. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 148]
199. 0 <=> main_query_1(sK8) [avatar definition]
200. main_query_1(sK8) <- {0} [avatar component clause 199]
202. 1 <=> ~main_query_1(sK8) [avatar definition]
203. ~main_query_1(sK8) <- {1} [avatar component clause 202]
205. 2 <=> main_query_2(sK8) [avatar definition]
206. main_query_2(sK8) <- {2} [avatar component clause 205]
208. 3 <=> ~main_query_2(sK8) [avatar definition]
209. ~main_query_2(sK8) <- {3} [avatar component clause 208]
210. ~1 | ~3 [avatar split clause 197,208,202]
211. 0 | 2 [avatar split clause 196,205,199]
212. select_3(sK8) <- {2} [resolution 162,206]
216. where_5(X1) [resolution 158,164]
250. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 151,149]
252. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 159,155]
254. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 250,156]
259. from_2(X3,sK4(X3),sK5(X3),sK6(X3),sK7(X3)) | ~select_3(X3) [resolution 252,150]
263. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 254,161]
264. ~select_1(X1) | select_3(X1) [subsumption resolution 263,216]
270. ~select_3(X1) | select_1(X1) [resolution 259,152]
273. select_1(sK8) <- {2} [resolution 270,212]
275. main_query_1(sK8) <- {2} [resolution 273,154]
276. $false <- {1, 2} [subsumption resolution 275,203]
277. 1 | ~2 [avatar contradiction clause 276]
278. select_1(sK8) <- {0} [resolution 200,153]
279. select_3(sK8) <- {0} [resolution 278,264]
282. main_query_2(sK8) <- {0} [resolution 279,163]
284. $false <- {0, 3} [subsumption resolution 282,209]
285. ~0 | 3 [avatar contradiction clause 284]
286. $false [avatar sat refutation 210,211,277,285]
% SZS output end Proof for query01_query04_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.012 s
% ------------------------------
% ------------------------------
