% Refutation found. Thanks to Tanya!
% SZS status Theorem for query01_query07_tptp4X
% SZS output start Proof for query01_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
6. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
8. ! [X5] : eq(X5,X5) [input]
14. ! [X5,X6] : (lt(X5,X6) <=> (~eq(X5,X6) & lte(X5,X6))) [input]
15. ! [X5,X6] : (gt(X5,X6) <=> lt(X6,X5)) [input]
41. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
42. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0] : eq(X0,X0) [rectify 8]
50. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 14]
51. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 15]
128. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 42]
129. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
130. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
131. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 130]
132. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
133. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 131,132]
134. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
135. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
136. ! [X0] : ((where_5(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [nnf transformation 5]
137. ! [X0] : ((where_5(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [flattening 136]
138. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 6]
139. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 138]
140. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
141. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 139,140]
142. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 7]
145. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 50]
146. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 145]
147. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 51]
148. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 128]
149. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
150. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 148,149]
151. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 129]
152. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 129]
153. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 133]
154. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 133]
155. ~main_query_1(X0) | select_1(X0) [cnf transformation 134]
156. ~select_1(X0) | main_query_1(X0) [cnf transformation 134]
157. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 135]
158. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 135]
161. gt(X0,X0) | where_5(X0) | lt(X0,X0) [cnf transformation 137]
162. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 141]
164. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 141]
165. ~main_query_2(X0) | select_3(X0) [cnf transformation 142]
166. ~select_3(X0) | main_query_2(X0) [cnf transformation 142]
167. eq(X0,X0) [cnf transformation 44]
176. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 146]
178. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 147]
203. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 150]
204. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 150]
229. 0 <=> main_query_1(sK8) [avatar definition]
230. main_query_1(sK8) <- {0} [avatar component clause 229]
232. 1 <=> ~main_query_1(sK8) [avatar definition]
233. ~main_query_1(sK8) <- {1} [avatar component clause 232]
235. 2 <=> main_query_2(sK8) [avatar definition]
236. main_query_2(sK8) <- {2} [avatar component clause 235]
238. 3 <=> ~main_query_2(sK8) [avatar definition]
239. ~main_query_2(sK8) <- {3} [avatar component clause 238]
240. ~1 | ~3 [avatar split clause 204,238,232]
241. 0 | 2 [avatar split clause 203,235,229]
242. select_3(sK8) <- {2} [resolution 165,236]
248. ~lt(X0,X0) [resolution 176,167]
250. where_5(X1) | lt(X1,X1) | lt(X1,X1) [resolution 161,178]
251. where_5(X1) | lt(X1,X1) [duplicate literal removal 250]
252. where_5(X1) [subsumption resolution 251,248]
257. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 153,151]
259. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 162,157]
261. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 257,158]
264. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 259,152]
266. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 261,164]
267. ~select_1(X1) | select_3(X1) [subsumption resolution 266,252]
269. ~select_3(X1) | select_1(X1) [resolution 264,154]
270. select_1(sK8) <- {2} [resolution 269,242]
272. main_query_1(sK8) <- {2} [resolution 270,156]
273. $false <- {1, 2} [subsumption resolution 272,233]
274. 1 | ~2 [avatar contradiction clause 273]
275. select_1(sK8) <- {0} [resolution 230,155]
276. select_3(sK8) <- {0} [resolution 275,267]
279. main_query_2(sK8) <- {0} [resolution 276,166]
281. $false <- {0, 3} [subsumption resolution 279,239]
282. ~0 | 3 [avatar contradiction clause 281]
283. $false [avatar sat refutation 240,241,274,282]
% SZS output end Proof for query01_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.010 s
% ------------------------------
% ------------------------------
