% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query02_tptp4X
% SZS output start Proof for query00_query02_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_4(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (select_3(X6) <=> ? [X7,X8,X9,X10] : ($true & from_4(X6,X7,X8,X9,X10))) [input]
6. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
44. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
45. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 44]
46. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
47. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
48. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
49. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [rectify 5]
50. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 49]
51. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 6]
89. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 45]
149. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 89]
150. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
151. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 46]
152. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 151]
153. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
154. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 152,153]
155. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 47]
156. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 48]
157. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 50]
158. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 157]
159. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
160. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0))) [skolemisation 158,159]
161. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 51]
164. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 149]
165. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
166. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 164,165]
167. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 150]
168. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 150]
169. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 154]
170. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 154]
171. ~main_query_1(X0) | select_1(X0) [cnf transformation 155]
172. ~select_1(X0) | main_query_1(X0) [cnf transformation 155]
173. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 156]
174. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 156]
175. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 160]
176. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 160]
177. ~main_query_2(X0) | select_3(X0) [cnf transformation 161]
178. ~select_3(X0) | main_query_2(X0) [cnf transformation 161]
214. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 166]
215. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 166]
244. 0 <=> main_query_1(sK8) [avatar definition]
245. main_query_1(sK8) <- {0} [avatar component clause 244]
247. 1 <=> ~main_query_1(sK8) [avatar definition]
248. ~main_query_1(sK8) <- {1} [avatar component clause 247]
250. 2 <=> main_query_2(sK8) [avatar definition]
251. main_query_2(sK8) <- {2} [avatar component clause 250]
253. 3 <=> ~main_query_2(sK8) [avatar definition]
254. ~main_query_2(sK8) <- {3} [avatar component clause 253]
255. ~1 | ~3 [avatar split clause 215,253,247]
256. 0 | 2 [avatar split clause 214,250,244]
257. select_3(sK8) <- {2} [resolution 177,251]
265. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 169,167]
267. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 175,173]
269. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 265,174]
272. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 267,168]
274. ~select_1(X1) | select_3(X1) [resolution 269,176]
276. ~select_3(X1) | select_1(X1) [resolution 272,170]
277. select_1(sK8) <- {2} [resolution 276,257]
279. main_query_1(sK8) <- {2} [resolution 277,172]
280. $false <- {1, 2} [subsumption resolution 279,248]
281. 1 | ~2 [avatar contradiction clause 280]
282. select_1(sK8) <- {0} [resolution 245,171]
283. select_3(sK8) <- {0} [resolution 282,274]
286. main_query_2(sK8) <- {0} [resolution 283,178]
287. $false <- {0, 3} [subsumption resolution 286,254]
288. ~0 | 3 [avatar contradiction clause 287]
289. $false [avatar sat refutation 255,256,281,288]
% SZS output end Proof for query00_query02_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
