% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_11_query_13_tptp4X
% SZS output start Proof for query_11_query_13_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1] : (where_6(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [input]
7. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
12. ! [X7] : lte(X7,X7) [input]
13. ! [X7,X8] : ((lte(X8,X7) & lte(X7,X8)) <=> eq(X7,X8)) [input]
34. ! [X7,X0,X1] : ((where_6(X0,X1) & eq(X1,X7)) => where_6(X0,X7)) [input]
51. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
52. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 51]
53. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
54. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
58. ! [X0] : lte(X0,X0) [rectify 12]
59. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
80. ! [X0,X1,X2] : ((where_6(X1,X2) & eq(X2,X0)) => where_6(X1,X0)) [rectify 34]
97. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 52]
133. ! [X0,X1,X2] : (where_6(X1,X0) | (~where_6(X1,X2) | ~eq(X2,X0))) [ennf transformation 80]
134. ! [X0,X1,X2] : (where_6(X1,X0) | ~where_6(X1,X2) | ~eq(X2,X0)) [flattening 133]
167. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 97]
168. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
169. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
170. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
171. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 170]
172. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
173. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 171,172]
174. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 53]
175. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
176. ! [X0,X1] : ((where_6(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 6]
177. ! [X0,X1] : ((where_6(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [flattening 176]
178. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 7]
179. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 178]
180. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
181. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 179,180]
182. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 54]
183. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 59]
184. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 183]
185. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 167]
186. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
187. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 185,186]
188. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 168]
189. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 168]
190. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 169]
191. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 169]
192. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 173]
193. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 173]
194. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 173]
195. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 174]
196. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 174]
197. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 175]
198. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 175]
199. ~where_6(X0,X1) | lte(X0,X1) [cnf transformation 177]
200. ~where_6(X0,X1) | lte(X1,X0) [cnf transformation 177]
201. ~lte(X1,X0) | where_6(X0,X1) | ~lte(X0,X1) [cnf transformation 177]
202. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 181]
203. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 181]
204. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 181]
205. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 182]
206. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 182]
210. lte(X0,X0) [cnf transformation 58]
211. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 184]
230. ~where_6(X1,X2) | where_6(X1,X0) | ~eq(X2,X0) [cnf transformation 134]
247. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 187]
248. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 187]
250. 0 <=> main_query_1(sK6,sK7) [avatar definition]
251. main_query_1(sK6,sK7) <- {0} [avatar component clause 250]
253. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
254. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 253]
256. 2 <=> main_query_2(sK6,sK7) [avatar definition]
257. main_query_2(sK6,sK7) <- {2} [avatar component clause 256]
259. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
260. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 259]
261. ~1 | ~3 [avatar split clause 248,259,253]
262. 0 | 2 [avatar split clause 247,256,250]
265. select_4(sK6,sK7) <- {2} [resolution 205,257]
266. where_6(sK6,sK7) <- {2} [resolution 265,203]
268. lte(sK7,sK6) <- {2} [resolution 266,200]
269. lte(sK6,sK7) <- {2} [resolution 266,199]
273. where_6(X0,X0) | ~lte(X0,X0) [resolution 201,210]
276. where_6(X0,X0) [subsumption resolution 273,210]
282. eq(sK6,sK7) | ~lte(sK6,sK7) <- {2} [resolution 211,268]
284. eq(sK6,sK7) <- {2} [subsumption resolution 282,269]
295. where_3(sK6,sK7) <- {2} [resolution 284,191]
318. ~eq(X1,X2) | where_6(X1,X2) [resolution 230,276]
358. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 192,188]
362. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 202,197]
417. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 358,198]
423. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 362,189]
431. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 417,204]
462. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 423,194]
475. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 462,265]
479. select_1(sK6,sK7) <- {2} [subsumption resolution 475,295]
492. main_query_1(sK6,sK7) <- {2} [resolution 479,196]
494. $false <- {1, 2} [subsumption resolution 492,254]
495. 1 | ~2 [avatar contradiction clause 494]
498. select_1(sK6,sK7) <- {0} [resolution 251,195]
502. where_3(sK6,sK7) <- {0} [resolution 498,193]
503. eq(sK6,sK7) <- {0} [resolution 502,190]
513. where_6(sK6,sK7) <- {0} [resolution 503,318]
578. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 513,431]
583. select_4(sK6,sK7) <- {0} [subsumption resolution 578,498]
619. main_query_2(sK6,sK7) <- {0} [resolution 583,206]
621. $false <- {0, 3} [subsumption resolution 619,260]
622. ~0 | 3 [avatar contradiction clause 621]
623. $false [avatar sat refutation 261,262,495,622]
% SZS output end Proof for query_11_query_13_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.015 s
% ------------------------------
% ------------------------------
