% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_15_query_28_tptp4X
% SZS output start Proof for query_15_query_28_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1,X2,X3,X4] : (from_6(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
3. ! [X0] : (select_5(X0) <=> ? [X1,X2,X3,X4] : ($true & from_6(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (where_3_subquery_4(X5) <=> select_5(X5)) [input]
5. ! [X0,X1] : (where_3(X0,X1) <=> (where_3_subquery_4(X0) & eq(X0,X1))) [input]
6. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
7. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
8. ! [X0,X7,X2,X3,X4] : (from_8(X0,X7,X2,X3,X4) <=> tab1(X0,X7,X2,X3,X4)) [input]
9. ! [X0,X1,X2,X3,X4] : (from_12(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
10. ! [X0,X7] : (where_13(X0,X7) <=> eq(X0,X7)) [input]
11. ! [X0,X7] : (select_11(X0,X7) <=> ? [X1,X2,X3,X4] : (where_13(X0,X7) & from_12(X0,X1,X2,X3,X4))) [input]
12. ! [X5,X7] : (where_9_subquery_10(X5,X7) <=> select_11(X5,X7)) [input]
13. ! [X0,X7] : (where_9(X0,X7) <=> where_9_subquery_10(X0,X7)) [input]
14. ! [X0,X7] : (select_7(X0,X7) <=> ? [X2,X3,X4] : (where_9(X0,X7) & from_8(X0,X7,X2,X3,X4))) [input]
15. ! [X5,X7] : (main_query_2(X5,X7) <=> select_7(X5,X7)) [input]
16. ! [X8] : eq(X8,X8) [input]
29. ! [X8,X0,X7] : ((where_9(X0,X7) & eq(X7,X8)) => where_9(X0,X8)) [input]
74. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
75. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 74]
76. ! [X0] : (select_5(X0) <=> ? [X1,X2,X3,X4] : from_6(X0,X1,X2,X3,X4)) [true and false elimination 3]
77. ! [X0] : (where_3_subquery_4(X0) <=> select_5(X0)) [rectify 4]
78. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 7]
79. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 8]
80. ! [X0,X1] : (where_13(X0,X1) <=> eq(X0,X1)) [rectify 10]
81. ! [X0,X1] : (select_11(X0,X1) <=> ? [X2,X3,X4,X5] : (where_13(X0,X1) & from_12(X0,X2,X3,X4,X5))) [rectify 11]
82. ! [X0,X1] : (where_9_subquery_10(X0,X1) <=> select_11(X0,X1)) [rectify 12]
83. ! [X0,X1] : (where_9(X0,X1) <=> where_9_subquery_10(X0,X1)) [rectify 13]
84. ! [X0,X1] : (select_7(X0,X1) <=> ? [X2,X3,X4] : (where_9(X0,X1) & from_8(X0,X1,X2,X3,X4))) [rectify 14]
85. ! [X0,X1] : (main_query_2(X0,X1) <=> select_7(X0,X1)) [rectify 15]
86. ! [X0] : eq(X0,X0) [rectify 16]
99. ! [X0,X1,X2] : ((where_9(X1,X2) & eq(X2,X0)) => where_9(X1,X0)) [rectify 29]
144. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 75]
156. ! [X0,X1,X2] : (where_9(X1,X0) | (~where_9(X1,X2) | ~eq(X2,X0))) [ennf transformation 99]
157. ! [X0,X1,X2] : (where_9(X1,X0) | ~where_9(X1,X2) | ~eq(X2,X0)) [flattening 156]
246. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 144]
247. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
248. ! [X0,X1,X2,X3,X4] : ((from_6(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_6(X0,X1,X2,X3,X4))) [nnf transformation 2]
249. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_6(X0,X1,X2,X3,X4) | ~select_5(X0))) [nnf transformation 76]
250. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_6(X0,X5,X6,X7,X8) | ~select_5(X0))) [rectify 249]
251. ! [X0] : (? [X5,X6,X7,X8] : from_6(X0,X5,X6,X7,X8) => from_6(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
252. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (from_6(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_5(X0))) [skolemisation 250,251]
253. ! [X0] : ((where_3_subquery_4(X0) | ~select_5(X0)) & (select_5(X0) | ~where_3_subquery_4(X0))) [nnf transformation 77]
254. ! [X0,X1] : ((where_3(X0,X1) | (~where_3_subquery_4(X0) | ~eq(X0,X1))) & ((where_3_subquery_4(X0) & eq(X0,X1)) | ~where_3(X0,X1))) [nnf transformation 5]
255. ! [X0,X1] : ((where_3(X0,X1) | ~where_3_subquery_4(X0) | ~eq(X0,X1)) & ((where_3_subquery_4(X0) & eq(X0,X1)) | ~where_3(X0,X1))) [flattening 254]
256. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 6]
257. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 256]
258. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1)))) [choice axiom]
259. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1))) | ~select_1(X0,X1))) [skolemisation 257,258]
260. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 78]
261. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 79]
262. ! [X0,X1,X2,X3,X4] : ((from_12(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_12(X0,X1,X2,X3,X4))) [nnf transformation 9]
263. ! [X0,X1] : ((where_13(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_13(X0,X1))) [nnf transformation 80]
264. ! [X0,X1] : ((select_11(X0,X1) | ! [X2,X3,X4,X5] : (~where_13(X0,X1) | ~from_12(X0,X2,X3,X4,X5))) & (? [X2,X3,X4,X5] : (where_13(X0,X1) & from_12(X0,X2,X3,X4,X5)) | ~select_11(X0,X1))) [nnf transformation 81]
265. ! [X0,X1] : ((select_11(X0,X1) | ! [X2,X3,X4,X5] : (~where_13(X0,X1) | ~from_12(X0,X2,X3,X4,X5))) & (? [X6,X7,X8,X9] : (where_13(X0,X1) & from_12(X0,X6,X7,X8,X9)) | ~select_11(X0,X1))) [rectify 264]
266. ! [X1,X0] : (? [X6,X7,X8,X9] : (where_13(X0,X1) & from_12(X0,X6,X7,X8,X9)) => (where_13(X0,X1) & from_12(X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)))) [choice axiom]
267. ! [X0,X1] : ((select_11(X0,X1) | ! [X2,X3,X4,X5] : (~where_13(X0,X1) | ~from_12(X0,X2,X3,X4,X5))) & ((where_13(X0,X1) & from_12(X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1))) | ~select_11(X0,X1))) [skolemisation 265,266]
268. ! [X0,X1] : ((where_9_subquery_10(X0,X1) | ~select_11(X0,X1)) & (select_11(X0,X1) | ~where_9_subquery_10(X0,X1))) [nnf transformation 82]
269. ! [X0,X1] : ((where_9(X0,X1) | ~where_9_subquery_10(X0,X1)) & (where_9_subquery_10(X0,X1) | ~where_9(X0,X1))) [nnf transformation 83]
270. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : (~where_9(X0,X1) | ~from_8(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_9(X0,X1) & from_8(X0,X1,X2,X3,X4)) | ~select_7(X0,X1))) [nnf transformation 84]
271. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : (~where_9(X0,X1) | ~from_8(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_9(X0,X1) & from_8(X0,X1,X5,X6,X7)) | ~select_7(X0,X1))) [rectify 270]
272. ! [X1,X0] : (? [X5,X6,X7] : (where_9(X0,X1) & from_8(X0,X1,X5,X6,X7)) => (where_9(X0,X1) & from_8(X0,X1,sK11(X0,X1),sK12(X0,X1),sK13(X0,X1)))) [choice axiom]
273. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : (~where_9(X0,X1) | ~from_8(X0,X1,X2,X3,X4))) & ((where_9(X0,X1) & from_8(X0,X1,sK11(X0,X1),sK12(X0,X1),sK13(X0,X1))) | ~select_7(X0,X1))) [skolemisation 271,272]
274. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_7(X0,X1)) & (select_7(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 85]
277. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 246]
278. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK14,sK15) | ~main_query_1(sK14,sK15)) & (main_query_2(sK14,sK15) | main_query_1(sK14,sK15))) [choice axiom]
279. (~main_query_2(sK14,sK15) | ~main_query_1(sK14,sK15)) & (main_query_2(sK14,sK15) | main_query_1(sK14,sK15)) [skolemisation 277,278]
280. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 247]
281. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 247]
282. ~from_6(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 248]
283. ~tab1(X0,X1,X2,X3,X4) | from_6(X0,X1,X2,X3,X4) [cnf transformation 248]
284. from_6(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_5(X0) [cnf transformation 252]
285. ~from_6(X0,X1,X2,X3,X4) | select_5(X0) [cnf transformation 252]
287. ~select_5(X0) | where_3_subquery_4(X0) [cnf transformation 253]
288. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 255]
290. ~eq(X0,X1) | ~where_3_subquery_4(X0) | where_3(X0,X1) [cnf transformation 255]
291. from_2(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1)) | ~select_1(X0,X1) [cnf transformation 259]
292. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 259]
293. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 259]
294. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 260]
295. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 260]
296. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 261]
297. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 261]
299. ~tab1(X0,X1,X2,X3,X4) | from_12(X0,X1,X2,X3,X4) [cnf transformation 262]
300. ~where_13(X0,X1) | eq(X0,X1) [cnf transformation 263]
301. ~eq(X0,X1) | where_13(X0,X1) [cnf transformation 263]
303. ~select_11(X0,X1) | where_13(X0,X1) [cnf transformation 267]
304. ~from_12(X0,X2,X3,X4,X5) | ~where_13(X0,X1) | select_11(X0,X1) [cnf transformation 267]
305. ~where_9_subquery_10(X0,X1) | select_11(X0,X1) [cnf transformation 268]
306. ~select_11(X0,X1) | where_9_subquery_10(X0,X1) [cnf transformation 268]
307. ~where_9(X0,X1) | where_9_subquery_10(X0,X1) [cnf transformation 269]
308. ~where_9_subquery_10(X0,X1) | where_9(X0,X1) [cnf transformation 269]
309. from_8(X0,X1,sK11(X0,X1),sK12(X0,X1),sK13(X0,X1)) | ~select_7(X0,X1) [cnf transformation 273]
310. ~select_7(X0,X1) | where_9(X0,X1) [cnf transformation 273]
311. ~from_8(X0,X1,X2,X3,X4) | ~where_9(X0,X1) | select_7(X0,X1) [cnf transformation 273]
312. ~main_query_2(X0,X1) | select_7(X0,X1) [cnf transformation 274]
313. ~select_7(X0,X1) | main_query_2(X0,X1) [cnf transformation 274]
314. eq(X0,X0) [cnf transformation 86]
325. ~where_9(X1,X2) | where_9(X1,X0) | ~eq(X2,X0) [cnf transformation 157]
370. main_query_2(sK14,sK15) | main_query_1(sK14,sK15) [cnf transformation 279]
371. ~main_query_2(sK14,sK15) | ~main_query_1(sK14,sK15) [cnf transformation 279]
373. 0 <=> main_query_1(sK14,sK15) [avatar definition]
374. main_query_1(sK14,sK15) <- {0} [avatar component clause 373]
376. 1 <=> ~main_query_1(sK14,sK15) [avatar definition]
377. ~main_query_1(sK14,sK15) <- {1} [avatar component clause 376]
379. 2 <=> main_query_2(sK14,sK15) [avatar definition]
380. main_query_2(sK14,sK15) <- {2} [avatar component clause 379]
382. 3 <=> ~main_query_2(sK14,sK15) [avatar definition]
383. ~main_query_2(sK14,sK15) <- {3} [avatar component clause 382]
384. ~1 | ~3 [avatar split clause 371,382,376]
385. 0 | 2 [avatar split clause 370,379,373]
386. where_13(X0,X0) [resolution 301,314]
388. select_7(sK14,sK15) <- {2} [resolution 312,380]
390. where_9(sK14,sK15) <- {2} [resolution 388,310]
392. where_9_subquery_10(sK14,sK15) <- {2} [resolution 390,307]
395. select_11(sK14,sK15) <- {2} [resolution 392,305]
398. where_13(sK14,sK15) <- {2} [resolution 395,303]
399. eq(sK14,sK15) <- {2} [resolution 398,300]
401. ~where_3_subquery_4(sK14) | where_3(sK14,sK15) <- {2} [resolution 399,290]
410. 4 <=> where_3(sK14,sK15) [avatar definition]
411. where_3(sK14,sK15) <- {4} [avatar component clause 410]
416. 7 <=> ~where_3_subquery_4(sK14) [avatar definition]
417. ~where_3_subquery_4(sK14) <- {7} [avatar component clause 416]
418. 4 | ~7 | ~2 [avatar split clause 401,379,416,410]
524. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_5(X0) [resolution 284,282]
541. tab1(X2,X3,sK4(X2,X3),sK5(X2,X3),sK6(X2,X3)) | ~select_1(X2,X3) [resolution 291,280]
545. tab1(X2,X3,sK11(X2,X3),sK12(X2,X3),sK13(X2,X3)) | ~select_7(X2,X3) [resolution 309,296]
647. from_12(X3,sK0(X3),sK1(X3),sK2(X3),sK3(X3)) | ~select_5(X3) [resolution 524,299]
656. from_6(X2,X3,sK4(X2,X3),sK5(X2,X3),sK6(X2,X3)) | ~select_1(X2,X3) [resolution 541,283]
657. from_8(X4,X5,sK4(X4,X5),sK5(X4,X5),sK6(X4,X5)) | ~select_1(X4,X5) [resolution 541,297]
666. from_2(X0,X1,sK11(X0,X1),sK12(X0,X1),sK13(X0,X1)) | ~select_7(X0,X1) [resolution 545,281]
667. from_6(X2,X3,sK11(X2,X3),sK12(X2,X3),sK13(X2,X3)) | ~select_7(X2,X3) [resolution 545,283]
759. 14 <=> select_5(sK14) [avatar definition]
760. select_5(sK14) <- {14} [avatar component clause 759]
830. ~where_13(X1,X2) | ~select_5(X1) | select_11(X1,X2) [resolution 647,304]
843. select_11(X0,X0) | ~select_5(X0) [resolution 830,386]
856. where_9_subquery_10(X6,X6) | ~select_5(X6) [resolution 843,306]
867. where_9(X4,X4) | ~select_5(X4) [resolution 856,308]
876. ~eq(X0,X1) | where_9(X0,X1) | ~select_5(X0) [resolution 867,325]
967. ~select_1(X2,X3) | select_5(X2) [resolution 656,285]
974. ~where_9(X2,X3) | ~select_1(X2,X3) | select_7(X2,X3) [resolution 657,311]
994. ~select_7(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 666,293]
1004. ~select_7(X2,X3) | select_5(X2) [resolution 667,285]
1010. select_5(sK14) <- {2} [resolution 1004,388]
1025. 14 | ~2 [avatar split clause 1010,379,759]
1038. select_1(sK14,sK15) <- {0} [resolution 374,294]
1039. where_3_subquery_4(sK14) <- {14} [resolution 760,287]
1040. $false <- {7, 14} [subsumption resolution 1039,417]
1041. 7 | ~14 [avatar contradiction clause 1040]
1042. select_5(sK14) <- {0} [resolution 1038,967]
1046. where_3(sK14,sK15) <- {0} [resolution 1038,292]
1049. 4 | ~0 [avatar split clause 1046,373,410]
1050. 14 | ~0 [avatar split clause 1042,373,759]
1055. eq(sK14,sK15) <- {4} [resolution 411,288]
1073. where_9(sK14,sK15) | ~select_5(sK14) <- {4} [resolution 1055,876]
1082. where_9(sK14,sK15) <- {4, 14} [subsumption resolution 1073,760]
1259. ~select_1(sK14,sK15) | select_7(sK14,sK15) <- {4, 14} [resolution 1082,974]
1263. select_7(sK14,sK15) <- {0, 4, 14} [subsumption resolution 1259,1038]
1330. main_query_2(sK14,sK15) <- {0, 4, 14} [resolution 1263,313]
1332. $false <- {0, 3, 4, 14} [subsumption resolution 1330,383]
1333. ~0 | 3 | ~4 | ~14 [avatar contradiction clause 1332]
1338. 20 <=> select_7(sK14,sK15) [avatar definition]
1339. select_7(sK14,sK15) <- {20} [avatar component clause 1338]
1341. 22 <=> select_1(sK14,sK15) [avatar definition]
1342. select_1(sK14,sK15) <- {22} [avatar component clause 1341]
1349. select_7(sK14,sK15) <- {2} [resolution 380,312]
1350. 20 | ~2 [avatar split clause 1349,379,1338]
1836. ~where_3(sK14,sK15) | select_1(sK14,sK15) <- {20} [resolution 994,1339]
1840. select_1(sK14,sK15) <- {4, 20} [subsumption resolution 1836,411]
1849. 22 | ~4 | ~20 [avatar split clause 1840,1338,410,1341]
1863. main_query_1(sK14,sK15) <- {22} [resolution 1342,295]
1870. $false <- {1, 22} [subsumption resolution 1863,377]
1871. 1 | ~22 [avatar contradiction clause 1870]
1872. $false [avatar sat refutation 384,385,418,1025,1041,1049,1050,1333,1350,1849,1871]
% SZS output end Proof for query_15_query_28_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5628
% Time elapsed: 0.112 s
% ------------------------------
% ------------------------------
