% Refutation found. Thanks to Tanya!
% SZS status Theorem for query08_query09_tptp4X
% SZS output start Proof for query08_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X6,X7,X8,X9,X10] : (from_7(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
6. ! [X6] : (select_6(X6) <=> ? [X7,X8,X9,X10] : ($true & from_7(X6,X7,X8,X9,X10))) [input]
7. ! [X11] : (query_5(X11) <=> select_6(X11)) [input]
8. ! [X12] : (from_8(X12) <=> query_5(X12)) [input]
9. ! [X12] : (select_4(X12) <=> ($true & from_8(X12))) [input]
10. ! [X11] : (main_query_2(X11) <=> select_4(X11)) [input]
14. ! [X13] : lte(X13,X13) [input]
20. ! [X13,X14] : (gte(X13,X14) <=> lte(X14,X13)) [input]
52. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
53. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 52]
54. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
55. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
56. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [rectify 6]
57. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 56]
58. ! [X0] : (query_5(X0) <=> select_6(X0)) [rectify 7]
59. ! [X0] : (from_8(X0) <=> query_5(X0)) [rectify 8]
60. ! [X0] : (select_4(X0) <=> ($true & from_8(X0))) [rectify 9]
61. ! [X0] : (select_4(X0) <=> from_8(X0)) [true and false elimination 60]
62. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 10]
66. ! [X0] : lte(X0,X0) [rectify 14]
72. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 20]
104. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 53]
172. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 104]
173. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
174. ! [X0] : ((where_3(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
175. ! [X0] : ((where_3(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_3(X0))) [flattening 174]
176. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
177. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 176]
178. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
179. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 177,178]
180. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 54]
181. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 55]
182. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 57]
183. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 182]
184. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
185. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 183,184]
186. ! [X0] : ((query_5(X0) | ~select_6(X0)) & (select_6(X0) | ~query_5(X0))) [nnf transformation 58]
187. ! [X0] : ((from_8(X0) | ~query_5(X0)) & (query_5(X0) | ~from_8(X0))) [nnf transformation 59]
188. ! [X0] : ((select_4(X0) | ~from_8(X0)) & (from_8(X0) | ~select_4(X0))) [nnf transformation 61]
189. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 62]
192. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 72]
193. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 172]
194. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
195. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 193,194]
196. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 173]
197. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 173]
200. ~gte(X0,X0) | where_3(X0) | ~lte(X0,X0) [cnf transformation 175]
201. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 179]
203. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 179]
204. ~main_query_1(X0) | select_1(X0) [cnf transformation 180]
205. ~select_1(X0) | main_query_1(X0) [cnf transformation 180]
206. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 181]
207. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 181]
208. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 185]
209. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 185]
210. ~query_5(X0) | select_6(X0) [cnf transformation 186]
211. ~select_6(X0) | query_5(X0) [cnf transformation 186]
212. ~from_8(X0) | query_5(X0) [cnf transformation 187]
213. ~query_5(X0) | from_8(X0) [cnf transformation 187]
214. ~select_4(X0) | from_8(X0) [cnf transformation 188]
215. ~from_8(X0) | select_4(X0) [cnf transformation 188]
216. ~main_query_2(X0) | select_4(X0) [cnf transformation 189]
217. ~select_4(X0) | main_query_2(X0) [cnf transformation 189]
221. lte(X0,X0) [cnf transformation 66]
227. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 192]
259. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 195]
260. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 195]
293. 0 <=> main_query_1(sK8) [avatar definition]
294. main_query_1(sK8) <- {0} [avatar component clause 293]
296. 1 <=> ~main_query_1(sK8) [avatar definition]
297. ~main_query_1(sK8) <- {1} [avatar component clause 296]
299. 2 <=> main_query_2(sK8) [avatar definition]
300. main_query_2(sK8) <- {2} [avatar component clause 299]
302. 3 <=> ~main_query_2(sK8) [avatar definition]
303. ~main_query_2(sK8) <- {3} [avatar component clause 302]
304. ~1 | ~3 [avatar split clause 260,302,296]
305. 0 | 2 [avatar split clause 259,299,293]
306. select_4(sK8) <- {2} [resolution 216,300]
307. from_8(sK8) <- {2} [resolution 306,214]
310. query_5(sK8) <- {2} [resolution 307,212]
312. select_6(sK8) <- {2} [resolution 310,210]
319. gte(X1,X1) [resolution 227,221]
321. where_3(X0) | ~lte(X0,X0) [resolution 200,319]
323. where_3(X0) [subsumption resolution 321,221]
329. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 201,196]
331. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 208,206]
333. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 329,207]
336. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 331,197]
338. ~select_1(X1) | select_6(X1) [resolution 333,209]
340. ~select_6(X1) | ~where_3(X1) | select_1(X1) [resolution 336,203]
341. ~select_6(X1) | select_1(X1) [subsumption resolution 340,323]
342. select_1(sK8) <- {2} [resolution 341,312]
344. main_query_1(sK8) <- {2} [resolution 342,205]
346. $false <- {1, 2} [subsumption resolution 344,297]
347. 1 | ~2 [avatar contradiction clause 346]
348. select_1(sK8) <- {0} [resolution 294,204]
349. select_6(sK8) <- {0} [resolution 348,338]
353. query_5(sK8) <- {0} [resolution 349,211]
354. from_8(sK8) <- {0} [resolution 353,213]
356. select_4(sK8) <- {0} [resolution 354,215]
358. main_query_2(sK8) <- {0} [resolution 356,217]
360. $false <- {0, 3} [subsumption resolution 358,303]
361. ~0 | 3 [avatar contradiction clause 360]
362. $false [avatar sat refutation 304,305,347,361]
% SZS output end Proof for query08_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
