% Refutation found. Thanks to Tanya!
% SZS status Theorem for query04_query08_tptp4X
% SZS output start Proof for query04_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X5,X1,X2,X3,X4] : (from_5(X5,X1,X2,X3,X4) <=> tab1(X5,X1,X2,X3,X4)) [input]
6. ! [X5] : (where_6(X5) <=> (gte(X5,X5) & lte(X5,X5))) [input]
7. ! [X5] : (select_4(X5) <=> ? [X1,X2,X3,X4] : (where_6(X5) & from_5(X5,X1,X2,X3,X4))) [input]
8. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
12. ! [X7] : lte(X7,X7) [input]
19. ! [X7,X8] : (gte(X7,X8) <=> lte(X8,X7)) [input]
40. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
41. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 40]
42. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
43. ! [X0] : (where_6(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 6]
44. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
45. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
46. ! [X0] : eq(X0,X0) [rectify 9]
49. ! [X0] : lte(X0,X0) [rectify 12]
56. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 19]
122. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 41]
123. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
124. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
125. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
126. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 125]
127. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
128. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 126,127]
129. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
130. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 42]
131. ! [X0] : ((where_6(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [nnf transformation 43]
132. ! [X0] : ((where_6(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [flattening 131]
133. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 44]
134. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 133]
135. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
136. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 134,135]
137. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 45]
140. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 56]
141. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 122]
142. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
143. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 141,142]
144. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 123]
145. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 123]
147. ~eq(X0,X0) | where_3(X0) [cnf transformation 124]
148. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 128]
150. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 128]
151. ~main_query_1(X0) | select_1(X0) [cnf transformation 129]
152. ~select_1(X0) | main_query_1(X0) [cnf transformation 129]
153. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 130]
154. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 130]
157. ~gte(X0,X0) | where_6(X0) | ~lte(X0,X0) [cnf transformation 132]
158. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 136]
160. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 136]
161. ~main_query_2(X0) | select_4(X0) [cnf transformation 137]
162. ~select_4(X0) | main_query_2(X0) [cnf transformation 137]
163. eq(X0,X0) [cnf transformation 46]
166. lte(X0,X0) [cnf transformation 49]
173. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 140]
194. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 143]
195. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 143]
197. 0 <=> main_query_1(sK8) [avatar definition]
198. main_query_1(sK8) <- {0} [avatar component clause 197]
200. 1 <=> ~main_query_1(sK8) [avatar definition]
201. ~main_query_1(sK8) <- {1} [avatar component clause 200]
203. 2 <=> main_query_2(sK8) [avatar definition]
204. main_query_2(sK8) <- {2} [avatar component clause 203]
206. 3 <=> ~main_query_2(sK8) [avatar definition]
207. ~main_query_2(sK8) <- {3} [avatar component clause 206]
208. ~1 | ~3 [avatar split clause 195,206,200]
209. 0 | 2 [avatar split clause 194,203,197]
210. select_4(sK8) <- {2} [resolution 161,204]
214. where_3(X1) [resolution 147,163]
224. gte(X1,X1) [resolution 173,166]
229. where_6(X1) | ~lte(X1,X1) [resolution 224,157]
230. where_6(X1) [subsumption resolution 229,166]
297. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 148,144]
299. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 158,153]
302. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 297,154]
306. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 299,145]
309. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 302,160]
310. ~select_1(X1) | select_4(X1) [subsumption resolution 309,230]
314. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 306,150]
320. ~select_4(X1) | select_1(X1) [subsumption resolution 314,214]
323. select_1(sK8) <- {2} [resolution 320,210]
327. main_query_1(sK8) <- {2} [resolution 323,152]
329. $false <- {1, 2} [subsumption resolution 327,201]
330. 1 | ~2 [avatar contradiction clause 329]
331. select_1(sK8) <- {0} [resolution 198,151]
334. select_4(sK8) <- {0} [resolution 331,310]
340. main_query_2(sK8) <- {0} [resolution 334,162]
342. $false <- {0, 3} [subsumption resolution 340,207]
343. ~0 | 3 [avatar contradiction clause 342]
344. $false [avatar sat refutation 208,209,330,343]
% SZS output end Proof for query04_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.025 s
% ------------------------------
% ------------------------------
