% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query08_tptp4X
% SZS output start Proof for query00_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
6. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
11. ! [X6] : lte(X6,X6) [input]
17. ! [X6,X7] : (gte(X6,X7) <=> lte(X7,X6)) [input]
46. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
47. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 46]
48. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
50. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
54. ! [X0] : lte(X0,X0) [rectify 11]
60. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 17]
89. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 47]
151. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 89]
152. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
153. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 48]
154. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 153]
155. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
156. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 154,155]
157. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 49]
158. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
159. ! [X0] : ((where_5(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [nnf transformation 5]
160. ! [X0] : ((where_5(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [flattening 159]
161. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 6]
162. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 161]
163. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
164. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 162,163]
165. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 50]
168. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 60]
169. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 151]
170. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
171. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 169,170]
172. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 152]
173. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 152]
174. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 156]
175. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 156]
176. ~main_query_1(X0) | select_1(X0) [cnf transformation 157]
177. ~select_1(X0) | main_query_1(X0) [cnf transformation 157]
178. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 158]
179. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 158]
182. ~gte(X0,X0) | where_5(X0) | ~lte(X0,X0) [cnf transformation 160]
183. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 164]
185. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 164]
186. ~main_query_2(X0) | select_3(X0) [cnf transformation 165]
187. ~select_3(X0) | main_query_2(X0) [cnf transformation 165]
191. lte(X0,X0) [cnf transformation 54]
197. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 168]
226. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 171]
227. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 171]
257. 0 <=> main_query_1(sK8) [avatar definition]
258. main_query_1(sK8) <- {0} [avatar component clause 257]
260. 1 <=> ~main_query_1(sK8) [avatar definition]
261. ~main_query_1(sK8) <- {1} [avatar component clause 260]
263. 2 <=> main_query_2(sK8) [avatar definition]
264. main_query_2(sK8) <- {2} [avatar component clause 263]
266. 3 <=> ~main_query_2(sK8) [avatar definition]
267. ~main_query_2(sK8) <- {3} [avatar component clause 266]
268. ~1 | ~3 [avatar split clause 227,266,260]
269. 0 | 2 [avatar split clause 226,263,257]
270. select_3(sK8) <- {2} [resolution 186,264]
278. gte(X1,X1) [resolution 197,191]
280. where_5(X0) | ~lte(X0,X0) [resolution 182,278]
282. where_5(X0) [subsumption resolution 280,191]
288. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 174,172]
290. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 183,178]
292. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 288,179]
295. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 290,173]
297. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 292,185]
298. ~select_1(X1) | select_3(X1) [subsumption resolution 297,282]
300. ~select_3(X1) | select_1(X1) [resolution 295,175]
301. select_1(sK8) <- {2} [resolution 300,270]
303. main_query_1(sK8) <- {2} [resolution 301,177]
304. $false <- {1, 2} [subsumption resolution 303,261]
305. 1 | ~2 [avatar contradiction clause 304]
306. select_1(sK8) <- {0} [resolution 258,176]
307. select_3(sK8) <- {0} [resolution 306,298]
310. main_query_2(sK8) <- {0} [resolution 307,187]
312. $false <- {0, 3} [subsumption resolution 310,267]
313. ~0 | 3 [avatar contradiction clause 312]
314. $false [avatar sat refutation 268,269,305,313]
% SZS output end Proof for query00_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.017 s
% ------------------------------
% ------------------------------
