% Refutation found. Thanks to Tanya!
% SZS status Theorem for query10_query12_tptp4X
% SZS output start Proof for query10_query12_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X1,X0] : (where_3(X1,X0) <=> eq(X0,X1)) [input]
3. ! [X1,X0] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X1,X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X1,X0] : (where_6(X1,X0) <=> (lte(X1,X0) & lte(X0,X1))) [input]
7. ! [X1,X0] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X1,X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
13. ! [X7,X8] : ((lte(X8,X7) & lte(X7,X8)) <=> eq(X7,X8)) [input]
54. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
55. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 54]
56. ! [X0,X1] : (where_3(X0,X1) <=> eq(X1,X0)) [rectify 2]
57. ! [X0,X1] : (select_1(X1,X0) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X1,X0,X2,X3,X4))) [rectify 3]
58. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
59. ! [X0,X1] : (where_6(X0,X1) <=> (lte(X0,X1) & lte(X1,X0))) [rectify 6]
60. ! [X0,X1] : (select_4(X1,X0) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X1,X0,X2,X3,X4))) [rectify 7]
61. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
66. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
107. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 55]
183. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 107]
184. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
185. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X1,X0)) & (eq(X1,X0) | ~where_3(X0,X1))) [nnf transformation 56]
186. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X1,X0,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X1,X0,X2,X3,X4)) | ~select_1(X1,X0))) [nnf transformation 57]
187. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X1,X0,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X1,X0,X5,X6,X7)) | ~select_1(X1,X0))) [rectify 186]
188. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X1,X0,X5,X6,X7)) => (where_3(X0,X1) & from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
189. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X1,X0,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X1,X0))) [skolemisation 187,188]
190. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 58]
191. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
192. ! [X0,X1] : ((where_6(X0,X1) | (~lte(X0,X1) | ~lte(X1,X0))) & ((lte(X0,X1) & lte(X1,X0)) | ~where_6(X0,X1))) [nnf transformation 59]
193. ! [X0,X1] : ((where_6(X0,X1) | ~lte(X0,X1) | ~lte(X1,X0)) & ((lte(X0,X1) & lte(X1,X0)) | ~where_6(X0,X1))) [flattening 192]
194. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X1,X0,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X1,X0,X2,X3,X4)) | ~select_4(X1,X0))) [nnf transformation 60]
195. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X1,X0,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X1,X0,X5,X6,X7)) | ~select_4(X1,X0))) [rectify 194]
196. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X1,X0,X5,X6,X7)) => (where_6(X0,X1) & from_5(X1,X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
197. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X1,X0,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X1,X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X1,X0))) [skolemisation 195,196]
198. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 61]
199. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 66]
200. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 199]
201. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 183]
202. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
203. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 201,202]
204. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 184]
205. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 184]
206. ~where_3(X0,X1) | eq(X1,X0) [cnf transformation 185]
207. ~eq(X1,X0) | where_3(X0,X1) [cnf transformation 185]
208. from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X1,X0) [cnf transformation 189]
209. ~select_1(X1,X0) | where_3(X0,X1) [cnf transformation 189]
210. ~from_2(X1,X0,X2,X3,X4) | ~where_3(X0,X1) | select_1(X1,X0) [cnf transformation 189]
211. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 190]
212. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 190]
213. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 191]
214. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 191]
215. ~where_6(X0,X1) | lte(X1,X0) [cnf transformation 193]
216. ~where_6(X0,X1) | lte(X0,X1) [cnf transformation 193]
217. ~lte(X1,X0) | ~lte(X0,X1) | where_6(X0,X1) [cnf transformation 193]
218. from_5(X1,X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X1,X0) [cnf transformation 197]
219. ~select_4(X1,X0) | where_6(X0,X1) [cnf transformation 197]
220. ~from_5(X1,X0,X2,X3,X4) | ~where_6(X0,X1) | select_4(X1,X0) [cnf transformation 197]
221. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 198]
222. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 198]
227. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 200]
228. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 200]
229. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 200]
266. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 203]
267. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 203]
304. 0 <=> main_query_1(sK6,sK7) [avatar definition]
305. main_query_1(sK6,sK7) <- {0} [avatar component clause 304]
307. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
308. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 307]
310. 2 <=> main_query_2(sK6,sK7) [avatar definition]
311. main_query_2(sK6,sK7) <- {2} [avatar component clause 310]
313. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
314. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 313]
315. ~1 | ~3 [avatar split clause 267,313,307]
316. 0 | 2 [avatar split clause 266,310,304]
319. select_4(sK6,sK7) <- {2} [resolution 221,311]
320. where_6(sK7,sK6) <- {2} [resolution 319,219]
322. lte(sK7,sK6) <- {2} [resolution 320,216]
323. lte(sK6,sK7) <- {2} [resolution 320,215]
336. eq(sK6,sK7) | ~lte(sK6,sK7) <- {2} [resolution 227,322]
338. eq(sK6,sK7) <- {2} [subsumption resolution 336,323]
349. where_3(sK7,sK6) <- {2} [resolution 338,207]
364. tab1(X2,X3,sK0(X3,X2),sK1(X3,X2),sK2(X3,X2)) | ~select_1(X2,X3) [resolution 208,204]
368. tab1(X2,X3,sK3(X3,X2),sK4(X3,X2),sK5(X3,X2)) | ~select_4(X2,X3) [resolution 218,213]
369. from_5(X0,X1,sK0(X1,X0),sK1(X1,X0),sK2(X1,X0)) | ~select_1(X0,X1) [resolution 364,214]
372. from_2(X2,X3,sK3(X3,X2),sK4(X3,X2),sK5(X3,X2)) | ~select_4(X2,X3) [resolution 368,205]
373. ~where_6(X1,X0) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 369,220]
375. ~select_4(X0,X1) | ~where_3(X1,X0) | select_1(X0,X1) [resolution 372,210]
393. ~where_3(sK7,sK6) | select_1(sK6,sK7) <- {2} [resolution 375,319]
394. select_1(sK6,sK7) <- {2} [subsumption resolution 393,349]
395. main_query_1(sK6,sK7) <- {2} [resolution 394,212]
397. $false <- {1, 2} [subsumption resolution 395,308]
398. 1 | ~2 [avatar contradiction clause 397]
399. select_1(sK6,sK7) <- {0} [resolution 305,211]
401. where_3(sK7,sK6) <- {0} [resolution 399,209]
402. eq(sK6,sK7) <- {0} [resolution 401,206]
404. lte(sK7,sK6) <- {0} [resolution 402,229]
405. lte(sK6,sK7) <- {0} [resolution 402,228]
414. ~lte(sK7,sK6) | where_6(sK7,sK6) <- {0} [resolution 405,217]
415. where_6(sK7,sK6) <- {0} [subsumption resolution 414,404]
426. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 415,373]
429. select_4(sK6,sK7) <- {0} [subsumption resolution 426,399]
434. main_query_2(sK6,sK7) <- {0} [resolution 429,222]
436. $false <- {0, 3} [subsumption resolution 434,314]
437. ~0 | 3 [avatar contradiction clause 436]
438. $false [avatar sat refutation 315,316,398,437]
% SZS output end Proof for query10_query12_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.033 s
% ------------------------------
% ------------------------------
