% Refutation found. Thanks to Tanya!
% SZS status Theorem for query06_query07_tptp4X
% SZS output start Proof for query06_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : ($true & from_5(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (subquery_from_3(X5) <=> select_4(X5)) [input]
4. ! [X6] : (from_2(X6) <=> subquery_from_3(X6)) [input]
5. ! [X6] : (select_1(X6) <=> ($true & from_2(X6))) [input]
6. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
7. ! [X7,X8,X9,X10,X11] : (from_7(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
8. ! [X7] : (where_8(X7) <=> (~gt(X7,X7) & ~lt(X7,X7))) [input]
9. ! [X7] : (select_6(X7) <=> ? [X8,X9,X10,X11] : (where_8(X7) & from_7(X7,X8,X9,X10,X11))) [input]
10. ! [X7] : (main_query_2(X7) <=> select_6(X7)) [input]
11. ! [X12] : eq(X12,X12) [input]
18. ! [X12,X13] : (lt(X12,X13) <=> (~eq(X12,X13) & lte(X12,X13))) [input]
19. ! [X12,X13] : (gt(X12,X13) <=> lt(X13,X12)) [input]
45. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
46. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 45]
47. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4)) [true and false elimination 2]
48. ! [X0] : (subquery_from_3(X0) <=> select_4(X0)) [rectify 3]
49. ! [X0] : (from_2(X0) <=> subquery_from_3(X0)) [rectify 4]
50. ! [X0] : (select_1(X0) <=> ($true & from_2(X0))) [rectify 5]
51. ! [X0] : (select_1(X0) <=> from_2(X0)) [true and false elimination 50]
52. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 6]
53. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 7]
54. ! [X0] : (where_8(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [rectify 8]
55. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4))) [rectify 9]
56. ! [X0] : (main_query_2(X0) <=> select_6(X0)) [rectify 10]
57. ! [X0] : eq(X0,X0) [rectify 11]
64. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 18]
65. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 19]
91. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 46]
143. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 91]
144. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 1]
145. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4) | ~select_4(X0))) [nnf transformation 47]
146. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) | ~select_4(X0))) [rectify 145]
147. ! [X0] : (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) => from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
148. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0))) [skolemisation 146,147]
149. ! [X0] : ((subquery_from_3(X0) | ~select_4(X0)) & (select_4(X0) | ~subquery_from_3(X0))) [nnf transformation 48]
150. ! [X0] : ((from_2(X0) | ~subquery_from_3(X0)) & (subquery_from_3(X0) | ~from_2(X0))) [nnf transformation 49]
151. ! [X0] : ((select_1(X0) | ~from_2(X0)) & (from_2(X0) | ~select_1(X0))) [nnf transformation 51]
152. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 52]
153. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 53]
154. ! [X0] : ((where_8(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_8(X0))) [nnf transformation 54]
155. ! [X0] : ((where_8(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_8(X0))) [flattening 154]
156. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4)) | ~select_6(X0))) [nnf transformation 55]
157. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) | ~select_6(X0))) [rectify 156]
158. ! [X0] : (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) => (where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
159. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & ((where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_6(X0))) [skolemisation 157,158]
160. ! [X0] : ((main_query_2(X0) | ~select_6(X0)) & (select_6(X0) | ~main_query_2(X0))) [nnf transformation 56]
163. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 64]
164. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 163]
165. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 65]
166. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 143]
167. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
168. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 166,167]
169. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 144]
170. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 144]
171. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [cnf transformation 148]
172. ~from_5(X0,X1,X2,X3,X4) | select_4(X0) [cnf transformation 148]
173. ~subquery_from_3(X0) | select_4(X0) [cnf transformation 149]
174. ~select_4(X0) | subquery_from_3(X0) [cnf transformation 149]
175. ~from_2(X0) | subquery_from_3(X0) [cnf transformation 150]
176. ~subquery_from_3(X0) | from_2(X0) [cnf transformation 150]
177. ~select_1(X0) | from_2(X0) [cnf transformation 151]
178. ~from_2(X0) | select_1(X0) [cnf transformation 151]
179. ~main_query_1(X0) | select_1(X0) [cnf transformation 152]
180. ~select_1(X0) | main_query_1(X0) [cnf transformation 152]
181. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 153]
182. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 153]
185. gt(X0,X0) | where_8(X0) | lt(X0,X0) [cnf transformation 155]
186. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 159]
188. ~from_7(X0,X1,X2,X3,X4) | ~where_8(X0) | select_6(X0) [cnf transformation 159]
189. ~main_query_2(X0) | select_6(X0) [cnf transformation 160]
190. ~select_6(X0) | main_query_2(X0) [cnf transformation 160]
191. eq(X0,X0) [cnf transformation 57]
201. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 164]
203. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 165]
228. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 168]
229. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 168]
231. 0 <=> main_query_1(sK8) [avatar definition]
232. main_query_1(sK8) <- {0} [avatar component clause 231]
234. 1 <=> ~main_query_1(sK8) [avatar definition]
235. ~main_query_1(sK8) <- {1} [avatar component clause 234]
237. 2 <=> main_query_2(sK8) [avatar definition]
238. main_query_2(sK8) <- {2} [avatar component clause 237]
240. 3 <=> ~main_query_2(sK8) [avatar definition]
241. ~main_query_2(sK8) <- {3} [avatar component clause 240]
242. ~1 | ~3 [avatar split clause 229,240,234]
243. 0 | 2 [avatar split clause 228,237,231]
244. select_6(sK8) <- {2} [resolution 189,238]
251. ~lt(X0,X0) [resolution 201,191]
253. where_8(X1) | lt(X1,X1) | lt(X1,X1) [resolution 185,203]
254. where_8(X1) | lt(X1,X1) [duplicate literal removal 253]
255. where_8(X1) [subsumption resolution 254,251]
319. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [resolution 171,169]
359. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 186,181]
435. from_7(X1,sK0(X1),sK1(X1),sK2(X1),sK3(X1)) | ~select_4(X1) [resolution 319,182]
441. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 359,170]
449. ~select_4(X1) | ~where_8(X1) | select_6(X1) [resolution 435,188]
450. ~select_4(X1) | select_6(X1) [subsumption resolution 449,255]
456. ~select_6(X5) | select_4(X5) [resolution 441,172]
459. select_4(sK8) <- {2} [resolution 456,244]
463. subquery_from_3(sK8) <- {2} [resolution 459,174]
466. from_2(sK8) <- {2} [resolution 463,176]
470. select_1(sK8) <- {2} [resolution 466,178]
474. main_query_1(sK8) <- {2} [resolution 470,180]
476. $false <- {1, 2} [subsumption resolution 474,235]
477. 1 | ~2 [avatar contradiction clause 476]
478. select_1(sK8) <- {0} [resolution 232,179]
482. from_2(sK8) <- {0} [resolution 478,177]
486. subquery_from_3(sK8) <- {0} [resolution 482,175]
490. select_4(sK8) <- {0} [resolution 486,173]
498. select_6(sK8) <- {0} [resolution 490,450]
508. main_query_2(sK8) <- {0} [resolution 498,190]
510. $false <- {0, 3} [subsumption resolution 508,241]
511. ~0 | 3 [avatar contradiction clause 510]
512. $false [avatar sat refutation 242,243,477,511]
% SZS output end Proof for query06_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.018 s
% ------------------------------
% ------------------------------
