% Refutation found. Thanks to Tanya!
% SZS status Theorem for query04_query07_tptp4X
% SZS output start Proof for query04_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
7. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [input]
9. ! [X5] : eq(X5,X5) [input]
16. ! [X5,X6] : (lt(X5,X6) <=> (~eq(X5,X6) & lte(X5,X6))) [input]
17. ! [X5,X6] : (gt(X5,X6) <=> lt(X6,X5)) [input]
43. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
44. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 43]
45. ! [X0] : eq(X0,X0) [rectify 9]
52. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 16]
53. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 17]
130. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 44]
131. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
132. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
133. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 133]
135. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 134,135]
137. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
138. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
139. ! [X0] : ((where_6(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [nnf transformation 6]
140. ! [X0] : ((where_6(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [flattening 139]
141. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 7]
142. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 141]
143. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
144. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 142,143]
145. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 8]
148. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 52]
149. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 148]
150. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 53]
151. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 130]
152. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
153. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 151,152]
154. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 131]
155. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 131]
157. ~eq(X0,X0) | where_3(X0) [cnf transformation 132]
158. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 136]
160. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 136]
161. ~main_query_1(X0) | select_1(X0) [cnf transformation 137]
162. ~select_1(X0) | main_query_1(X0) [cnf transformation 137]
163. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 138]
164. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 138]
167. gt(X0,X0) | where_6(X0) | lt(X0,X0) [cnf transformation 140]
168. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 144]
170. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 144]
171. ~main_query_2(X0) | select_4(X0) [cnf transformation 145]
172. ~select_4(X0) | main_query_2(X0) [cnf transformation 145]
173. eq(X0,X0) [cnf transformation 45]
183. ~lt(X0,X1) | ~eq(X0,X1) [cnf transformation 149]
185. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 150]
210. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 153]
211. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 153]
213. 0 <=> main_query_1(sK8) [avatar definition]
214. main_query_1(sK8) <- {0} [avatar component clause 213]
216. 1 <=> ~main_query_1(sK8) [avatar definition]
217. ~main_query_1(sK8) <- {1} [avatar component clause 216]
219. 2 <=> main_query_2(sK8) [avatar definition]
220. main_query_2(sK8) <- {2} [avatar component clause 219]
222. 3 <=> ~main_query_2(sK8) [avatar definition]
223. ~main_query_2(sK8) <- {3} [avatar component clause 222]
224. ~1 | ~3 [avatar split clause 211,222,216]
225. 0 | 2 [avatar split clause 210,219,213]
226. select_4(sK8) <- {2} [resolution 171,220]
230. where_3(X1) [resolution 157,173]
239. where_6(X1) | lt(X1,X1) | lt(X1,X1) [resolution 167,185]
240. lt(X1,X1) | where_6(X1) [duplicate literal removal 239]
243. where_6(X2) | ~eq(X2,X2) [resolution 240,183]
245. where_6(X2) [subsumption resolution 243,173]
318. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 158,154]
349. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 168,163]
444. from_5(X1,sK0(X1),sK1(X1),sK2(X1),sK3(X1)) | ~select_1(X1) [resolution 318,164]
451. from_2(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 349,155]
464. ~select_1(X9) | ~where_6(X9) | select_4(X9) [resolution 444,170]
465. ~select_1(X9) | select_4(X9) [subsumption resolution 464,245]
469. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 451,160]
475. ~select_4(X1) | select_1(X1) [subsumption resolution 469,230]
476. select_1(sK8) <- {2} [resolution 475,226]
478. main_query_1(sK8) <- {2} [resolution 476,162]
480. $false <- {1, 2} [subsumption resolution 478,217]
481. 1 | ~2 [avatar contradiction clause 480]
482. select_1(sK8) <- {0} [resolution 214,161]
483. select_4(sK8) <- {0} [resolution 482,465]
487. main_query_2(sK8) <- {0} [resolution 483,172]
489. $false <- {0, 3} [subsumption resolution 487,223]
490. ~0 | 3 [avatar contradiction clause 489]
491. $false [avatar sat refutation 224,225,481,490]
% SZS output end Proof for query04_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.013 s
% ------------------------------
% ------------------------------
