% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_03_query_09_tptp4X
% SZS output start Proof for query_03_query_09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
4. ! [X7,X8,X9,X10,X11] : (from_4(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_8(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7] : (select_7(X7) <=> ? [X8,X9,X10,X11] : ($true & from_8(X7,X8,X9,X10,X11))) [input]
7. ! [X5] : (where_5_subquery_6(X5) <=> select_7(X5)) [input]
8. ! [X7] : (where_5(X7) <=> where_5_subquery_6(X7)) [input]
9. ! [X7,X8] : (select_3(X7,X8) <=> ? [X9,X10,X11] : (where_5(X7) & from_4(X7,X8,X9,X10,X11))) [input]
10. ! [X5,X6] : (main_query_2(X5,X6) <=> select_3(X5,X6)) [input]
54. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
55. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 54]
56. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 3]
58. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
59. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
60. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
61. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 60]
62. ! [X0] : (where_5_subquery_6(X0) <=> select_7(X0)) [rectify 7]
63. ! [X0] : (where_5(X0) <=> where_5_subquery_6(X0)) [rectify 8]
64. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 9]
65. ! [X0,X1] : (main_query_2(X0,X1) <=> select_3(X0,X1)) [rectify 10]
109. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 55]
181. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 109]
182. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
183. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0,X1))) [nnf transformation 56]
184. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) | ~select_1(X0,X1))) [rectify 183]
185. ! [X1,X0] : (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) => from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) [choice axiom]
186. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1))) [skolemisation 184,185]
187. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
188. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 58]
189. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 59]
190. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 61]
191. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 190]
192. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
193. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_7(X0))) [skolemisation 191,192]
194. ! [X0] : ((where_5_subquery_6(X0) | ~select_7(X0)) & (select_7(X0) | ~where_5_subquery_6(X0))) [nnf transformation 62]
195. ! [X0] : ((where_5(X0) | ~where_5_subquery_6(X0)) & (where_5_subquery_6(X0) | ~where_5(X0))) [nnf transformation 63]
196. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0,X1))) [nnf transformation 64]
197. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) | ~select_3(X0,X1))) [rectify 196]
198. ! [X1,X0] : (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) => (where_5(X0) & from_4(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
199. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_3(X0,X1))) [skolemisation 197,198]
200. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 65]
203. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 181]
204. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
205. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 203,204]
206. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 182]
207. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 182]
208. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 186]
209. ~from_2(X0,X1,X2,X3,X4) | select_1(X0,X1) [cnf transformation 186]
210. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 187]
211. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 187]
212. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 188]
213. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 188]
215. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 189]
217. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 193]
219. ~select_7(X0) | where_5_subquery_6(X0) [cnf transformation 194]
221. ~where_5_subquery_6(X0) | where_5(X0) [cnf transformation 195]
222. from_4(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_3(X0,X1) [cnf transformation 199]
224. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0,X1) [cnf transformation 199]
225. ~main_query_2(X0,X1) | select_3(X0,X1) [cnf transformation 200]
226. ~select_3(X0,X1) | main_query_2(X0,X1) [cnf transformation 200]
268. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 205]
269. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 205]
271. 0 <=> main_query_1(sK10,sK11) [avatar definition]
272. main_query_1(sK10,sK11) <- {0} [avatar component clause 271]
274. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
275. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 274]
277. 2 <=> main_query_2(sK10,sK11) [avatar definition]
278. main_query_2(sK10,sK11) <- {2} [avatar component clause 277]
280. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
281. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 280]
282. ~1 | ~3 [avatar split clause 269,280,274]
283. 0 | 2 [avatar split clause 268,277,271]
284. select_3(sK10,sK11) <- {2} [resolution 225,278]
307. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 208,206]
312. tab1(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_3(X0,X1) [resolution 222,212]
354. from_4(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 307,213]
355. from_8(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 307,215]
361. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_3(X0,X1) [resolution 312,207]
488. ~select_1(X2,X3) | ~where_5(X2) | select_3(X2,X3) [resolution 354,224]
510. ~select_1(X2,X3) | select_7(X2) [resolution 355,217]
532. ~select_3(X2,X3) | select_1(X2,X3) [resolution 361,209]
555. select_1(sK10,sK11) <- {2} [resolution 532,284]
567. main_query_1(sK10,sK11) <- {2} [resolution 555,211]
568. $false <- {1, 2} [subsumption resolution 567,275]
569. 1 | ~2 [avatar contradiction clause 568]
570. select_1(sK10,sK11) <- {0} [resolution 272,210]
571. select_7(sK10) <- {0} [resolution 570,510]
572. ~where_5(sK10) | select_3(sK10,sK11) <- {0} [resolution 570,488]
578. 4 <=> select_3(sK10,sK11) [avatar definition]
579. select_3(sK10,sK11) <- {4} [avatar component clause 578]
581. 6 <=> where_5(sK10) [avatar definition]
584. 7 <=> ~where_5(sK10) [avatar definition]
586. 4 | ~7 | ~0 [avatar split clause 572,271,584,578]
587. where_5_subquery_6(sK10) <- {0} [resolution 571,219]
604. where_5(sK10) <- {0} [resolution 587,221]
608. 6 | ~0 [avatar split clause 604,271,581]
628. main_query_2(sK10,sK11) <- {4} [resolution 579,226]
630. $false <- {3, 4} [subsumption resolution 628,281]
631. 3 | ~4 [avatar contradiction clause 630]
632. $false [avatar sat refutation 282,283,569,586,608,631]
% SZS output end Proof for query_03_query_09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.012 s
% ------------------------------
% ------------------------------
