% Refutation found. Thanks to Tanya!
% SZS status Theorem for query01_query02_tptp4X
% SZS output start Proof for query01_query02_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X5,X6,X7,X8,X9] : (from_4(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
5. ! [X5] : (select_3(X5) <=> ? [X6,X7,X8,X9] : ($true & from_4(X5,X6,X7,X8,X9))) [input]
6. ! [X10] : (main_query_2(X10) <=> select_3(X10)) [input]
41. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
42. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
45. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [rectify 5]
46. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 45]
47. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 6]
133. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 42]
134. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 135]
137. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 136,137]
139. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
140. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 44]
141. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 46]
142. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 141]
143. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
144. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0))) [skolemisation 142,143]
145. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 47]
148. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 133]
149. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
150. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 148,149]
151. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 134]
152. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 134]
153. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 138]
154. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 138]
155. ~main_query_1(X0) | select_1(X0) [cnf transformation 139]
156. ~select_1(X0) | main_query_1(X0) [cnf transformation 139]
157. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 140]
158. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 140]
159. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 144]
160. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 144]
161. ~main_query_2(X0) | select_3(X0) [cnf transformation 145]
162. ~select_3(X0) | main_query_2(X0) [cnf transformation 145]
195. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 150]
196. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 150]
198. 0 <=> main_query_1(sK8) [avatar definition]
199. main_query_1(sK8) <- {0} [avatar component clause 198]
201. 1 <=> ~main_query_1(sK8) [avatar definition]
202. ~main_query_1(sK8) <- {1} [avatar component clause 201]
204. 2 <=> main_query_2(sK8) [avatar definition]
205. main_query_2(sK8) <- {2} [avatar component clause 204]
207. 3 <=> ~main_query_2(sK8) [avatar definition]
208. ~main_query_2(sK8) <- {3} [avatar component clause 207]
209. ~1 | ~3 [avatar split clause 196,207,201]
210. 0 | 2 [avatar split clause 195,204,198]
211. select_3(sK8) <- {2} [resolution 161,205]
242. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 153,151]
244. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 159,157]
246. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 242,158]
249. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 244,152]
251. ~select_1(X1) | select_3(X1) [resolution 246,160]
253. ~select_3(X1) | select_1(X1) [resolution 249,154]
256. select_1(sK8) <- {2} [resolution 253,211]
260. main_query_1(sK8) <- {2} [resolution 256,156]
261. $false <- {1, 2} [subsumption resolution 260,202]
262. 1 | ~2 [avatar contradiction clause 261]
263. select_1(sK8) <- {0} [resolution 199,155]
266. select_3(sK8) <- {0} [resolution 263,251]
271. main_query_2(sK8) <- {0} [resolution 266,162]
272. $false <- {0, 3} [subsumption resolution 271,208]
273. ~0 | 3 [avatar contradiction clause 272]
274. $false [avatar sat refutation 209,210,262,273]
% SZS output end Proof for query01_query02_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.021 s
% ------------------------------
% ------------------------------
