% Refutation found. Thanks to Tanya!
% SZS status Theorem for query03_query08_tptp4X
% SZS output start Proof for query03_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_4(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (where_5(X6) <=> (gte(X6,X6) & lte(X6,X6))) [input]
6. ! [X6] : (select_3(X6) <=> ? [X7,X8,X9,X10] : (where_5(X6) & from_4(X6,X7,X8,X9,X10))) [input]
7. ! [X11] : (main_query_2(X11) <=> select_3(X11)) [input]
11. ! [X12] : lte(X12,X12) [input]
17. ! [X12,X13] : (gte(X12,X13) <=> lte(X13,X12)) [input]
41. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
42. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
45. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
46. ! [X0] : (where_5(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 5]
47. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
48. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
52. ! [X0] : lte(X0,X0) [rectify 11]
58. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 17]
82. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 42]
134. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 82]
135. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 136]
138. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
139. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 137,138]
140. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 44]
141. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 45]
142. ! [X0] : ((where_5(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [nnf transformation 46]
143. ! [X0] : ((where_5(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [flattening 142]
144. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 47]
145. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 144]
146. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
147. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 145,146]
148. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 48]
151. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 58]
152. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 134]
153. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
154. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 152,153]
155. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 135]
156. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 135]
157. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 139]
158. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 139]
159. ~main_query_1(X0) | select_1(X0) [cnf transformation 140]
160. ~select_1(X0) | main_query_1(X0) [cnf transformation 140]
161. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 141]
162. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 141]
165. ~gte(X0,X0) | where_5(X0) | ~lte(X0,X0) [cnf transformation 143]
166. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 147]
168. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 147]
169. ~main_query_2(X0) | select_3(X0) [cnf transformation 148]
170. ~select_3(X0) | main_query_2(X0) [cnf transformation 148]
174. lte(X0,X0) [cnf transformation 52]
180. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 151]
204. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 154]
205. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 154]
230. 0 <=> main_query_1(sK8) [avatar definition]
231. main_query_1(sK8) <- {0} [avatar component clause 230]
233. 1 <=> ~main_query_1(sK8) [avatar definition]
234. ~main_query_1(sK8) <- {1} [avatar component clause 233]
236. 2 <=> main_query_2(sK8) [avatar definition]
237. main_query_2(sK8) <- {2} [avatar component clause 236]
239. 3 <=> ~main_query_2(sK8) [avatar definition]
240. ~main_query_2(sK8) <- {3} [avatar component clause 239]
241. ~1 | ~3 [avatar split clause 205,239,233]
242. 0 | 2 [avatar split clause 204,236,230]
243. select_3(sK8) <- {2} [resolution 169,237]
251. gte(X1,X1) [resolution 180,174]
253. where_5(X0) | ~lte(X0,X0) [resolution 165,251]
255. where_5(X0) [subsumption resolution 253,174]
261. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 157,155]
263. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 166,161]
265. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 261,162]
268. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 263,156]
270. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 265,168]
271. ~select_1(X1) | select_3(X1) [subsumption resolution 270,255]
273. ~select_3(X1) | select_1(X1) [resolution 268,158]
274. select_1(sK8) <- {2} [resolution 273,243]
276. main_query_1(sK8) <- {2} [resolution 274,160]
277. $false <- {1, 2} [subsumption resolution 276,234]
278. 1 | ~2 [avatar contradiction clause 277]
279. select_1(sK8) <- {0} [resolution 231,159]
280. select_3(sK8) <- {0} [resolution 279,271]
283. main_query_2(sK8) <- {0} [resolution 280,170]
285. $false <- {0, 3} [subsumption resolution 283,240]
286. ~0 | 3 [avatar contradiction clause 285]
287. $false [avatar sat refutation 241,242,278,286]
% SZS output end Proof for query03_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.014 s
% ------------------------------
% ------------------------------
