% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_04_query_09_tptp4X
% SZS output start Proof for query_04_query_09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X5,X6,X7,X8,X9] : (from_4(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
5. ! [X5] : (where_5(X5) <=> (gte(X5,X5) & lte(X5,X5))) [input]
6. ! [X5] : (select_3(X5) <=> ? [X6,X7,X8,X9] : (where_5(X5) & from_4(X5,X6,X7,X8,X9))) [input]
7. ! [X10] : (main_query_2(X10) <=> select_3(X10)) [input]
11. ! [X11] : lte(X11,X11) [input]
18. ! [X11,X12] : (gte(X11,X12) <=> lte(X12,X11)) [input]
39. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
40. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 39]
41. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
42. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
43. ! [X0] : (where_5(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 5]
44. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
45. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
49. ! [X0] : lte(X0,X0) [rectify 11]
56. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 18]
122. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 40]
123. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
124. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 41]
125. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 124]
126. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
127. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 125,126]
128. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
129. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 42]
130. ! [X0] : ((where_5(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [nnf transformation 43]
131. ! [X0] : ((where_5(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [flattening 130]
132. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 44]
133. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 132]
134. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
135. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 133,134]
136. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 45]
139. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 56]
140. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 122]
141. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
142. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 140,141]
143. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 123]
144. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 123]
145. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 127]
146. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 127]
147. ~main_query_1(X0) | select_1(X0) [cnf transformation 128]
148. ~select_1(X0) | main_query_1(X0) [cnf transformation 128]
149. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 129]
150. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 129]
153. ~gte(X0,X0) | where_5(X0) | ~lte(X0,X0) [cnf transformation 131]
154. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 135]
156. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 135]
157. ~main_query_2(X0) | select_3(X0) [cnf transformation 136]
158. ~select_3(X0) | main_query_2(X0) [cnf transformation 136]
162. lte(X0,X0) [cnf transformation 49]
169. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 139]
190. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 142]
191. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 142]
193. 0 <=> main_query_1(sK8) [avatar definition]
194. main_query_1(sK8) <- {0} [avatar component clause 193]
196. 1 <=> ~main_query_1(sK8) [avatar definition]
197. ~main_query_1(sK8) <- {1} [avatar component clause 196]
199. 2 <=> main_query_2(sK8) [avatar definition]
200. main_query_2(sK8) <- {2} [avatar component clause 199]
202. 3 <=> ~main_query_2(sK8) [avatar definition]
203. ~main_query_2(sK8) <- {3} [avatar component clause 202]
204. ~1 | ~3 [avatar split clause 191,202,196]
205. 0 | 2 [avatar split clause 190,199,193]
206. select_3(sK8) <- {2} [resolution 157,200]
215. gte(X1,X1) [resolution 169,162]
219. where_5(X0) | ~lte(X0,X0) [resolution 153,215]
221. where_5(X0) [subsumption resolution 219,162]
266. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 145,143]
289. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 154,149]
291. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 266,150]
294. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 289,144]
296. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 291,156]
297. ~select_1(X1) | select_3(X1) [subsumption resolution 296,221]
299. ~select_3(X1) | select_1(X1) [resolution 294,146]
302. select_1(sK8) <- {2} [resolution 299,206]
306. main_query_1(sK8) <- {2} [resolution 302,148]
307. $false <- {1, 2} [subsumption resolution 306,197]
308. 1 | ~2 [avatar contradiction clause 307]
309. select_1(sK8) <- {0} [resolution 194,147]
312. select_3(sK8) <- {0} [resolution 309,297]
317. main_query_2(sK8) <- {0} [resolution 312,158]
319. $false <- {0, 3} [subsumption resolution 317,203]
320. ~0 | 3 [avatar contradiction clause 319]
321. $false [avatar sat refutation 204,205,308,320]
% SZS output end Proof for query_04_query_09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.029 s
% ------------------------------
% ------------------------------
