% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_21_query_30_tptp4X
% SZS output start Proof for query_21_query_30_tptp4X
1. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) <=> (tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4))) [input]
2. ! [X0,X5] : (where_3(X0,X5) <=> eq(X0,X5)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_3(X0,X5) & from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [input]
4. ! [X10,X11] : (main_query_1(X10,X11) <=> select_1(X10,X11)) [input]
5. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) <=> (tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4))) [input]
6. ! [X5,X6,X7,X8,X9] : (from_9(X5,X6,X7,X8,X9) <=> tab2(X5,X6,X7,X8,X9)) [input]
7. ! [X5,X0] : (where_10(X5,X0) <=> neq(X5,X0)) [input]
8. ! [X5,X0] : (select_8(X5,X0) <=> ? [X6,X7,X8,X9] : (where_10(X5,X0) & from_9(X5,X6,X7,X8,X9))) [input]
9. ! [X12,X0] : (where_6_subquery_7(X12,X0) <=> select_8(X12,X0)) [input]
10. ! [X5,X6,X7,X8,X9] : (from_13(X5,X6,X7,X8,X9) <=> tab2(X5,X6,X7,X8,X9)) [input]
11. ! [X5,X0] : (where_14(X5,X0) <=> eq(X5,X0)) [input]
12. ! [X5,X0] : (select_12(X5,X0) <=> ? [X6,X7,X8,X9] : (where_14(X5,X0) & from_13(X5,X6,X7,X8,X9))) [input]
13. ! [X12,X0] : (where_6_subquery_11(X12,X0) <=> select_12(X12,X0)) [input]
14. ! [X0] : (where_6(X0) <=> (where_6_subquery_11(X0,X0) & ~where_6_subquery_7(X0,X0))) [input]
15. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_6(X0) & from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [input]
16. ! [X10,X11] : (main_query_2(X10,X11) <=> select_4(X10,X11)) [input]
17. ! [X13] : eq(X13,X13) [input]
18. ! [X13,X14] : (eq(X13,X14) => eq(X14,X13)) [input]
19. ! [X13,X14,X15] : ((eq(X14,X15) & eq(X13,X14)) => eq(X13,X15)) [input]
25. ! [X13,X14] : (neq(X13,X14) <=> ~eq(X13,X14)) [input]
27. ! [X13,X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) & eq(X0,X13)) => from_2(X13,X1,X2,X3,X4,X5,X6,X7,X8,X9)) [input]
28. ! [X13,X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) & eq(X1,X13)) => from_2(X0,X13,X2,X3,X4,X5,X6,X7,X8,X9)) [input]
37. ! [X13,X5,X0] : ((where_10(X5,X0) & eq(X5,X13)) => where_10(X13,X0)) [input]
38. ! [X13,X5,X0] : ((where_10(X5,X0) & eq(X0,X13)) => where_10(X5,X13)) [input]
39. ! [X13,X5,X0] : ((where_14(X5,X0) & eq(X5,X13)) => where_14(X13,X0)) [input]
40. ! [X13,X5,X0] : ((where_14(X5,X0) & eq(X0,X13)) => where_14(X5,X13)) [input]
41. ! [X13,X12,X0] : ((select_12(X12,X0) & eq(X12,X13)) => select_12(X13,X0)) [input]
52. ! [X13,X5,X6,X7,X8,X9] : ((from_13(X5,X6,X7,X8,X9) & eq(X5,X13)) => from_13(X13,X6,X7,X8,X9)) [input]
59. ! [X13,X10,X11] : ((select_1(X10,X11) & eq(X10,X13)) => select_1(X13,X11)) [input]
60. ! [X13,X10,X11] : ((select_1(X10,X11) & eq(X11,X13)) => select_1(X10,X13)) [input]
66. ! [X13,X12,X0] : ((select_8(X12,X0) & eq(X0,X13)) => select_8(X12,X13)) [input]
67. ! [X13,X0,X5] : ((where_3(X0,X5) & eq(X0,X13)) => where_3(X13,X5)) [input]
68. ! [X13,X0,X5] : ((where_3(X0,X5) & eq(X5,X13)) => where_3(X0,X13)) [input]
89. ! [X10,X11] : (main_query_1(X10,X11) <=> main_query_2(X10,X11)) [input]
90. ~! [X10,X11] : (main_query_1(X10,X11) <=> main_query_2(X10,X11)) [negated conjecture 89]
91. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [rectify 2]
92. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
93. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab2(X0,X1,X2,X3,X4)) [rectify 6]
94. ! [X0,X1] : (where_10(X0,X1) <=> neq(X0,X1)) [rectify 7]
95. ! [X0,X1] : (select_8(X0,X1) <=> ? [X2,X3,X4,X5] : (where_10(X0,X1) & from_9(X0,X2,X3,X4,X5))) [rectify 8]
96. ! [X0,X1] : (where_6_subquery_7(X0,X1) <=> select_8(X0,X1)) [rectify 9]
97. ! [X0,X1,X2,X3,X4] : (from_13(X0,X1,X2,X3,X4) <=> tab2(X0,X1,X2,X3,X4)) [rectify 10]
98. ! [X0,X1] : (where_14(X0,X1) <=> eq(X0,X1)) [rectify 11]
99. ! [X0,X1] : (select_12(X0,X1) <=> ? [X2,X3,X4,X5] : (where_14(X0,X1) & from_13(X0,X2,X3,X4,X5))) [rectify 12]
100. ! [X0,X1] : (where_6_subquery_11(X0,X1) <=> select_12(X0,X1)) [rectify 13]
101. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 16]
102. ! [X0] : eq(X0,X0) [rectify 17]
103. ! [X0,X1] : (eq(X0,X1) => eq(X1,X0)) [rectify 18]
104. ! [X0,X1,X2] : ((eq(X1,X2) & eq(X0,X1)) => eq(X0,X2)) [rectify 19]
110. ! [X0,X1] : (neq(X0,X1) <=> ~eq(X0,X1)) [rectify 25]
112. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ((from_2(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10) & eq(X1,X0)) => from_2(X0,X2,X3,X4,X5,X6,X7,X8,X9,X10)) [rectify 27]
113. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ((from_2(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10) & eq(X2,X0)) => from_2(X1,X0,X3,X4,X5,X6,X7,X8,X9,X10)) [rectify 28]
122. ! [X0,X1,X2] : ((where_10(X1,X2) & eq(X1,X0)) => where_10(X0,X2)) [rectify 37]
123. ! [X0,X1,X2] : ((where_10(X1,X2) & eq(X2,X0)) => where_10(X1,X0)) [rectify 38]
124. ! [X0,X1,X2] : ((where_14(X1,X2) & eq(X1,X0)) => where_14(X0,X2)) [rectify 39]
125. ! [X0,X1,X2] : ((where_14(X1,X2) & eq(X2,X0)) => where_14(X1,X0)) [rectify 40]
126. ! [X0,X1,X2] : ((select_12(X1,X2) & eq(X1,X0)) => select_12(X0,X2)) [rectify 41]
137. ! [X0,X1,X2,X3,X4,X5] : ((from_13(X1,X2,X3,X4,X5) & eq(X1,X0)) => from_13(X0,X2,X3,X4,X5)) [rectify 52]
144. ! [X0,X1,X2] : ((select_1(X1,X2) & eq(X1,X0)) => select_1(X0,X2)) [rectify 59]
145. ! [X0,X1,X2] : ((select_1(X1,X2) & eq(X2,X0)) => select_1(X1,X0)) [rectify 60]
151. ! [X0,X1,X2] : ((select_8(X1,X2) & eq(X2,X0)) => select_8(X1,X0)) [rectify 66]
152. ! [X0,X1,X2] : ((where_3(X1,X2) & eq(X1,X0)) => where_3(X0,X2)) [rectify 67]
153. ! [X0,X1,X2] : ((where_3(X1,X2) & eq(X2,X0)) => where_3(X1,X0)) [rectify 68]
174. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 90]
175. ! [X0,X1] : (eq(X1,X0) | ~eq(X0,X1)) [ennf transformation 103]
176. ! [X0,X1,X2] : (eq(X0,X2) | (~eq(X1,X2) | ~eq(X0,X1))) [ennf transformation 104]
177. ! [X0,X1,X2] : (eq(X0,X2) | ~eq(X1,X2) | ~eq(X0,X1)) [flattening 176]
180. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (from_2(X0,X2,X3,X4,X5,X6,X7,X8,X9,X10) | (~from_2(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10) | ~eq(X1,X0))) [ennf transformation 112]
181. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (from_2(X0,X2,X3,X4,X5,X6,X7,X8,X9,X10) | ~from_2(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10) | ~eq(X1,X0)) [flattening 180]
182. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (from_2(X1,X0,X3,X4,X5,X6,X7,X8,X9,X10) | (~from_2(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10) | ~eq(X2,X0))) [ennf transformation 113]
183. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (from_2(X1,X0,X3,X4,X5,X6,X7,X8,X9,X10) | ~from_2(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10) | ~eq(X2,X0)) [flattening 182]
200. ! [X0,X1,X2] : (where_10(X0,X2) | (~where_10(X1,X2) | ~eq(X1,X0))) [ennf transformation 122]
201. ! [X0,X1,X2] : (where_10(X0,X2) | ~where_10(X1,X2) | ~eq(X1,X0)) [flattening 200]
202. ! [X0,X1,X2] : (where_10(X1,X0) | (~where_10(X1,X2) | ~eq(X2,X0))) [ennf transformation 123]
203. ! [X0,X1,X2] : (where_10(X1,X0) | ~where_10(X1,X2) | ~eq(X2,X0)) [flattening 202]
204. ! [X0,X1,X2] : (where_14(X0,X2) | (~where_14(X1,X2) | ~eq(X1,X0))) [ennf transformation 124]
205. ! [X0,X1,X2] : (where_14(X0,X2) | ~where_14(X1,X2) | ~eq(X1,X0)) [flattening 204]
206. ! [X0,X1,X2] : (where_14(X1,X0) | (~where_14(X1,X2) | ~eq(X2,X0))) [ennf transformation 125]
207. ! [X0,X1,X2] : (where_14(X1,X0) | ~where_14(X1,X2) | ~eq(X2,X0)) [flattening 206]
208. ! [X0,X1,X2] : (select_12(X0,X2) | (~select_12(X1,X2) | ~eq(X1,X0))) [ennf transformation 126]
209. ! [X0,X1,X2] : (select_12(X0,X2) | ~select_12(X1,X2) | ~eq(X1,X0)) [flattening 208]
230. ! [X0,X1,X2,X3,X4,X5] : (from_13(X0,X2,X3,X4,X5) | (~from_13(X1,X2,X3,X4,X5) | ~eq(X1,X0))) [ennf transformation 137]
231. ! [X0,X1,X2,X3,X4,X5] : (from_13(X0,X2,X3,X4,X5) | ~from_13(X1,X2,X3,X4,X5) | ~eq(X1,X0)) [flattening 230]
244. ! [X0,X1,X2] : (select_1(X0,X2) | (~select_1(X1,X2) | ~eq(X1,X0))) [ennf transformation 144]
245. ! [X0,X1,X2] : (select_1(X0,X2) | ~select_1(X1,X2) | ~eq(X1,X0)) [flattening 244]
246. ! [X0,X1,X2] : (select_1(X1,X0) | (~select_1(X1,X2) | ~eq(X2,X0))) [ennf transformation 145]
247. ! [X0,X1,X2] : (select_1(X1,X0) | ~select_1(X1,X2) | ~eq(X2,X0)) [flattening 246]
258. ! [X0,X1,X2] : (select_8(X1,X0) | (~select_8(X1,X2) | ~eq(X2,X0))) [ennf transformation 151]
259. ! [X0,X1,X2] : (select_8(X1,X0) | ~select_8(X1,X2) | ~eq(X2,X0)) [flattening 258]
260. ! [X0,X1,X2] : (where_3(X0,X2) | (~where_3(X1,X2) | ~eq(X1,X0))) [ennf transformation 152]
261. ! [X0,X1,X2] : (where_3(X0,X2) | ~where_3(X1,X2) | ~eq(X1,X0)) [flattening 260]
262. ! [X0,X1,X2] : (where_3(X1,X0) | (~where_3(X1,X2) | ~eq(X2,X0))) [ennf transformation 153]
263. ! [X0,X1,X2] : (where_3(X1,X0) | ~where_3(X1,X2) | ~eq(X2,X0)) [flattening 262]
304. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 174]
305. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | (~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4))) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [nnf transformation 1]
306. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4)) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [flattening 305]
307. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 91]
308. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_3(X0,X5) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & (? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_3(X0,X5) & from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9)) | ~select_1(X0,X1))) [nnf transformation 3]
309. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_3(X0,X5) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_3(X0,X13) & from_2(X0,X1,X10,X11,X12,X13,X14,X15,X16,X17)) | ~select_1(X0,X1))) [rectify 308]
310. ! [X1,X0] : (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_3(X0,X13) & from_2(X0,X1,X10,X11,X12,X13,X14,X15,X16,X17)) => (where_3(X0,sK3(X0,X1)) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1),sK5(X0,X1),sK6(X0,X1),sK7(X0,X1)))) [choice axiom]
311. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_3(X0,X5) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & ((where_3(X0,sK3(X0,X1)) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1),sK5(X0,X1),sK6(X0,X1),sK7(X0,X1))) | ~select_1(X0,X1))) [skolemisation 309,310]
312. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 92]
313. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | (~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4))) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [nnf transformation 5]
314. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4)) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [flattening 313]
315. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab2(X0,X1,X2,X3,X4)) & (tab2(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 93]
316. ! [X0,X1] : ((where_10(X0,X1) | ~neq(X0,X1)) & (neq(X0,X1) | ~where_10(X0,X1))) [nnf transformation 94]
317. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5] : (~where_10(X0,X1) | ~from_9(X0,X2,X3,X4,X5))) & (? [X2,X3,X4,X5] : (where_10(X0,X1) & from_9(X0,X2,X3,X4,X5)) | ~select_8(X0,X1))) [nnf transformation 95]
318. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5] : (~where_10(X0,X1) | ~from_9(X0,X2,X3,X4,X5))) & (? [X6,X7,X8,X9] : (where_10(X0,X1) & from_9(X0,X6,X7,X8,X9)) | ~select_8(X0,X1))) [rectify 317]
319. ! [X1,X0] : (? [X6,X7,X8,X9] : (where_10(X0,X1) & from_9(X0,X6,X7,X8,X9)) => (where_10(X0,X1) & from_9(X0,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1),sK11(X0,X1)))) [choice axiom]
320. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5] : (~where_10(X0,X1) | ~from_9(X0,X2,X3,X4,X5))) & ((where_10(X0,X1) & from_9(X0,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1),sK11(X0,X1))) | ~select_8(X0,X1))) [skolemisation 318,319]
321. ! [X0,X1] : ((where_6_subquery_7(X0,X1) | ~select_8(X0,X1)) & (select_8(X0,X1) | ~where_6_subquery_7(X0,X1))) [nnf transformation 96]
322. ! [X0,X1,X2,X3,X4] : ((from_13(X0,X1,X2,X3,X4) | ~tab2(X0,X1,X2,X3,X4)) & (tab2(X0,X1,X2,X3,X4) | ~from_13(X0,X1,X2,X3,X4))) [nnf transformation 97]
323. ! [X0,X1] : ((where_14(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_14(X0,X1))) [nnf transformation 98]
324. ! [X0,X1] : ((select_12(X0,X1) | ! [X2,X3,X4,X5] : (~where_14(X0,X1) | ~from_13(X0,X2,X3,X4,X5))) & (? [X2,X3,X4,X5] : (where_14(X0,X1) & from_13(X0,X2,X3,X4,X5)) | ~select_12(X0,X1))) [nnf transformation 99]
325. ! [X0,X1] : ((select_12(X0,X1) | ! [X2,X3,X4,X5] : (~where_14(X0,X1) | ~from_13(X0,X2,X3,X4,X5))) & (? [X6,X7,X8,X9] : (where_14(X0,X1) & from_13(X0,X6,X7,X8,X9)) | ~select_12(X0,X1))) [rectify 324]
326. ! [X1,X0] : (? [X6,X7,X8,X9] : (where_14(X0,X1) & from_13(X0,X6,X7,X8,X9)) => (where_14(X0,X1) & from_13(X0,sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1)))) [choice axiom]
327. ! [X0,X1] : ((select_12(X0,X1) | ! [X2,X3,X4,X5] : (~where_14(X0,X1) | ~from_13(X0,X2,X3,X4,X5))) & ((where_14(X0,X1) & from_13(X0,sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1))) | ~select_12(X0,X1))) [skolemisation 325,326]
328. ! [X0,X1] : ((where_6_subquery_11(X0,X1) | ~select_12(X0,X1)) & (select_12(X0,X1) | ~where_6_subquery_11(X0,X1))) [nnf transformation 100]
329. ! [X0] : ((where_6(X0) | (~where_6_subquery_11(X0,X0) | where_6_subquery_7(X0,X0))) & ((where_6_subquery_11(X0,X0) & ~where_6_subquery_7(X0,X0)) | ~where_6(X0))) [nnf transformation 14]
330. ! [X0] : ((where_6(X0) | ~where_6_subquery_11(X0,X0) | where_6_subquery_7(X0,X0)) & ((where_6_subquery_11(X0,X0) & ~where_6_subquery_7(X0,X0)) | ~where_6(X0))) [flattening 329]
331. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & (? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_6(X0) & from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9)) | ~select_4(X0,X1))) [nnf transformation 15]
332. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_6(X0) & from_5(X0,X1,X10,X11,X12,X13,X14,X15,X16,X17)) | ~select_4(X0,X1))) [rectify 331]
333. ! [X1,X0] : (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_6(X0) & from_5(X0,X1,X10,X11,X12,X13,X14,X15,X16,X17)) => (where_6(X0) & from_5(X0,X1,sK16(X0,X1),sK17(X0,X1),sK18(X0,X1),sK19(X0,X1),sK20(X0,X1),sK21(X0,X1),sK22(X0,X1),sK23(X0,X1)))) [choice axiom]
334. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & ((where_6(X0) & from_5(X0,X1,sK16(X0,X1),sK17(X0,X1),sK18(X0,X1),sK19(X0,X1),sK20(X0,X1),sK21(X0,X1),sK22(X0,X1),sK23(X0,X1))) | ~select_4(X0,X1))) [skolemisation 332,333]
335. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 101]
338. ! [X0,X1] : ((neq(X0,X1) | eq(X0,X1)) & (~eq(X0,X1) | ~neq(X0,X1))) [nnf transformation 110]
339. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 304]
340. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK24,sK25) | ~main_query_1(sK24,sK25)) & (main_query_2(sK24,sK25) | main_query_1(sK24,sK25))) [choice axiom]
341. (~main_query_2(sK24,sK25) | ~main_query_1(sK24,sK25)) & (main_query_2(sK24,sK25) | main_query_1(sK24,sK25)) [skolemisation 339,340]
342. ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | tab1(X0,X1,X2,X3,X4) [cnf transformation 306]
343. ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | tab2(X5,X6,X7,X8,X9) [cnf transformation 306]
344. from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4) [cnf transformation 306]
345. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 307]
346. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 307]
347. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1),sK5(X0,X1),sK6(X0,X1),sK7(X0,X1)) | ~select_1(X0,X1) [cnf transformation 311]
348. where_3(X0,sK3(X0,X1)) | ~select_1(X0,X1) [cnf transformation 311]
349. ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~where_3(X0,X5) | select_1(X0,X1) [cnf transformation 311]
350. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 312]
351. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 312]
352. ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | tab1(X0,X1,X2,X3,X4) [cnf transformation 314]
354. from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4) [cnf transformation 314]
355. ~from_9(X0,X1,X2,X3,X4) | tab2(X0,X1,X2,X3,X4) [cnf transformation 315]
356. ~tab2(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 315]
357. ~where_10(X0,X1) | neq(X0,X1) [cnf transformation 316]
358. ~neq(X0,X1) | where_10(X0,X1) [cnf transformation 316]
359. from_9(X0,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1),sK11(X0,X1)) | ~select_8(X0,X1) [cnf transformation 320]
360. ~select_8(X0,X1) | where_10(X0,X1) [cnf transformation 320]
361. ~from_9(X0,X2,X3,X4,X5) | ~where_10(X0,X1) | select_8(X0,X1) [cnf transformation 320]
362. ~where_6_subquery_7(X0,X1) | select_8(X0,X1) [cnf transformation 321]
364. ~from_13(X0,X1,X2,X3,X4) | tab2(X0,X1,X2,X3,X4) [cnf transformation 322]
365. ~tab2(X0,X1,X2,X3,X4) | from_13(X0,X1,X2,X3,X4) [cnf transformation 322]
366. ~where_14(X0,X1) | eq(X0,X1) [cnf transformation 323]
367. ~eq(X0,X1) | where_14(X0,X1) [cnf transformation 323]
368. from_13(X0,sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1)) | ~select_12(X0,X1) [cnf transformation 327]
370. ~from_13(X0,X2,X3,X4,X5) | ~where_14(X0,X1) | select_12(X0,X1) [cnf transformation 327]
371. ~where_6_subquery_11(X0,X1) | select_12(X0,X1) [cnf transformation 328]
372. ~select_12(X0,X1) | where_6_subquery_11(X0,X1) [cnf transformation 328]
374. where_6_subquery_11(X0,X0) | ~where_6(X0) [cnf transformation 330]
375. ~where_6_subquery_11(X0,X0) | where_6(X0) | where_6_subquery_7(X0,X0) [cnf transformation 330]
376. from_5(X0,X1,sK16(X0,X1),sK17(X0,X1),sK18(X0,X1),sK19(X0,X1),sK20(X0,X1),sK21(X0,X1),sK22(X0,X1),sK23(X0,X1)) | ~select_4(X0,X1) [cnf transformation 334]
377. ~select_4(X0,X1) | where_6(X0) [cnf transformation 334]
378. ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~where_6(X0) | select_4(X0,X1) [cnf transformation 334]
379. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 335]
380. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 335]
381. eq(X0,X0) [cnf transformation 102]
382. ~eq(X0,X1) | eq(X1,X0) [cnf transformation 175]
383. ~eq(X1,X2) | eq(X0,X2) | ~eq(X0,X1) [cnf transformation 177]
389. ~neq(X0,X1) | ~eq(X0,X1) [cnf transformation 338]
390. neq(X0,X1) | eq(X0,X1) [cnf transformation 338]
391. ~from_2(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10) | from_2(X0,X2,X3,X4,X5,X6,X7,X8,X9,X10) | ~eq(X1,X0) [cnf transformation 181]
392. ~from_2(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10) | from_2(X1,X0,X3,X4,X5,X6,X7,X8,X9,X10) | ~eq(X2,X0) [cnf transformation 183]
401. ~where_10(X1,X2) | where_10(X0,X2) | ~eq(X1,X0) [cnf transformation 201]
402. ~where_10(X1,X2) | where_10(X1,X0) | ~eq(X2,X0) [cnf transformation 203]
403. ~where_14(X1,X2) | where_14(X0,X2) | ~eq(X1,X0) [cnf transformation 205]
404. ~where_14(X1,X2) | where_14(X1,X0) | ~eq(X2,X0) [cnf transformation 207]
405. ~select_12(X1,X2) | select_12(X0,X2) | ~eq(X1,X0) [cnf transformation 209]
416. ~from_13(X1,X2,X3,X4,X5) | from_13(X0,X2,X3,X4,X5) | ~eq(X1,X0) [cnf transformation 231]
423. ~select_1(X1,X2) | select_1(X0,X2) | ~eq(X1,X0) [cnf transformation 245]
424. ~select_1(X1,X2) | select_1(X1,X0) | ~eq(X2,X0) [cnf transformation 247]
430. ~select_8(X1,X2) | select_8(X1,X0) | ~eq(X2,X0) [cnf transformation 259]
431. ~eq(X1,X0) | ~where_3(X1,X2) | where_3(X0,X2) [cnf transformation 261]
432. ~eq(X2,X0) | ~where_3(X1,X2) | where_3(X1,X0) [cnf transformation 263]
453. main_query_2(sK24,sK25) | main_query_1(sK24,sK25) [cnf transformation 341]
454. ~main_query_2(sK24,sK25) | ~main_query_1(sK24,sK25) [cnf transformation 341]
456. 0 <=> main_query_1(sK24,sK25) [avatar definition]
457. main_query_1(sK24,sK25) <- {0} [avatar component clause 456]
459. 1 <=> ~main_query_1(sK24,sK25) [avatar definition]
460. ~main_query_1(sK24,sK25) <- {1} [avatar component clause 459]
462. 2 <=> main_query_2(sK24,sK25) [avatar definition]
463. main_query_2(sK24,sK25) <- {2} [avatar component clause 462]
465. 3 <=> ~main_query_2(sK24,sK25) [avatar definition]
466. ~main_query_2(sK24,sK25) <- {3} [avatar component clause 465]
467. ~1 | ~3 [avatar split clause 454,465,459]
468. 0 | 2 [avatar split clause 453,462,456]
469. where_3(X0,X0) [resolution 346,381]
471. where_14(X0,X0) [resolution 367,381]
473. select_12(X0,X0) | ~where_6(X0) [resolution 371,374]
483. where_10(X2,X3) | eq(X2,X3) [resolution 390,358]
484. eq(X0,sK3(X0,X1)) | ~select_1(X0,X1) [resolution 348,345]
489. eq(sK3(X4,X5),X4) | ~select_1(X4,X5) [resolution 484,382]
490. where_14(X6,sK3(X6,X7)) | ~select_1(X6,X7) [resolution 484,367]
493. eq(X2,sK3(X3,X4)) | ~eq(X2,X3) | ~select_1(X3,X4) [resolution 383,484]
504. where_14(sK3(X9,X10),X9) | ~select_1(X9,X10) [resolution 489,367]
506. ~eq(X2,X0) | where_10(X0,X1) | eq(X2,X1) [resolution 401,483]
510. ~eq(X1,X0) | where_14(X0,X1) [resolution 403,471]
511. where_14(X2,sK3(X3,X4)) | ~eq(X3,X2) | ~select_1(X3,X4) [resolution 403,490]
517. ~eq(sK3(X2,X4),X3) | where_14(X2,X3) | ~select_1(X2,X4) [resolution 404,490]
560. where_3(sK3(X2,X4),X3) | ~where_3(X2,X3) | ~select_1(X2,X4) [resolution 431,484]
566. where_3(X2,sK3(X3,X4)) | ~where_3(X2,X3) | ~select_1(X3,X4) [resolution 432,484]
567. ~where_3(X5,sK3(X6,X7)) | where_3(X5,X6) | ~select_1(X6,X7) [resolution 432,489]
569. where_10(sK3(X2,X3),X4) | eq(X2,X4) | ~select_1(X2,X3) [resolution 506,484]
570. eq(sK3(X5,X7),X6) | where_10(X5,X6) | ~select_1(X5,X7) [resolution 506,489]
591. ~select_1(X10,X11) | ~eq(X9,X10) | eq(X12,sK3(X10,X11)) | ~eq(X12,X9) [resolution 493,383]
607. eq(sK3(X0,X2),X1) | ~select_1(X0,X2) | ~where_3(X0,X1) [resolution 560,345]
608. eq(X0,sK3(X1,X2)) | ~select_1(X1,X2) | ~where_3(X0,X1) [resolution 566,345]
647. eq(X9,sK3(X8,X10)) | ~select_1(X8,X10) | ~eq(X8,X9) [resolution 511,366]
671. where_10(X2,X3) | ~select_1(X2,X4) | where_14(X2,X3) | ~select_1(X2,X4) [resolution 570,517]
688. ~select_1(X2,X4) | where_10(X2,X3) | where_14(X2,X3) [duplicate literal removal 671]
691. 4 <=> ! [X1] : ~select_1(sK25,X1) [avatar definition]
692. ~select_1(sK25,X1) <- {4} [avatar component clause 691]
754. tab2(X0,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1),sK11(X0,X1)) | ~select_8(X0,X1) [resolution 359,355]
761. tab2(X0,sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1)) | ~select_12(X0,X1) [resolution 368,364]
763. from_13(X7,sK12(X5,X6),sK13(X5,X6),sK14(X5,X6),sK15(X5,X6)) | ~select_12(X5,X6) | ~eq(X5,X7) [resolution 368,416]
771. ~tab2(X0,X1,X2,X3,X4) | ~tab1(X5,X6,X7,X8,X9) | ~where_3(X5,X0) | select_1(X5,X6) [resolution 344,349]
781. ~tab2(X20,X21,X22,X23,X24) | ~tab1(X25,X26,X27,X28,X29) | ~where_6(X25) | select_4(X25,X26) [resolution 354,378]
783. 8 <=> ! [X25,X27,X29,X26,X28] : (~tab1(X25,X26,X27,X28,X29) | select_4(X25,X26) | ~where_6(X25)) [avatar definition]
784. ~tab1(X25,X26,X27,X28,X29) | select_4(X25,X26) | ~where_6(X25) <- {8} [avatar component clause 783]
786. 10 <=> ! [X20,X22,X21,X23,X24] : ~tab2(X20,X21,X22,X23,X24) [avatar definition]
787. ~tab2(X20,X21,X22,X23,X24) <- {10} [avatar component clause 786]
788. 8 | 10 [avatar split clause 781,786,783]
800. ~select_1(X2,X3) | ~where_3(X2,X4) | where_14(X2,X4) | ~select_1(X2,X3) [resolution 607,517]
817. ~select_1(X2,X3) | ~where_3(X2,X4) | where_14(X2,X4) [duplicate literal removal 800]
824. eq(sK3(X9,X10),X11) | ~where_3(X11,X9) | ~select_1(X9,X10) [resolution 608,382]
864. 12 <=> ! [X2] : ~select_1(sK24,X2) [avatar definition]
865. ~select_1(sK24,X2) <- {12} [avatar component clause 864]
922. 20 <=> where_6(sK24) [avatar definition]
923. where_6(sK24) <- {20} [avatar component clause 922]
925. 21 <=> ~where_6(sK24) [avatar definition]
926. ~where_6(sK24) <- {21} [avatar component clause 925]
928. 22 <=> ! [X0] : (where_10(sK24,X0) | select_12(sK24,X0)) [avatar definition]
929. select_12(sK24,X0) | where_10(sK24,X0) <- {22} [avatar component clause 928]
962. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 347,342]
963. tab2(sK3(X2,X3),sK4(X2,X3),sK5(X2,X3),sK6(X2,X3),sK7(X2,X3)) | ~select_1(X2,X3) [resolution 347,343]
965. from_2(X8,X7,sK0(X6,X7),sK1(X6,X7),sK2(X6,X7),sK3(X6,X7),sK4(X6,X7),sK5(X6,X7),sK6(X6,X7),sK7(X6,X7)) | ~select_1(X6,X7) | ~eq(X6,X8) [resolution 347,391]
978. tab1(X0,X1,sK16(X0,X1),sK17(X0,X1),sK18(X0,X1)) | ~select_4(X0,X1) [resolution 376,352]
996. ~select_1(X0,X1) | select_4(X0,X1) | ~where_6(X0) <- {8} [resolution 962,784]
1018. from_13(X2,sK8(X2,X3),sK9(X2,X3),sK10(X2,X3),sK11(X2,X3)) | ~select_8(X2,X3) [resolution 754,365]
1027. from_9(X0,sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1)) | ~select_12(X0,X1) [resolution 761,356]
1061. select_4(sK24,sK25) <- {2} [resolution 463,379]
1065. where_6(sK24) <- {2} [resolution 1061,377]
1068. 20 | ~2 [avatar split clause 1065,462,922]
1085. ~tab1(X7,X8,X9,X10,X11) | ~where_3(X7,X12) | select_1(X7,X8) | ~select_12(X12,X13) [resolution 771,761]
1091. from_9(sK3(X0,X1),sK4(X0,X1),sK5(X0,X1),sK6(X0,X1),sK7(X0,X1)) | ~select_1(X0,X1) [resolution 963,356]
1092. from_13(sK3(X2,X3),sK4(X2,X3),sK5(X2,X3),sK6(X2,X3),sK7(X2,X3)) | ~select_1(X2,X3) [resolution 963,365]
1163. ~select_12(X3,X4) | ~eq(X3,X5) | ~where_14(X5,X6) | select_12(X5,X6) [resolution 763,370]
1269. where_3(X5,sK3(X3,X4)) | ~eq(X3,X5) | ~select_1(X3,X4) [resolution 647,346]
1393. from_2(X15,X16,sK0(X13,X14),sK1(X13,X14),sK2(X13,X14),sK3(X13,X14),sK4(X13,X14),sK5(X13,X14),sK6(X13,X14),sK7(X13,X14)) | ~eq(X13,X15) | ~select_1(X13,X14) | ~eq(X14,X16) [resolution 965,392]
1750. ~where_14(X1,X2) | ~eq(X0,X1) | select_12(X1,X2) | ~where_6(X0) [resolution 1163,473]
1770. ~select_4(X4,X6) | select_1(X4,X6) | ~select_12(X5,X7) | ~where_3(X4,X5) [resolution 1085,978]
1786. ~where_14(X2,X4) | ~select_8(X2,X3) | select_12(X2,X4) [resolution 1018,370]
1807. ~select_12(X2,X3) | ~where_10(X2,X4) | select_8(X2,X4) [resolution 1027,361]
1815. ~where_10(X0,X1) | select_8(X0,X1) | ~where_6(X0) [resolution 1807,473]
1820. select_8(X0,X1) | ~where_6(X0) | eq(X0,X1) [resolution 1815,483]
1836. ~eq(X8,X9) | eq(X7,X8) | select_8(X7,X9) | ~where_6(X7) [resolution 1820,430]
1946. ~where_10(sK3(X2,X3),X4) | ~select_1(X2,X3) | select_8(sK3(X2,X3),X4) [resolution 1091,361]
1956. ~where_14(sK3(X2,X3),X4) | ~select_1(X2,X3) | select_12(sK3(X2,X3),X4) [resolution 1092,370]
2584. ~eq(X0,X1) | ~select_1(X0,X2) | where_3(X1,X0) | ~select_1(X0,X2) [resolution 1269,567]
2587. ~select_1(X0,X2) | ~eq(X0,X1) | where_3(X1,X0) [duplicate literal removal 2584]
3004. 32 <=> ! [X0] : select_1(X0,sK25) [avatar definition]
3005. select_1(X0,sK25) <- {32} [avatar component clause 3004]
3020. 34 <=> ! [X11,X10] : ~select_1(X10,X11) [avatar definition]
3021. ~select_1(X10,X11) <- {34} [avatar component clause 3020]
3476. $false <- {4, 32} [resolution 3005,692]
3479. ~4 | ~32 [avatar contradiction clause 3476]
3857. select_1(sK24,sK25) <- {0} [resolution 457,350]
3858. $false <- {0, 12} [subsumption resolution 3857,865]
3859. ~0 | ~12 [avatar contradiction clause 3858]
3860. $false <- {0, 34} [subsumption resolution 3857,3021]
3861. ~0 | ~34 [avatar contradiction clause 3860]
3863. ~eq(sK24,X0) | select_1(X0,sK25) <- {0} [resolution 3857,423]
3864. ~eq(sK25,X1) | select_1(sK24,X1) <- {0} [resolution 3857,424]
3865. where_14(sK24,X2) | where_10(sK24,X2) <- {0} [resolution 3857,688]
3867. select_4(sK24,sK25) | ~where_6(sK24) <- {0, 8} [resolution 3857,996]
3885. ~eq(sK24,X21) | where_3(X21,sK24) <- {0} [resolution 3857,2587]
3892. select_4(sK24,sK25) <- {0, 8, 20} [subsumption resolution 3867,923]
3898. main_query_2(sK24,sK25) <- {0, 8, 20} [resolution 3892,380]
3900. $false <- {0, 3, 8, 20} [subsumption resolution 3898,466]
3901. ~0 | 3 | ~8 | ~20 [avatar contradiction clause 3900]
3952. 38 <=> ! [X7,X8] : eq(X7,X8) [avatar definition]
3953. eq(X7,X8) <- {38} [avatar component clause 3952]
3955. 40 <=> ! [X6] : ~where_6(X6) [avatar definition]
3956. ~where_6(X6) <- {40} [avatar component clause 3955]
3960. 42 <=> ! [X15,X14] : where_3(X15,X14) [avatar definition]
3961. where_3(X15,X14) <- {42} [avatar component clause 3960]
4002. ~select_1(X40,X41) <- {10} [resolution 787,963]
4013. 34 | ~10 [avatar split clause 4002,786,3020]
4027. select_1(sK3(sK24,X0),sK25) | ~select_1(sK24,X0) <- {0} [resolution 3863,484]
4035. select_1(sK24,sK3(sK25,X0)) | ~select_1(sK25,X0) <- {0} [resolution 3864,484]
4091. where_10(sK24,X2) | ~eq(X3,sK24) | select_12(sK24,X2) | ~where_6(X3) <- {0} [resolution 3865,1750]
4145. where_6_subquery_11(sK24,X1) | where_10(sK24,X1) <- {22} [resolution 929,372]
4156. 52 <=> ! [X6] : where_10(sK24,X6) [avatar definition]
4157. where_10(sK24,X6) <- {52} [avatar component clause 4156]
4300. 62 <=> ! [X3] : (~eq(X3,sK24) | ~where_6(X3)) [avatar definition]
4301. ~eq(X3,sK24) | ~where_6(X3) <- {62} [avatar component clause 4300]
4604. where_10(sK24,sK24) | where_6(sK24) | where_6_subquery_7(sK24,sK24) <- {22} [resolution 4145,375]
4605. where_10(sK24,sK24) | where_6_subquery_7(sK24,sK24) <- {21, 22} [subsumption resolution 4604,926]
4607. 65 <=> ~where_6_subquery_7(sK24,sK24) [avatar definition]
4608. ~where_6_subquery_7(sK24,sK24) <- {65} [avatar component clause 4607]
4610. 64 <=> where_6_subquery_7(sK24,sK24) [avatar definition]
4611. where_6_subquery_7(sK24,sK24) <- {64} [avatar component clause 4610]
4613. 67 <=> ~where_10(sK24,sK24) [avatar definition]
4614. ~where_10(sK24,sK24) <- {67} [avatar component clause 4613]
4616. 66 <=> where_10(sK24,sK24) [avatar definition]
4617. where_10(sK24,sK24) <- {66} [avatar component clause 4616]
4618. 64 | 66 | 21 | ~22 [avatar split clause 4605,928,925,4616,4610]
4633. where_10(sK24,X0) | ~eq(sK24,X0) <- {66} [resolution 4617,402]
4636. where_10(sK24,X0) <- {66} [subsumption resolution 4633,483]
4637. 52 | ~66 [avatar split clause 4636,4616,4156]
4655. neq(sK24,X5) <- {52} [resolution 4157,357]
4669. ~eq(sK24,X0) <- {52} [resolution 4655,389]
4715. $false <- {52} [resolution 4669,381]
4723. ~52 [avatar contradiction clause 4715]
4740. select_8(sK24,sK24) <- {64} [resolution 4611,362]
4773. where_10(sK24,sK24) <- {64} [resolution 4740,360]
4782. $false <- {64, 67} [subsumption resolution 4773,4614]
4783. ~64 | 67 [avatar contradiction clause 4782]
4784. 62 | 22 | ~0 [avatar split clause 4091,456,928,4300]
5125. ~eq(sK3(sK24,X1),X2) | select_1(X2,sK25) | ~select_1(sK24,X1) <- {0} [resolution 4027,423]
5164. ~eq(X8,X9) | ~select_1(X8,X10) | ~eq(X10,X11) | ~where_3(X9,sK3(X8,X10)) | select_1(X9,X11) [resolution 1393,349]
5175. ~select_1(X8,X10) | ~eq(X8,X9) | ~eq(X10,X11) | select_1(X9,X11) [subsumption resolution 5164,1269]
6673. select_1(X9,sK25) | ~select_1(sK24,X10) | ~where_3(X9,sK24) | ~select_1(sK24,X10) <- {0} [resolution 5125,824]
6684. select_1(X9,sK25) | ~select_1(sK24,X10) | ~where_3(X9,sK24) <- {0} [duplicate literal removal 6673]
6707. 86 <=> ! [X9] : (select_1(X9,sK25) | ~where_3(X9,sK24)) [avatar definition]
6708. ~where_3(X9,sK24) | select_1(X9,sK25) <- {86} [avatar component clause 6707]
6709. 12 | 86 | ~0 [avatar split clause 6684,456,6707,864]
6801. select_1(sK24,sK25) <- {86} [resolution 6708,469]
7881. ~select_1(X0,X1) | select_8(sK3(X0,X1),X2) | eq(X0,X2) | ~select_1(X0,X1) [resolution 1946,569]
7899. select_8(sK3(X0,X1),X2) | ~select_1(X0,X1) | eq(X0,X2) [duplicate literal removal 7881]
7973. ~eq(X1,X0) | eq(X1,sK3(sK24,sK25)) | ~eq(X0,sK24) <- {0} [resolution 591,3857]
8049. eq(X0,sK3(sK24,sK25)) | ~eq(X0,sK24) <- {0} [resolution 7973,381]
8090. where_14(sK3(sK24,sK25),X16) | ~eq(X16,sK24) <- {0} [resolution 8049,510]
8175. 89 <=> ~where_14(sK24,sK3(sK24,sK25)) [avatar definition]
8176. ~where_14(sK24,sK3(sK24,sK25)) <- {89} [avatar component clause 8175]
8178. 88 <=> where_14(sK24,sK3(sK24,sK25)) [avatar definition]
8179. where_14(sK24,sK3(sK24,sK25)) <- {88} [avatar component clause 8178]
8183. 91 <=> ~where_14(sK3(sK24,sK25),sK24) [avatar definition]
8184. ~where_14(sK3(sK24,sK25),sK24) <- {91} [avatar component clause 8183]
8186. 90 <=> where_14(sK3(sK24,sK25),sK24) [avatar definition]
8187. where_14(sK3(sK24,sK25),sK24) <- {90} [avatar component clause 8186]
8232. 97 <=> ~where_6(sK3(sK24,sK25)) [avatar definition]
8233. ~where_6(sK3(sK24,sK25)) <- {97} [avatar component clause 8232]
8346. eq(sK24,sK3(sK24,sK25)) <- {88} [resolution 8179,366]
8353. eq(sK3(sK24,sK25),sK24) <- {90} [resolution 8187,366]
8366. 110 <=> ! [X2] : (~eq(X2,sK3(sK24,sK25)) | ~where_6(X2)) [avatar definition]
8367. ~eq(X2,sK3(sK24,sK25)) | ~where_6(X2) <- {110} [avatar component clause 8366]
8370. 112 <=> ! [X3] : ~select_8(sK3(sK24,sK25),X3) [avatar definition]
8371. ~select_8(sK3(sK24,sK25),X3) <- {112} [avatar component clause 8370]
8442. eq(X22,sK24) | select_8(X22,sK3(sK24,sK25)) | ~where_6(X22) <- {88} [resolution 8346,1836]
8464. select_8(X22,sK3(sK24,sK25)) | ~where_6(X22) <- {62, 88} [subsumption resolution 8442,4301]
8483. ~where_6(sK3(sK24,sK25)) <- {62, 90} [resolution 8353,4301]
8625. where_10(X0,sK3(sK24,sK25)) | ~where_6(X0) <- {62, 88} [resolution 8464,360]
8659. neq(X5,sK3(sK24,sK25)) | ~where_6(X5) <- {62, 88} [resolution 8625,357]
8670. ~where_6(X0) | ~eq(X0,sK3(sK24,sK25)) <- {62, 88} [resolution 8659,389]
8672. 110 | ~62 | ~88 [avatar split clause 8670,8178,4300,8366]
8684. ~where_6(X3) | ~select_1(sK24,sK25) | ~eq(sK24,X3) <- {110} [resolution 8367,647]
8699. ~eq(sK24,X3) | ~where_6(X3) <- {0, 110} [subsumption resolution 8684,3857]
9307. ~eq(X9,sK24) | ~select_8(sK3(sK24,sK25),X10) | select_12(sK3(sK24,sK25),X9) <- {0} [resolution 8090,1786]
9309. 114 <=> ! [X9] : (~eq(X9,sK24) | select_12(sK3(sK24,sK25),X9)) [avatar definition]
9310. select_12(sK3(sK24,sK25),X9) | ~eq(X9,sK24) <- {114} [avatar component clause 9309]
9311. 112 | 114 | ~0 [avatar split clause 9307,456,9309,8370]
9366. ~select_1(sK24,sK25) | eq(sK24,X0) <- {112} [resolution 8371,7899]
9375. eq(sK24,X0) <- {0, 112} [subsumption resolution 9366,3857]
9431. select_1(X1,sK25) <- {0, 112} [resolution 9375,3863]
9432. where_3(X2,sK24) <- {0, 112} [resolution 9375,3885]
9450. ~where_6(X36) <- {0, 110, 112} [resolution 9375,8699]
9453. eq(X39,sK24) <- {0, 112} [resolution 9375,382]
9454. eq(X40,X41) | ~eq(X40,sK24) <- {0, 112} [resolution 9375,383]
9458. ~where_3(X46,sK24) | where_3(X46,X47) <- {0, 112} [resolution 9375,432]
9496. 32 | ~0 | ~112 [avatar split clause 9431,8370,456,3004]
9497. 40 | ~0 | ~110 | ~112 [avatar split clause 9450,8370,8366,456,3955]
9498. eq(X40,X41) <- {0, 112} [subsumption resolution 9454,9453]
9499. 38 | ~0 | ~112 [avatar split clause 9498,8370,456,3952]
9502. where_3(X46,X47) <- {0, 112} [subsumption resolution 9458,9432]
9503. 42 | ~0 | ~112 [avatar split clause 9502,8370,456,3960]
10154. where_6_subquery_11(sK3(sK24,sK25),X1) | ~eq(X1,sK24) <- {114} [resolution 9310,372]
10530. ~eq(sK3(sK24,sK25),sK24) | where_6(sK3(sK24,sK25)) | where_6_subquery_7(sK3(sK24,sK25),sK3(sK24,sK25)) <- {114} [resolution 10154,375]
10531. where_6(sK3(sK24,sK25)) | where_6_subquery_7(sK3(sK24,sK25),sK3(sK24,sK25)) <- {90, 114} [subsumption resolution 10530,8353]
10532. where_6_subquery_7(sK3(sK24,sK25),sK3(sK24,sK25)) <- {90, 97, 114} [subsumption resolution 10531,8233]
10543. select_8(sK3(sK24,sK25),sK3(sK24,sK25)) <- {90, 97, 114} [resolution 10532,362]
10551. where_10(sK3(sK24,sK25),sK3(sK24,sK25)) <- {90, 97, 114} [resolution 10543,360]
10571. where_10(sK3(sK24,sK25),X0) | ~eq(sK3(sK24,sK25),X0) <- {90, 97, 114} [resolution 10551,402]
10574. where_10(sK3(sK24,sK25),X0) <- {90, 97, 114} [subsumption resolution 10571,483]
10587. neq(sK3(sK24,sK25),X6) <- {90, 97, 114} [resolution 10574,357]
10609. ~eq(sK3(sK24,sK25),X0) <- {90, 97, 114} [resolution 10587,389]
10645. $false <- {90, 97, 114} [resolution 10609,8353]
10664. ~90 | 97 | ~114 [avatar contradiction clause 10645]
10719. ~select_1(sK25,X85) | ~eq(sK24,X86) | ~eq(sK3(sK25,X85),X87) | select_1(X86,X87) <- {0} [resolution 4035,5175]
10740. ~select_1(sK25,X85) | ~eq(sK3(sK25,X85),X87) | select_1(X86,X87) <- {0, 38} [subsumption resolution 10719,3953]
10741. ~select_1(sK25,X85) | select_1(X86,X87) <- {0, 38} [subsumption resolution 10740,3953]
10743. 132 <=> ! [X86,X87] : select_1(X86,X87) [avatar definition]
10744. select_1(X86,X87) <- {132} [avatar component clause 10743]
10745. 132 | 4 | ~0 | ~38 [avatar split clause 10741,3952,456,691,10743]
11329. ~97 | ~62 | ~90 [avatar split clause 8483,8186,4300,8232]
11559. ~where_3(X30,X31) | where_14(X30,X31) <- {132} [resolution 10744,817]
11597. where_14(X30,X31) <- {42, 132} [subsumption resolution 11559,3961]
11685. ~select_1(X16,X17) | select_12(sK3(X16,X17),X18) <- {42, 132} [resolution 11597,1956]
11687. select_12(sK3(X16,X17),X18) <- {42, 132} [subsumption resolution 11685,10744]
12372. select_12(X6,X7) | ~eq(sK3(X8,X9),X6) <- {42, 132} [resolution 11687,405]
12378. select_12(X6,X7) <- {38, 42, 132} [subsumption resolution 12372,3953]
12404. where_6_subquery_11(X2,X3) <- {38, 42, 132} [resolution 12378,372]
12527. where_6(X8) | where_6_subquery_7(X8,X8) <- {38, 42, 132} [resolution 12404,375]
12528. where_6_subquery_7(X8,X8) <- {38, 40, 42, 132} [subsumption resolution 12527,3956]
12625. $false <- {38, 40, 42, 65, 132} [resolution 12528,4608]
12626. ~38 | ~40 | ~42 | 65 | ~132 [avatar contradiction clause 12625]
12704. 188 <=> select_1(sK24,sK25) [avatar definition]
12705. select_1(sK24,sK25) <- {188} [avatar component clause 12704]
12707. 189 <=> ~select_1(sK24,sK25) [avatar definition]
12708. ~select_1(sK24,sK25) <- {189} [avatar component clause 12707]
12721. 188 | ~86 [avatar split clause 6801,6707,12704]
12774. 198 <=> ! [X1] : (~where_6(X1) | ~where_3(sK24,X1)) [avatar definition]
12775. ~where_3(sK24,X1) | ~where_6(X1) <- {198} [avatar component clause 12774]
12798. select_4(sK24,sK25) <- {2} [resolution 463,379]
12807. select_1(sK24,sK25) | ~select_12(X8,X9) | ~where_3(sK24,X8) <- {2} [resolution 12798,1770]
12818. ~select_12(X8,X9) | ~where_3(sK24,X8) <- {2, 189} [subsumption resolution 12807,12708]
12999. ~select_1(sK24,sK25) <- {89} [resolution 8176,490]
13040. ~select_1(sK24,sK25) <- {91} [resolution 8184,504]
13190. ~where_3(sK24,X0) | ~where_6(X0) <- {2, 189} [resolution 12818,473]
13200. 198 | ~2 | 189 [avatar split clause 13190,12707,462,12774]
13231. ~where_6(sK24) <- {198} [resolution 12775,469]
13239. $false <- {20, 198} [subsumption resolution 13231,923]
13240. ~20 | ~198 [avatar contradiction clause 13239]
13283. ~189 | 91 [avatar split clause 13040,8183,12707]
13298. ~189 | 89 [avatar split clause 12999,8175,12707]
13377. main_query_1(sK24,sK25) <- {188} [resolution 12705,351]
13424. $false <- {1, 188} [subsumption resolution 13377,460]
13425. 1 | ~188 [avatar contradiction clause 13424]
13432. $false [avatar sat refutation 467,468,788,1068,3479,3859,3861,3901,4013,4618,4637,4723,4783,4784,6709,8672,9311,9496,9497,9499,9503,10664,10745,11329,12626,12721,13200,13240,13283,13298,13425]
% SZS output end Proof for query_21_query_30_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 11897
% Time elapsed: 0.766 s
% ------------------------------
% ------------------------------
