% Refutation found. Thanks to Tanya!
% SZS status Theorem for query02_query08_tptp4X
% SZS output start Proof for query02_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_4(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (where_5(X6) <=> (gte(X6,X6) & lte(X6,X6))) [input]
6. ! [X6] : (select_3(X6) <=> ? [X7,X8,X9,X10] : (where_5(X6) & from_4(X6,X7,X8,X9,X10))) [input]
7. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
11. ! [X11] : lte(X11,X11) [input]
17. ! [X11,X12] : (gte(X11,X12) <=> lte(X12,X11)) [input]
46. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
47. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 46]
48. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
50. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
51. ! [X0] : (where_5(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 5]
52. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
53. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
57. ! [X0] : lte(X0,X0) [rectify 11]
63. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 17]
92. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 47]
154. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 92]
155. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
156. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 48]
157. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 156]
158. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
159. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 157,158]
160. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 49]
161. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 50]
162. ! [X0] : ((where_5(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [nnf transformation 51]
163. ! [X0] : ((where_5(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [flattening 162]
164. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 52]
165. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 164]
166. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
167. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 165,166]
168. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 53]
171. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 63]
172. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 154]
173. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
174. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 172,173]
175. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 155]
176. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 155]
177. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 159]
178. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 159]
179. ~main_query_1(X0) | select_1(X0) [cnf transformation 160]
180. ~select_1(X0) | main_query_1(X0) [cnf transformation 160]
181. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 161]
182. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 161]
185. ~gte(X0,X0) | where_5(X0) | ~lte(X0,X0) [cnf transformation 163]
186. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 167]
188. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 167]
189. ~main_query_2(X0) | select_3(X0) [cnf transformation 168]
190. ~select_3(X0) | main_query_2(X0) [cnf transformation 168]
194. lte(X0,X0) [cnf transformation 57]
200. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 171]
229. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 174]
230. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 174]
260. 0 <=> main_query_1(sK8) [avatar definition]
261. main_query_1(sK8) <- {0} [avatar component clause 260]
263. 1 <=> ~main_query_1(sK8) [avatar definition]
264. ~main_query_1(sK8) <- {1} [avatar component clause 263]
266. 2 <=> main_query_2(sK8) [avatar definition]
267. main_query_2(sK8) <- {2} [avatar component clause 266]
269. 3 <=> ~main_query_2(sK8) [avatar definition]
270. ~main_query_2(sK8) <- {3} [avatar component clause 269]
271. ~1 | ~3 [avatar split clause 230,269,263]
272. 0 | 2 [avatar split clause 229,266,260]
273. select_3(sK8) <- {2} [resolution 189,267]
281. gte(X1,X1) [resolution 200,194]
283. where_5(X0) | ~lte(X0,X0) [resolution 185,281]
285. where_5(X0) [subsumption resolution 283,194]
291. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 177,175]
293. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 186,181]
295. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 291,182]
298. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 293,176]
300. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 295,188]
301. ~select_1(X1) | select_3(X1) [subsumption resolution 300,285]
303. ~select_3(X1) | select_1(X1) [resolution 298,178]
304. select_1(sK8) <- {2} [resolution 303,273]
306. main_query_1(sK8) <- {2} [resolution 304,180]
307. $false <- {1, 2} [subsumption resolution 306,264]
308. 1 | ~2 [avatar contradiction clause 307]
309. select_1(sK8) <- {0} [resolution 261,179]
310. select_3(sK8) <- {0} [resolution 309,301]
313. main_query_2(sK8) <- {0} [resolution 310,190]
315. $false <- {0, 3} [subsumption resolution 313,270]
316. ~0 | 3 [avatar contradiction clause 315]
317. $false [avatar sat refutation 271,272,308,316]
% SZS output end Proof for query02_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
