% Refutation found. Thanks to Tanya!
% SZS status Theorem for query05_query06_tptp4X
% SZS output start Proof for query05_query06_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X6,X7,X8,X9,X10] : (from_7(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
6. ! [X6] : (select_6(X6) <=> ? [X7,X8,X9,X10] : ($true & from_7(X6,X7,X8,X9,X10))) [input]
7. ! [X5] : (query_5(X5) <=> select_6(X5)) [input]
8. ! [X11] : (from_8(X11) <=> query_5(X11)) [input]
9. ! [X11] : (select_4(X11) <=> ($true & from_8(X11))) [input]
10. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
11. ! [X12] : eq(X12,X12) [input]
52. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
53. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 52]
54. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
55. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
56. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [rectify 6]
57. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 56]
58. ! [X0] : (query_5(X0) <=> select_6(X0)) [rectify 7]
59. ! [X0] : (from_8(X0) <=> query_5(X0)) [rectify 8]
60. ! [X0] : (select_4(X0) <=> ($true & from_8(X0))) [rectify 9]
61. ! [X0] : (select_4(X0) <=> from_8(X0)) [true and false elimination 60]
62. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 10]
63. ! [X0] : eq(X0,X0) [rectify 11]
104. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 53]
172. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 104]
173. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
174. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
175. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
176. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 175]
177. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
178. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 176,177]
179. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 54]
180. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 55]
181. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 57]
182. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 181]
183. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
184. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 182,183]
185. ! [X0] : ((query_5(X0) | ~select_6(X0)) & (select_6(X0) | ~query_5(X0))) [nnf transformation 58]
186. ! [X0] : ((from_8(X0) | ~query_5(X0)) & (query_5(X0) | ~from_8(X0))) [nnf transformation 59]
187. ! [X0] : ((select_4(X0) | ~from_8(X0)) & (from_8(X0) | ~select_4(X0))) [nnf transformation 61]
188. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 62]
191. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 172]
192. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
193. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 191,192]
194. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 173]
195. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 173]
197. ~eq(X0,X0) | where_3(X0) [cnf transformation 174]
198. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 178]
200. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 178]
201. ~main_query_1(X0) | select_1(X0) [cnf transformation 179]
202. ~select_1(X0) | main_query_1(X0) [cnf transformation 179]
203. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 180]
204. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 180]
205. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 184]
206. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 184]
207. ~query_5(X0) | select_6(X0) [cnf transformation 185]
208. ~select_6(X0) | query_5(X0) [cnf transformation 185]
209. ~from_8(X0) | query_5(X0) [cnf transformation 186]
210. ~query_5(X0) | from_8(X0) [cnf transformation 186]
211. ~select_4(X0) | from_8(X0) [cnf transformation 187]
212. ~from_8(X0) | select_4(X0) [cnf transformation 187]
213. ~main_query_2(X0) | select_4(X0) [cnf transformation 188]
214. ~select_4(X0) | main_query_2(X0) [cnf transformation 188]
215. eq(X0,X0) [cnf transformation 63]
254. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 193]
255. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 193]
288. 0 <=> main_query_1(sK8) [avatar definition]
289. main_query_1(sK8) <- {0} [avatar component clause 288]
291. 1 <=> ~main_query_1(sK8) [avatar definition]
292. ~main_query_1(sK8) <- {1} [avatar component clause 291]
294. 2 <=> main_query_2(sK8) [avatar definition]
295. main_query_2(sK8) <- {2} [avatar component clause 294]
297. 3 <=> ~main_query_2(sK8) [avatar definition]
298. ~main_query_2(sK8) <- {3} [avatar component clause 297]
299. ~1 | ~3 [avatar split clause 255,297,291]
300. 0 | 2 [avatar split clause 254,294,288]
301. select_4(sK8) <- {2} [resolution 213,295]
302. from_8(sK8) <- {2} [resolution 301,211]
305. query_5(sK8) <- {2} [resolution 302,209]
307. select_6(sK8) <- {2} [resolution 305,207]
309. where_3(X1) [resolution 197,215]
321. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 198,194]
323. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 205,203]
325. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 321,204]
328. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 323,195]
330. ~select_1(X1) | select_6(X1) [resolution 325,206]
332. ~select_6(X1) | ~where_3(X1) | select_1(X1) [resolution 328,200]
333. ~select_6(X1) | select_1(X1) [subsumption resolution 332,309]
334. select_1(sK8) <- {2} [resolution 333,307]
336. main_query_1(sK8) <- {2} [resolution 334,202]
338. $false <- {1, 2} [subsumption resolution 336,292]
339. 1 | ~2 [avatar contradiction clause 338]
340. select_1(sK8) <- {0} [resolution 289,201]
341. select_6(sK8) <- {0} [resolution 340,330]
345. query_5(sK8) <- {0} [resolution 341,208]
346. from_8(sK8) <- {0} [resolution 345,210]
348. select_4(sK8) <- {0} [resolution 346,212]
350. main_query_2(sK8) <- {0} [resolution 348,214]
352. $false <- {0, 3} [subsumption resolution 350,298]
353. ~0 | 3 [avatar contradiction clause 352]
354. $false [avatar sat refutation 299,300,339,353]
% SZS output end Proof for query05_query06_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.011 s
% ------------------------------
% ------------------------------
