% Refutation found. Thanks to Tanya!
% SZS status Theorem for query30_query33_tptp4X
% SZS output start Proof for query30_query33_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> lt(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1] : (where_6(X0,X1) <=> lt(X0,X1)) [input]
7. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
52. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
53. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 52]
54. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
55. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
99. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 53]
169. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 99]
170. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
171. ! [X0,X1] : ((where_3(X0,X1) | ~lt(X0,X1)) & (lt(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
172. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
173. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 172]
174. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
175. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 173,174]
176. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 54]
177. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
178. ! [X0,X1] : ((where_6(X0,X1) | ~lt(X0,X1)) & (lt(X0,X1) | ~where_6(X0,X1))) [nnf transformation 6]
179. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 7]
180. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 179]
181. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
182. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 180,181]
183. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 55]
188. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 169]
189. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
190. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 188,189]
191. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 170]
192. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 170]
193. ~where_3(X0,X1) | lt(X0,X1) [cnf transformation 171]
194. ~lt(X0,X1) | where_3(X0,X1) [cnf transformation 171]
195. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 175]
196. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 175]
197. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 175]
198. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 176]
199. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 176]
200. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 177]
201. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 177]
202. ~where_6(X0,X1) | lt(X0,X1) [cnf transformation 178]
203. ~lt(X0,X1) | where_6(X0,X1) [cnf transformation 178]
204. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 182]
205. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 182]
206. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 182]
207. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 183]
208. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 183]
253. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 190]
254. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 190]
256. 0 <=> main_query_1(sK6,sK7) [avatar definition]
257. main_query_1(sK6,sK7) <- {0} [avatar component clause 256]
259. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
260. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 259]
262. 2 <=> main_query_2(sK6,sK7) [avatar definition]
263. main_query_2(sK6,sK7) <- {2} [avatar component clause 262]
265. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
266. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 265]
267. ~1 | ~3 [avatar split clause 254,265,259]
268. 0 | 2 [avatar split clause 253,262,256]
269. select_4(sK6,sK7) <- {2} [resolution 207,263]
271. where_6(sK6,sK7) <- {2} [resolution 269,205]
273. lt(sK6,sK7) <- {2} [resolution 271,202]
276. where_3(sK6,sK7) <- {2} [resolution 273,194]
428. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 195,191]
446. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 204,200]
606. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 428,201]
614. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 446,192]
870. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 606,206]
892. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 614,197]
898. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 892,269]
899. select_1(sK6,sK7) <- {2} [subsumption resolution 898,276]
907. main_query_1(sK6,sK7) <- {2} [resolution 899,199]
909. $false <- {1, 2} [subsumption resolution 907,260]
910. 1 | ~2 [avatar contradiction clause 909]
911. select_1(sK6,sK7) <- {0} [resolution 257,198]
913. where_3(sK6,sK7) <- {0} [resolution 911,196]
959. lt(sK6,sK7) <- {0} [resolution 913,193]
961. where_6(sK6,sK7) <- {0} [resolution 959,203]
984. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 961,870]
986. select_4(sK6,sK7) <- {0} [subsumption resolution 984,911]
1032. main_query_2(sK6,sK7) <- {0} [resolution 986,208]
1034. $false <- {0, 3} [subsumption resolution 1032,266]
1035. ~0 | 3 [avatar contradiction clause 1034]
1036. $false [avatar sat refutation 267,268,910,1035]
% SZS output end Proof for query30_query33_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.017 s
% ------------------------------
% ------------------------------
