% Refutation found. Thanks to Tanya!
% SZS status Theorem for query04_query07_tptp4X
% SZS output start Proof for query04_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
7. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [input]
9. ! [X5] : eq(X5,X5) [input]
15. ! [X5,X6] : (lt(X5,X6) <=> (~eq(X5,X6) & lte(X5,X6))) [input]
16. ! [X5,X6] : (gt(X5,X6) <=> lt(X6,X5)) [input]
42. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
43. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 42]
44. ! [X0] : eq(X0,X0) [rectify 9]
50. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 15]
51. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 16]
128. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 43]
129. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
130. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
131. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
132. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 131]
133. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 132,133]
135. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
136. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
137. ! [X0] : ((where_6(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [nnf transformation 6]
138. ! [X0] : ((where_6(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [flattening 137]
139. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 7]
140. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 139]
141. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
142. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 140,141]
143. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 8]
146. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 50]
147. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 146]
148. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 51]
149. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 128]
150. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
151. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 149,150]
152. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 129]
153. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 129]
155. ~eq(X0,X0) | where_3(X0) [cnf transformation 130]
156. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 134]
158. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 134]
159. ~main_query_1(X0) | select_1(X0) [cnf transformation 135]
160. ~select_1(X0) | main_query_1(X0) [cnf transformation 135]
161. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 136]
162. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 136]
165. gt(X0,X0) | where_6(X0) | lt(X0,X0) [cnf transformation 138]
166. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 142]
168. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 142]
169. ~main_query_2(X0) | select_4(X0) [cnf transformation 143]
170. ~select_4(X0) | main_query_2(X0) [cnf transformation 143]
171. eq(X0,X0) [cnf transformation 44]
180. ~lt(X0,X1) | ~eq(X0,X1) [cnf transformation 147]
182. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 148]
207. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 151]
208. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 151]
233. 0 <=> main_query_1(sK8) [avatar definition]
234. main_query_1(sK8) <- {0} [avatar component clause 233]
236. 1 <=> ~main_query_1(sK8) [avatar definition]
237. ~main_query_1(sK8) <- {1} [avatar component clause 236]
239. 2 <=> main_query_2(sK8) [avatar definition]
240. main_query_2(sK8) <- {2} [avatar component clause 239]
242. 3 <=> ~main_query_2(sK8) [avatar definition]
243. ~main_query_2(sK8) <- {3} [avatar component clause 242]
244. ~1 | ~3 [avatar split clause 208,242,236]
245. 0 | 2 [avatar split clause 207,239,233]
246. select_4(sK8) <- {2} [resolution 169,240]
250. where_3(X1) [resolution 155,171]
258. where_6(X1) | lt(X1,X1) | lt(X1,X1) [resolution 165,182]
259. lt(X1,X1) | where_6(X1) [duplicate literal removal 258]
264. where_6(X2) | ~eq(X2,X2) [resolution 259,180]
266. where_6(X2) [subsumption resolution 264,171]
272. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 156,152]
274. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 166,161]
276. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 272,162]
279. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 274,153]
281. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 276,168]
282. ~select_1(X1) | select_4(X1) [subsumption resolution 281,266]
284. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 279,158]
285. ~select_4(X1) | select_1(X1) [subsumption resolution 284,250]
286. select_1(sK8) <- {2} [resolution 285,246]
288. main_query_1(sK8) <- {2} [resolution 286,160]
290. $false <- {1, 2} [subsumption resolution 288,237]
291. 1 | ~2 [avatar contradiction clause 290]
292. select_1(sK8) <- {0} [resolution 234,159]
293. select_4(sK8) <- {0} [resolution 292,282]
297. main_query_2(sK8) <- {0} [resolution 293,170]
299. $false <- {0, 3} [subsumption resolution 297,243]
300. ~0 | 3 [avatar contradiction clause 299]
301. $false [avatar sat refutation 244,245,291,300]
% SZS output end Proof for query04_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
