% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_01_query_07_tptp4X
% SZS output start Proof for query_01_query_07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
6. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X5,X6] : (main_query_2(X5,X6) <=> select_3(X5,X6)) [input]
8. ! [X7] : eq(X7,X7) [input]
14. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
15. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
46. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
47. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 46]
48. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 3]
50. ! [X0,X1] : (main_query_2(X0,X1) <=> select_3(X0,X1)) [rectify 7]
51. ! [X0] : eq(X0,X0) [rectify 8]
57. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 14]
58. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 15]
89. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 47]
151. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 89]
152. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
153. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0,X1))) [nnf transformation 48]
154. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) | ~select_1(X0,X1))) [rectify 153]
155. ! [X1,X0] : (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) => from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) [choice axiom]
156. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1))) [skolemisation 154,155]
157. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 49]
158. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
159. ! [X0] : ((where_5(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [nnf transformation 5]
160. ! [X0] : ((where_5(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [flattening 159]
161. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0,X1))) [nnf transformation 6]
162. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) | ~select_3(X0,X1))) [rectify 161]
163. ! [X1,X0] : (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) => (where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
164. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_3(X0,X1))) [skolemisation 162,163]
165. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 50]
168. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 57]
169. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 168]
170. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 58]
171. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 151]
172. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
173. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 171,172]
174. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 152]
175. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 152]
176. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 156]
177. ~from_2(X0,X1,X2,X3,X4) | select_1(X0,X1) [cnf transformation 156]
178. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 157]
179. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 157]
180. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 158]
181. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 158]
184. gt(X0,X0) | where_5(X0) | lt(X0,X0) [cnf transformation 160]
185. from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [cnf transformation 164]
187. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0,X1) [cnf transformation 164]
188. ~main_query_2(X0,X1) | select_3(X0,X1) [cnf transformation 165]
189. ~select_3(X0,X1) | main_query_2(X0,X1) [cnf transformation 165]
190. eq(X0,X0) [cnf transformation 51]
199. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 169]
201. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 170]
231. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 173]
232. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 173]
234. 0 <=> main_query_1(sK6,sK7) [avatar definition]
235. main_query_1(sK6,sK7) <- {0} [avatar component clause 234]
237. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
238. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 237]
240. 2 <=> main_query_2(sK6,sK7) [avatar definition]
241. main_query_2(sK6,sK7) <- {2} [avatar component clause 240]
243. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
244. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 243]
245. ~1 | ~3 [avatar split clause 232,243,237]
246. 0 | 2 [avatar split clause 231,240,234]
247. select_3(sK6,sK7) <- {2} [resolution 188,241]
253. ~lt(X0,X0) [resolution 199,190]
255. where_5(X1) | lt(X1,X1) | lt(X1,X1) [resolution 184,201]
256. where_5(X1) | lt(X1,X1) [duplicate literal removal 255]
257. where_5(X1) [subsumption resolution 256,253]
270. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 176,174]
272. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 185,180]
285. from_4(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 270,181]
291. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 272,175]
306. ~select_1(X2,X3) | ~where_5(X2) | select_3(X2,X3) [resolution 285,187]
312. ~select_1(X2,X3) | select_3(X2,X3) [subsumption resolution 306,257]
328. ~select_3(X2,X3) | select_1(X2,X3) [resolution 291,177]
341. select_1(sK6,sK7) <- {2} [resolution 328,247]
350. main_query_1(sK6,sK7) <- {2} [resolution 341,179]
351. $false <- {1, 2} [subsumption resolution 350,238]
352. 1 | ~2 [avatar contradiction clause 351]
353. select_1(sK6,sK7) <- {0} [resolution 235,178]
354. select_3(sK6,sK7) <- {0} [resolution 353,312]
357. main_query_2(sK6,sK7) <- {0} [resolution 354,189]
359. $false <- {0, 3} [subsumption resolution 357,244]
360. ~0 | 3 [avatar contradiction clause 359]
361. $false [avatar sat refutation 245,246,352,360]
% SZS output end Proof for query_01_query_07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
