% Refutation found. Thanks to Tanya!
% SZS status Theorem for query20_query22_tptp4X
% SZS output start Proof for query20_query22_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
3. ! [X5,X6] : (where_9(X5,X6) <=> eq(X6,X5)) [input]
4. ! [X5,X6] : (select_7(X6,X5) <=> ? [X0,X1,X2,X3,X4] : (where_9(X5,X6) & from_8(X0,X1,X2,X3,X4))) [input]
5. ! [X6,X5] : (query_6(X6,X5) <=> select_7(X6,X5)) [input]
6. ! [X7,X8] : (from_10(X7,X8) <=> query_6(X7,X8)) [input]
7. ! [X7] : (select_5(X7) <=> ? [X8] : ($true & from_10(X7,X8))) [input]
8. ! [X6] : (query_4(X6) <=> select_5(X6)) [input]
9. ! [X6] : (where_3(X6) <=> query_4(X6)) [input]
10. ! [X6] : (select_1(X6) <=> ? [X0,X1,X2,X3,X4] : (where_3(X6) & from_2(X0,X1,X2,X3,X4))) [input]
11. ! [X6] : (main_query_1(X6) <=> select_1(X6)) [input]
12. ! [X0,X1,X2,X3,X4] : (from_14(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
13. ! [X5,X6] : (where_15(X5,X6) <=> eq(X6,X5)) [input]
14. ! [X5,X6] : (select_13(X6,X5) <=> ? [X0,X1,X2,X3,X4] : (where_15(X5,X6) & from_14(X0,X1,X2,X3,X4))) [input]
15. ! [X9,X5] : (query_12(X9,X5) <=> select_13(X9,X5)) [input]
16. ! [X10,X8] : (from_16(X10,X8) <=> query_12(X10,X8)) [input]
17. ! [X10] : (select_11(X10) <=> ? [X8] : ($true & from_16(X10,X8))) [input]
18. ! [X9] : (main_query_2(X9) <=> select_11(X9)) [input]
19. ! [X11] : eq(X11,X11) [input]
74. ! [X6] : (main_query_1(X6) <=> main_query_2(X6)) [input]
75. ~! [X6] : (main_query_1(X6) <=> main_query_2(X6)) [negated conjecture 74]
76. ! [X0,X1] : (where_9(X0,X1) <=> eq(X1,X0)) [rectify 3]
77. ! [X0,X1] : (select_7(X1,X0) <=> ? [X2,X3,X4,X5,X6] : (where_9(X0,X1) & from_8(X2,X3,X4,X5,X6))) [rectify 4]
78. ! [X0,X1] : (query_6(X0,X1) <=> select_7(X0,X1)) [rectify 5]
79. ! [X0,X1] : (from_10(X0,X1) <=> query_6(X0,X1)) [rectify 6]
80. ! [X0] : (select_5(X0) <=> ? [X1] : ($true & from_10(X0,X1))) [rectify 7]
81. ! [X0] : (select_5(X0) <=> ? [X1] : from_10(X0,X1)) [true and false elimination 80]
82. ! [X0] : (query_4(X0) <=> select_5(X0)) [rectify 8]
83. ! [X0] : (where_3(X0) <=> query_4(X0)) [rectify 9]
84. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : (where_3(X0) & from_2(X1,X2,X3,X4,X5))) [rectify 10]
85. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 11]
86. ! [X0,X1] : (where_15(X0,X1) <=> eq(X1,X0)) [rectify 13]
87. ! [X0,X1] : (select_13(X1,X0) <=> ? [X2,X3,X4,X5,X6] : (where_15(X0,X1) & from_14(X2,X3,X4,X5,X6))) [rectify 14]
88. ! [X0,X1] : (query_12(X0,X1) <=> select_13(X0,X1)) [rectify 15]
89. ! [X0,X1] : (from_16(X0,X1) <=> query_12(X0,X1)) [rectify 16]
90. ! [X0] : (select_11(X0) <=> ? [X1] : ($true & from_16(X0,X1))) [rectify 17]
91. ! [X0] : (select_11(X0) <=> ? [X1] : from_16(X0,X1)) [true and false elimination 90]
92. ! [X0] : (main_query_2(X0) <=> select_11(X0)) [rectify 18]
93. ! [X0] : eq(X0,X0) [rectify 19]
148. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 75]
244. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 148]
245. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
246. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 2]
247. ! [X0,X1] : ((where_9(X0,X1) | ~eq(X1,X0)) & (eq(X1,X0) | ~where_9(X0,X1))) [nnf transformation 76]
248. ! [X0,X1] : ((select_7(X1,X0) | ! [X2,X3,X4,X5,X6] : (~where_9(X0,X1) | ~from_8(X2,X3,X4,X5,X6))) & (? [X2,X3,X4,X5,X6] : (where_9(X0,X1) & from_8(X2,X3,X4,X5,X6)) | ~select_7(X1,X0))) [nnf transformation 77]
249. ! [X0,X1] : ((select_7(X1,X0) | ! [X2,X3,X4,X5,X6] : (~where_9(X0,X1) | ~from_8(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_9(X0,X1) & from_8(X7,X8,X9,X10,X11)) | ~select_7(X1,X0))) [rectify 248]
250. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_9(X0,X1) & from_8(X7,X8,X9,X10,X11)) => (where_9(X0,X1) & from_8(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)))) [choice axiom]
251. ! [X0,X1] : ((select_7(X1,X0) | ! [X2,X3,X4,X5,X6] : (~where_9(X0,X1) | ~from_8(X2,X3,X4,X5,X6))) & ((where_9(X0,X1) & from_8(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1))) | ~select_7(X1,X0))) [skolemisation 249,250]
252. ! [X0,X1] : ((query_6(X0,X1) | ~select_7(X0,X1)) & (select_7(X0,X1) | ~query_6(X0,X1))) [nnf transformation 78]
253. ! [X0,X1] : ((from_10(X0,X1) | ~query_6(X0,X1)) & (query_6(X0,X1) | ~from_10(X0,X1))) [nnf transformation 79]
254. ! [X0] : ((select_5(X0) | ! [X1] : ~from_10(X0,X1)) & (? [X1] : from_10(X0,X1) | ~select_5(X0))) [nnf transformation 81]
255. ! [X0] : ((select_5(X0) | ! [X1] : ~from_10(X0,X1)) & (? [X2] : from_10(X0,X2) | ~select_5(X0))) [rectify 254]
256. ! [X0] : (? [X2] : from_10(X0,X2) => from_10(X0,sK5(X0))) [choice axiom]
257. ! [X0] : ((select_5(X0) | ! [X1] : ~from_10(X0,X1)) & (from_10(X0,sK5(X0)) | ~select_5(X0))) [skolemisation 255,256]
258. ! [X0] : ((query_4(X0) | ~select_5(X0)) & (select_5(X0) | ~query_4(X0))) [nnf transformation 82]
259. ! [X0] : ((where_3(X0) | ~query_4(X0)) & (query_4(X0) | ~where_3(X0))) [nnf transformation 83]
260. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : (~where_3(X0) | ~from_2(X1,X2,X3,X4,X5))) & (? [X1,X2,X3,X4,X5] : (where_3(X0) & from_2(X1,X2,X3,X4,X5)) | ~select_1(X0))) [nnf transformation 84]
261. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : (~where_3(X0) | ~from_2(X1,X2,X3,X4,X5))) & (? [X6,X7,X8,X9,X10] : (where_3(X0) & from_2(X6,X7,X8,X9,X10)) | ~select_1(X0))) [rectify 260]
262. ! [X0] : (? [X6,X7,X8,X9,X10] : (where_3(X0) & from_2(X6,X7,X8,X9,X10)) => (where_3(X0) & from_2(sK6(X0),sK7(X0),sK8(X0),sK9(X0),sK10(X0)))) [choice axiom]
263. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : (~where_3(X0) | ~from_2(X1,X2,X3,X4,X5))) & ((where_3(X0) & from_2(sK6(X0),sK7(X0),sK8(X0),sK9(X0),sK10(X0))) | ~select_1(X0))) [skolemisation 261,262]
264. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 85]
265. ! [X0,X1,X2,X3,X4] : ((from_14(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_14(X0,X1,X2,X3,X4))) [nnf transformation 12]
266. ! [X0,X1] : ((where_15(X0,X1) | ~eq(X1,X0)) & (eq(X1,X0) | ~where_15(X0,X1))) [nnf transformation 86]
267. ! [X0,X1] : ((select_13(X1,X0) | ! [X2,X3,X4,X5,X6] : (~where_15(X0,X1) | ~from_14(X2,X3,X4,X5,X6))) & (? [X2,X3,X4,X5,X6] : (where_15(X0,X1) & from_14(X2,X3,X4,X5,X6)) | ~select_13(X1,X0))) [nnf transformation 87]
268. ! [X0,X1] : ((select_13(X1,X0) | ! [X2,X3,X4,X5,X6] : (~where_15(X0,X1) | ~from_14(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_15(X0,X1) & from_14(X7,X8,X9,X10,X11)) | ~select_13(X1,X0))) [rectify 267]
269. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_15(X0,X1) & from_14(X7,X8,X9,X10,X11)) => (where_15(X0,X1) & from_14(sK11(X0,X1),sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1)))) [choice axiom]
270. ! [X0,X1] : ((select_13(X1,X0) | ! [X2,X3,X4,X5,X6] : (~where_15(X0,X1) | ~from_14(X2,X3,X4,X5,X6))) & ((where_15(X0,X1) & from_14(sK11(X0,X1),sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1))) | ~select_13(X1,X0))) [skolemisation 268,269]
271. ! [X0,X1] : ((query_12(X0,X1) | ~select_13(X0,X1)) & (select_13(X0,X1) | ~query_12(X0,X1))) [nnf transformation 88]
272. ! [X0,X1] : ((from_16(X0,X1) | ~query_12(X0,X1)) & (query_12(X0,X1) | ~from_16(X0,X1))) [nnf transformation 89]
273. ! [X0] : ((select_11(X0) | ! [X1] : ~from_16(X0,X1)) & (? [X1] : from_16(X0,X1) | ~select_11(X0))) [nnf transformation 91]
274. ! [X0] : ((select_11(X0) | ! [X1] : ~from_16(X0,X1)) & (? [X2] : from_16(X0,X2) | ~select_11(X0))) [rectify 273]
275. ! [X0] : (? [X2] : from_16(X0,X2) => from_16(X0,sK16(X0))) [choice axiom]
276. ! [X0] : ((select_11(X0) | ! [X1] : ~from_16(X0,X1)) & (from_16(X0,sK16(X0)) | ~select_11(X0))) [skolemisation 274,275]
277. ! [X0] : ((main_query_2(X0) | ~select_11(X0)) & (select_11(X0) | ~main_query_2(X0))) [nnf transformation 92]
280. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 244]
281. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK17) | ~main_query_1(sK17)) & (main_query_2(sK17) | main_query_1(sK17))) [choice axiom]
282. (~main_query_2(sK17) | ~main_query_1(sK17)) & (main_query_2(sK17) | main_query_1(sK17)) [skolemisation 280,281]
283. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 245]
284. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 245]
285. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 246]
286. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 246]
288. ~eq(X1,X0) | where_9(X0,X1) [cnf transformation 247]
289. from_8(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_7(X1,X0) [cnf transformation 251]
291. select_7(X1,X0) | ~where_9(X0,X1) | ~from_8(X2,X3,X4,X5,X6) [cnf transformation 251]
293. ~select_7(X0,X1) | query_6(X0,X1) [cnf transformation 252]
295. ~query_6(X0,X1) | from_10(X0,X1) [cnf transformation 253]
297. ~from_10(X0,X1) | select_5(X0) [cnf transformation 257]
299. ~select_5(X0) | query_4(X0) [cnf transformation 258]
301. ~query_4(X0) | where_3(X0) [cnf transformation 259]
302. from_2(sK6(X0),sK7(X0),sK8(X0),sK9(X0),sK10(X0)) | ~select_1(X0) [cnf transformation 263]
304. select_1(X0) | ~where_3(X0) | ~from_2(X1,X2,X3,X4,X5) [cnf transformation 263]
305. ~main_query_1(X0) | select_1(X0) [cnf transformation 264]
306. ~select_1(X0) | main_query_1(X0) [cnf transformation 264]
307. ~from_14(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 265]
308. ~tab1(X0,X1,X2,X3,X4) | from_14(X0,X1,X2,X3,X4) [cnf transformation 265]
310. ~eq(X1,X0) | where_15(X0,X1) [cnf transformation 266]
311. from_14(sK11(X0,X1),sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1)) | ~select_13(X1,X0) [cnf transformation 270]
313. select_13(X1,X0) | ~where_15(X0,X1) | ~from_14(X2,X3,X4,X5,X6) [cnf transformation 270]
314. ~query_12(X0,X1) | select_13(X0,X1) [cnf transformation 271]
315. ~select_13(X0,X1) | query_12(X0,X1) [cnf transformation 271]
316. ~from_16(X0,X1) | query_12(X0,X1) [cnf transformation 272]
317. ~query_12(X0,X1) | from_16(X0,X1) [cnf transformation 272]
318. from_16(X0,sK16(X0)) | ~select_11(X0) [cnf transformation 276]
319. ~from_16(X0,X1) | select_11(X0) [cnf transformation 276]
320. ~main_query_2(X0) | select_11(X0) [cnf transformation 277]
321. ~select_11(X0) | main_query_2(X0) [cnf transformation 277]
322. eq(X0,X0) [cnf transformation 93]
375. main_query_2(sK17) | main_query_1(sK17) [cnf transformation 282]
376. ~main_query_2(sK17) | ~main_query_1(sK17) [cnf transformation 282]
423. 0 <=> ! [X3,X5,X2,X4,X6] : ~from_8(X2,X3,X4,X5,X6) [avatar definition]
424. ~from_8(X2,X3,X4,X5,X6) <- {0} [avatar component clause 423]
426. 2 <=> ! [X1,X0] : (select_7(X1,X0) | ~where_9(X0,X1)) [avatar definition]
427. ~where_9(X0,X1) | select_7(X1,X0) <- {2} [avatar component clause 426]
428. 0 | 2 [avatar split clause 291,426,423]
430. 4 <=> ! [X1,X3,X5,X2,X4] : ~from_2(X1,X2,X3,X4,X5) [avatar definition]
431. ~from_2(X1,X2,X3,X4,X5) <- {4} [avatar component clause 430]
433. 6 <=> ! [X0] : (select_1(X0) | ~where_3(X0)) [avatar definition]
434. ~where_3(X0) | select_1(X0) <- {6} [avatar component clause 433]
435. 4 | 6 [avatar split clause 304,433,430]
437. 8 <=> ! [X3,X5,X2,X4,X6] : ~from_14(X2,X3,X4,X5,X6) [avatar definition]
438. ~from_14(X2,X3,X4,X5,X6) <- {8} [avatar component clause 437]
440. 10 <=> ! [X1,X0] : (select_13(X1,X0) | ~where_15(X0,X1)) [avatar definition]
441. ~where_15(X0,X1) | select_13(X1,X0) <- {10} [avatar component clause 440]
442. 8 | 10 [avatar split clause 313,440,437]
444. 12 <=> main_query_1(sK17) [avatar definition]
445. main_query_1(sK17) <- {12} [avatar component clause 444]
447. 13 <=> ~main_query_1(sK17) [avatar definition]
448. ~main_query_1(sK17) <- {13} [avatar component clause 447]
450. 14 <=> main_query_2(sK17) [avatar definition]
451. main_query_2(sK17) <- {14} [avatar component clause 450]
453. 15 <=> ~main_query_2(sK17) [avatar definition]
454. ~main_query_2(sK17) <- {15} [avatar component clause 453]
455. ~13 | ~15 [avatar split clause 376,453,447]
456. 12 | 14 [avatar split clause 375,450,444]
464. where_9(X0,X0) [resolution 288,322]
468. where_15(X0,X0) [resolution 310,322]
470. query_12(X0,sK16(X0)) | ~select_11(X0) [resolution 318,316]
475. select_7(X0,X0) <- {2} [resolution 427,464]
476. query_6(X0,X0) <- {2} [resolution 475,293]
478. from_10(X0,X0) <- {2} [resolution 476,295]
482. select_5(X1) <- {2} [resolution 478,297]
484. query_4(X0) <- {2} [resolution 482,299]
486. where_3(X0) <- {2} [resolution 484,301]
500. tab1(sK6(X0),sK7(X0),sK8(X0),sK9(X0),sK10(X0)) | ~select_1(X0) [resolution 302,283]
504. select_11(sK17) <- {14} [resolution 451,320]
505. tab1(sK0(X1,X0),sK1(X1,X0),sK2(X1,X0),sK3(X1,X0),sK4(X1,X0)) | ~select_7(X0,X1) [resolution 289,285]
507. tab1(sK11(X1,X0),sK12(X1,X0),sK13(X1,X0),sK14(X1,X0),sK15(X1,X0)) | ~select_13(X0,X1) [resolution 311,307]
508. ~select_13(X2,X3) <- {8} [resolution 311,438]
510. select_13(X1,sK16(X1)) | ~select_11(X1) [resolution 470,314]
511. ~select_11(X1) <- {8} [subsumption resolution 510,508]
512. $false <- {8, 14} [resolution 511,504]
513. ~8 | ~14 [avatar contradiction clause 512]
514. select_13(X0,X0) <- {10} [resolution 441,468]
515. query_12(X0,X0) <- {10} [resolution 514,315]
517. from_14(sK6(X0),sK7(X0),sK8(X0),sK9(X0),sK10(X0)) | ~select_1(X0) [resolution 500,308]
520. from_16(X0,X0) <- {10} [resolution 515,317]
524. from_2(sK0(X5,X4),sK1(X5,X4),sK2(X5,X4),sK3(X5,X4),sK4(X5,X4)) | ~select_7(X4,X5) [resolution 505,284]
525. ~select_7(X4,X5) <- {4} [subsumption resolution 524,431]
526. $false <- {2, 4} [resolution 525,475]
529. ~2 | ~4 [avatar contradiction clause 526]
530. select_1(X0) <- {2, 6} [subsumption resolution 434,486]
531. main_query_1(X0) <- {2, 6} [resolution 530,306]
534. $false <- {2, 6, 13} [resolution 531,448]
535. ~2 | ~6 | 13 [avatar contradiction clause 534]
537. select_11(X1) <- {10} [resolution 520,319]
539. from_8(sK11(X3,X2),sK12(X3,X2),sK13(X3,X2),sK14(X3,X2),sK15(X3,X2)) | ~select_13(X2,X3) [resolution 507,286]
541. main_query_2(X0) <- {10} [resolution 537,321]
546. $false <- {10, 15} [resolution 541,454]
547. ~10 | 15 [avatar contradiction clause 546]
573. ~select_13(X0,X1) <- {0} [resolution 539,424]
577. ~select_11(X0) <- {0} [resolution 573,510]
580. $false <- {0, 10} [subsumption resolution 577,537]
581. ~0 | ~10 [avatar contradiction clause 580]
583. select_1(sK17) <- {12} [resolution 445,305]
592. ~select_1(X2) <- {8} [resolution 438,517]
597. $false <- {8, 12} [resolution 592,583]
598. ~8 | ~12 [avatar contradiction clause 597]
599. $false [avatar sat refutation 428,435,442,455,456,513,529,535,547,581,598]
% SZS output end Proof for query20_query22_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
