% Refutation found. Thanks to Tanya!
% SZS status Theorem for query03_query08_tptp4X
% SZS output start Proof for query03_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X5,X6,X7,X8,X9] : (from_4(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
5. ! [X5] : (where_5(X5) <=> (gte(X5,X5) & lte(X5,X5))) [input]
6. ! [X5] : (select_3(X5) <=> ? [X6,X7,X8,X9] : (where_5(X5) & from_4(X5,X6,X7,X8,X9))) [input]
7. ! [X10] : (main_query_2(X10) <=> select_3(X10)) [input]
11. ! [X11] : lte(X11,X11) [input]
18. ! [X11,X12] : (gte(X11,X12) <=> lte(X12,X11)) [input]
42. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
43. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 42]
44. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
45. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
46. ! [X0] : (where_5(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 5]
47. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
48. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
52. ! [X0] : lte(X0,X0) [rectify 11]
59. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 18]
134. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 43]
135. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 44]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 136]
138. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
139. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 137,138]
140. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
141. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 45]
142. ! [X0] : ((where_5(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [nnf transformation 46]
143. ! [X0] : ((where_5(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [flattening 142]
144. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 47]
145. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 144]
146. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
147. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 145,146]
148. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 48]
151. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 59]
152. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 134]
153. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
154. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 152,153]
155. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 135]
156. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 135]
157. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 139]
158. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 139]
159. ~main_query_1(X0) | select_1(X0) [cnf transformation 140]
160. ~select_1(X0) | main_query_1(X0) [cnf transformation 140]
161. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 141]
162. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 141]
165. ~gte(X0,X0) | where_5(X0) | ~lte(X0,X0) [cnf transformation 143]
166. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 147]
168. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 147]
169. ~main_query_2(X0) | select_3(X0) [cnf transformation 148]
170. ~select_3(X0) | main_query_2(X0) [cnf transformation 148]
174. lte(X0,X0) [cnf transformation 52]
181. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 151]
205. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 154]
206. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 154]
208. 0 <=> main_query_1(sK8) [avatar definition]
209. main_query_1(sK8) <- {0} [avatar component clause 208]
211. 1 <=> ~main_query_1(sK8) [avatar definition]
212. ~main_query_1(sK8) <- {1} [avatar component clause 211]
214. 2 <=> main_query_2(sK8) [avatar definition]
215. main_query_2(sK8) <- {2} [avatar component clause 214]
217. 3 <=> ~main_query_2(sK8) [avatar definition]
218. ~main_query_2(sK8) <- {3} [avatar component clause 217]
219. ~1 | ~3 [avatar split clause 206,217,211]
220. 0 | 2 [avatar split clause 205,214,208]
221. select_3(sK8) <- {2} [resolution 169,215]
230. gte(X1,X1) [resolution 181,174]
234. where_5(X0) | ~lte(X0,X0) [resolution 165,230]
236. where_5(X0) [subsumption resolution 234,174]
302. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 157,155]
304. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 166,161]
306. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 302,162]
309. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 304,156]
311. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 306,168]
312. ~select_1(X1) | select_3(X1) [subsumption resolution 311,236]
314. ~select_3(X1) | select_1(X1) [resolution 309,158]
315. select_1(sK8) <- {2} [resolution 314,221]
319. main_query_1(sK8) <- {2} [resolution 315,160]
320. $false <- {1, 2} [subsumption resolution 319,212]
321. 1 | ~2 [avatar contradiction clause 320]
322. select_1(sK8) <- {0} [resolution 209,159]
325. select_3(sK8) <- {0} [resolution 322,312]
330. main_query_2(sK8) <- {0} [resolution 325,170]
332. $false <- {0, 3} [subsumption resolution 330,218]
333. ~0 | 3 [avatar contradiction clause 332]
334. $false [avatar sat refutation 219,220,321,333]
% SZS output end Proof for query03_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.025 s
% ------------------------------
% ------------------------------
