% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_04_query_09_tptp4X
% SZS output start Proof for query_04_query_09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [input]
4. ! [X5,X6,X7,X8,X9] : (from_4(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
5. ! [X5,X6,X7,X8,X9] : (from_8(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
6. ! [X5] : (select_7(X5) <=> ? [X6,X7,X8,X9] : ($true & from_8(X5,X6,X7,X8,X9))) [input]
7. ! [X10] : (where_5_subquery_6(X10) <=> select_7(X10)) [input]
8. ! [X5] : (where_5(X5) <=> where_5_subquery_6(X5)) [input]
9. ! [X5,X6] : (select_3(X5,X6) <=> ? [X7,X8,X9] : (where_5(X5) & from_4(X5,X6,X7,X8,X9))) [input]
10. ! [X10,X11] : (main_query_2(X10,X11) <=> select_3(X10,X11)) [input]
54. ! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [input]
55. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [negated conjecture 54]
56. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
57. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
58. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
59. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
60. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 59]
61. ! [X0] : (where_5_subquery_6(X0) <=> select_7(X0)) [rectify 7]
62. ! [X0] : (where_5(X0) <=> where_5_subquery_6(X0)) [rectify 8]
63. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 9]
64. ! [X0,X1] : (main_query_2(X0,X1) <=> select_3(X0,X1)) [rectify 10]
179. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 55]
180. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
181. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0,X1))) [nnf transformation 56]
182. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) | ~select_1(X0,X1))) [rectify 181]
183. ! [X1,X0] : (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) => from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) [choice axiom]
184. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1))) [skolemisation 182,183]
185. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 3]
186. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 57]
187. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 58]
188. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 60]
189. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 188]
190. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
191. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_7(X0))) [skolemisation 189,190]
192. ! [X0] : ((where_5_subquery_6(X0) | ~select_7(X0)) & (select_7(X0) | ~where_5_subquery_6(X0))) [nnf transformation 61]
193. ! [X0] : ((where_5(X0) | ~where_5_subquery_6(X0)) & (where_5_subquery_6(X0) | ~where_5(X0))) [nnf transformation 62]
194. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0,X1))) [nnf transformation 63]
195. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) | ~select_3(X0,X1))) [rectify 194]
196. ! [X1,X0] : (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) => (where_5(X0) & from_4(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
197. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_3(X0,X1))) [skolemisation 195,196]
198. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 64]
201. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 179]
202. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
203. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 201,202]
204. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 180]
205. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 180]
206. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 184]
207. ~from_2(X0,X1,X2,X3,X4) | select_1(X0,X1) [cnf transformation 184]
208. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 185]
209. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 185]
210. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 186]
211. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 186]
213. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 187]
215. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 191]
217. ~select_7(X0) | where_5_subquery_6(X0) [cnf transformation 192]
219. ~where_5_subquery_6(X0) | where_5(X0) [cnf transformation 193]
220. from_4(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_3(X0,X1) [cnf transformation 197]
222. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0,X1) [cnf transformation 197]
223. ~main_query_2(X0,X1) | select_3(X0,X1) [cnf transformation 198]
224. ~select_3(X0,X1) | main_query_2(X0,X1) [cnf transformation 198]
266. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 203]
267. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 203]
269. 0 <=> main_query_1(sK10,sK11) [avatar definition]
270. main_query_1(sK10,sK11) <- {0} [avatar component clause 269]
272. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
273. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 272]
275. 2 <=> main_query_2(sK10,sK11) [avatar definition]
276. main_query_2(sK10,sK11) <- {2} [avatar component clause 275]
278. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
279. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 278]
280. ~1 | ~3 [avatar split clause 267,278,272]
281. 0 | 2 [avatar split clause 266,275,269]
282. select_3(sK10,sK11) <- {2} [resolution 223,276]
305. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 206,204]
310. tab1(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_3(X0,X1) [resolution 220,210]
352. from_4(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 305,211]
353. from_8(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 305,213]
359. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_3(X0,X1) [resolution 310,205]
486. ~select_1(X2,X3) | ~where_5(X2) | select_3(X2,X3) [resolution 352,222]
508. ~select_1(X2,X3) | select_7(X2) [resolution 353,215]
530. ~select_3(X2,X3) | select_1(X2,X3) [resolution 359,207]
553. select_1(sK10,sK11) <- {2} [resolution 530,282]
565. main_query_1(sK10,sK11) <- {2} [resolution 553,209]
566. $false <- {1, 2} [subsumption resolution 565,273]
567. 1 | ~2 [avatar contradiction clause 566]
568. select_1(sK10,sK11) <- {0} [resolution 270,208]
569. select_7(sK10) <- {0} [resolution 568,508]
570. ~where_5(sK10) | select_3(sK10,sK11) <- {0} [resolution 568,486]
576. 4 <=> select_3(sK10,sK11) [avatar definition]
577. select_3(sK10,sK11) <- {4} [avatar component clause 576]
579. 6 <=> where_5(sK10) [avatar definition]
582. 7 <=> ~where_5(sK10) [avatar definition]
584. 4 | ~7 | ~0 [avatar split clause 570,269,582,576]
585. where_5_subquery_6(sK10) <- {0} [resolution 569,217]
602. where_5(sK10) <- {0} [resolution 585,219]
606. 6 | ~0 [avatar split clause 602,269,579]
626. main_query_2(sK10,sK11) <- {4} [resolution 577,224]
628. $false <- {3, 4} [subsumption resolution 626,279]
629. 3 | ~4 [avatar contradiction clause 628]
630. $false [avatar sat refutation 280,281,567,584,606,629]
% SZS output end Proof for query_04_query_09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
