% Refutation found. Thanks to Tanya!
% SZS status Theorem for query12_query13_tptp4X
% SZS output start Proof for query12_query13_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1] : (where_6(X0,X1) <=> (~gt(X0,X1) & ~lt(X0,X1))) [input]
7. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
10. ! [X7,X8] : (eq(X7,X8) => eq(X8,X7)) [input]
12. ! [X7] : lte(X7,X7) [input]
13. ! [X7,X8] : ((lte(X8,X7) & lte(X7,X8)) <=> eq(X7,X8)) [input]
15. ! [X7,X8] : (lte(X8,X7) | lte(X7,X8)) [input]
16. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
17. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
49. ! [X7,X0,X1] : ((where_3(X0,X1) & eq(X1,X7)) => where_3(X0,X7)) [input]
55. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
56. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 55]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
58. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
60. ! [X0,X1] : (eq(X0,X1) => eq(X1,X0)) [rectify 10]
62. ! [X0] : lte(X0,X0) [rectify 12]
63. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
65. ! [X0,X1] : (lte(X1,X0) | lte(X0,X1)) [rectify 15]
66. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 16]
67. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 17]
99. ! [X0,X1,X2] : ((where_3(X1,X2) & eq(X2,X0)) => where_3(X1,X0)) [rectify 49]
105. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 56]
106. ! [X0,X1] : (eq(X1,X0) | ~eq(X0,X1)) [ennf transformation 60]
169. ! [X0,X1,X2] : (where_3(X1,X0) | (~where_3(X1,X2) | ~eq(X2,X0))) [ennf transformation 99]
170. ! [X0,X1,X2] : (where_3(X1,X0) | ~where_3(X1,X2) | ~eq(X2,X0)) [flattening 169]
181. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 105]
182. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
183. ! [X0,X1] : ((where_3(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [nnf transformation 2]
184. ! [X0,X1] : ((where_3(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [flattening 183]
185. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
186. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 185]
187. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
188. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 186,187]
189. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
190. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
191. ! [X0,X1] : ((where_6(X0,X1) | (gt(X0,X1) | lt(X0,X1))) & ((~gt(X0,X1) & ~lt(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 6]
192. ! [X0,X1] : ((where_6(X0,X1) | gt(X0,X1) | lt(X0,X1)) & ((~gt(X0,X1) & ~lt(X0,X1)) | ~where_6(X0,X1))) [flattening 191]
193. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 7]
194. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 193]
195. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
196. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 194,195]
197. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 58]
198. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 63]
199. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 198]
200. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 66]
201. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 200]
202. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 67]
203. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 181]
204. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
205. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 203,204]
206. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 182]
207. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 182]
208. ~where_3(X0,X1) | lte(X0,X1) [cnf transformation 184]
209. ~where_3(X0,X1) | lte(X1,X0) [cnf transformation 184]
210. ~lte(X1,X0) | where_3(X0,X1) | ~lte(X0,X1) [cnf transformation 184]
211. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 188]
212. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 188]
213. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 188]
214. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 189]
215. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 189]
216. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 190]
217. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 190]
218. ~lt(X0,X1) | ~where_6(X0,X1) [cnf transformation 192]
219. ~gt(X0,X1) | ~where_6(X0,X1) [cnf transformation 192]
220. gt(X0,X1) | where_6(X0,X1) | lt(X0,X1) [cnf transformation 192]
221. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 196]
222. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 196]
223. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 196]
224. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 197]
225. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 197]
227. ~eq(X0,X1) | eq(X1,X0) [cnf transformation 106]
229. lte(X0,X0) [cnf transformation 62]
230. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 199]
231. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 199]
232. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 199]
234. lte(X1,X0) | lte(X0,X1) [cnf transformation 65]
236. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 201]
237. ~lte(X0,X1) | eq(X0,X1) | lt(X0,X1) [cnf transformation 201]
238. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 202]
239. ~lt(X1,X0) | gt(X0,X1) [cnf transformation 202]
269. ~eq(X2,X0) | ~where_3(X1,X2) | where_3(X1,X0) [cnf transformation 170]
275. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 205]
276. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 205]
278. 0 <=> main_query_1(sK6,sK7) [avatar definition]
279. main_query_1(sK6,sK7) <- {0} [avatar component clause 278]
281. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
282. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 281]
284. 2 <=> main_query_2(sK6,sK7) [avatar definition]
285. main_query_2(sK6,sK7) <- {2} [avatar component clause 284]
287. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
288. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 287]
289. ~1 | ~3 [avatar split clause 276,287,281]
290. 0 | 2 [avatar split clause 275,284,278]
291. select_4(sK6,sK7) <- {2} [resolution 224,285]
293. where_6(sK6,sK7) <- {2} [resolution 291,222]
299. where_3(X0,X0) | ~lte(X0,X0) [resolution 210,229]
302. where_3(X0,X0) [subsumption resolution 299,229]
305. lt(X1,X0) | lt(X0,X1) | where_6(X0,X1) [resolution 220,238]
315. eq(X1,X2) | lt(X1,X2) | lte(X2,X1) [resolution 237,234]
317. lt(X1,X2) | lte(X2,X1) [subsumption resolution 315,232]
320. ~where_6(X5,X4) | lte(X4,X5) [resolution 317,218]
326. lte(sK7,sK6) <- {2} [resolution 320,293]
328. eq(sK7,sK6) | lt(sK7,sK6) <- {2} [resolution 326,237]
333. 5 <=> ~lt(sK7,sK6) [avatar definition]
334. ~lt(sK7,sK6) <- {5} [avatar component clause 333]
336. 4 <=> lt(sK7,sK6) [avatar definition]
337. lt(sK7,sK6) <- {4} [avatar component clause 336]
342. 6 <=> eq(sK7,sK6) [avatar definition]
343. eq(sK7,sK6) <- {6} [avatar component clause 342]
344. 4 | 6 | ~2 [avatar split clause 328,284,342,336]
346. 8 <=> lte(sK6,sK7) [avatar definition]
347. lte(sK6,sK7) <- {8} [avatar component clause 346]
355. 10 <=> eq(sK6,sK7) [avatar definition]
356. eq(sK6,sK7) <- {10} [avatar component clause 355]
362. 12 <=> where_3(sK6,sK7) [avatar definition]
363. where_3(sK6,sK7) <- {12} [avatar component clause 362]
366. gt(sK6,sK7) <- {4} [resolution 337,239]
374. ~where_6(sK6,sK7) <- {4} [resolution 366,219]
375. $false <- {2, 4} [subsumption resolution 374,293]
376. ~2 | ~4 [avatar contradiction clause 375]
377. select_1(sK6,sK7) <- {0} [resolution 279,214]
380. where_3(sK6,sK7) <- {0} [resolution 377,212]
381. 12 | ~0 [avatar split clause 380,278,362]
383. lte(sK7,sK6) <- {12} [resolution 363,209]
384. lte(sK6,sK7) <- {12} [resolution 363,208]
387. 8 | ~12 [avatar split clause 384,362,346]
391. eq(sK7,sK6) | ~lte(sK7,sK6) <- {8} [resolution 347,230]
394. 15 <=> ~lt(sK6,sK7) [avatar definition]
395. ~lt(sK6,sK7) <- {15} [avatar component clause 394]
400. eq(sK7,sK6) <- {8, 12} [subsumption resolution 391,383]
401. 6 | ~8 | ~12 [avatar split clause 400,362,346,342]
404. eq(sK6,sK7) <- {6} [resolution 343,227]
408. ~lt(sK7,sK6) <- {6} [resolution 343,236]
421. ~5 | ~6 [avatar split clause 408,342,333]
431. ~lt(sK6,sK7) <- {10} [resolution 356,236]
444. ~15 | ~10 [avatar split clause 431,355,394]
447. eq(sK6,sK7) | ~lte(sK6,sK7) <- {12} [resolution 383,230]
459. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 211,206]
461. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 221,216]
503. lt(sK6,sK7) | where_6(sK6,sK7) <- {5} [resolution 305,334]
513. where_6(sK6,sK7) <- {5, 15} [subsumption resolution 503,395]
531. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 459,217]
541. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 461,207]
695. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 531,223]
702. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {5, 15} [resolution 695,513]
704. select_4(sK6,sK7) <- {0, 5, 15} [subsumption resolution 702,377]
707. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 541,213]
715. main_query_2(sK6,sK7) <- {0, 5, 15} [resolution 704,225]
717. $false <- {0, 3, 5, 15} [subsumption resolution 715,288]
718. ~0 | 3 | 5 | 15 [avatar contradiction clause 717]
719. eq(sK6,sK7) <- {8, 12} [subsumption resolution 447,347]
720. 10 | ~8 | ~12 [avatar split clause 719,362,346,355]
721. 10 | ~6 [avatar split clause 404,342,355]
723. select_4(sK6,sK7) <- {2} [resolution 285,224]
730. lte(sK6,sK7) <- {10} [resolution 356,231]
744. ~where_3(X12,sK6) | where_3(X12,sK7) <- {10} [resolution 356,269]
953. where_3(sK6,sK7) <- {10} [resolution 744,302]
956. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 707,723]
960. select_1(sK6,sK7) <- {2, 12} [subsumption resolution 956,363]
972. main_query_1(sK6,sK7) <- {2, 12} [resolution 960,215]
974. $false <- {1, 2, 12} [subsumption resolution 972,282]
975. 1 | ~2 | ~12 [avatar contradiction clause 974]
978. 8 | ~10 [avatar split clause 730,355,346]
988. 12 | ~10 [avatar split clause 953,355,362]
990. $false [avatar sat refutation 289,290,344,376,381,387,401,421,444,718,720,721,975,978,988]
% SZS output end Proof for query12_query13_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.027 s
% ------------------------------
% ------------------------------
