% Refutation found. Thanks to Tanya!
% SZS status Theorem for query10_query11_tptp4X
% SZS output start Proof for query10_query11_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X1,X0] : (where_3(X1,X0) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X1,X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8] : (where_6(X7,X8) <=> eq(X7,X8)) [input]
7. ! [X7,X8] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X7,X8) & from_5(X7,X8,X9,X10,X11))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
54. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
55. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 54]
56. ! [X0,X1] : (where_3(X0,X1) <=> eq(X1,X0)) [rectify 2]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
58. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
59. ! [X0,X1] : (where_6(X0,X1) <=> eq(X0,X1)) [rectify 6]
60. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
61. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
107. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 55]
183. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 107]
184. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
185. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X1,X0)) & (eq(X1,X0) | ~where_3(X0,X1))) [nnf transformation 56]
186. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X1,X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X1,X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
187. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X1,X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X1,X0) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 186]
188. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X1,X0) & from_2(X0,X1,X5,X6,X7)) => (where_3(X1,X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
189. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X1,X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X1,X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 187,188]
190. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
191. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 58]
192. ! [X0,X1] : ((where_6(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_6(X0,X1))) [nnf transformation 59]
193. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 60]
194. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 193]
195. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
196. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 194,195]
197. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 61]
200. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 183]
201. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
202. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 200,201]
203. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 184]
204. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 184]
205. ~where_3(X0,X1) | eq(X1,X0) [cnf transformation 185]
206. ~eq(X1,X0) | where_3(X0,X1) [cnf transformation 185]
207. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 189]
208. ~select_1(X0,X1) | where_3(X1,X0) [cnf transformation 189]
209. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X1,X0) | select_1(X0,X1) [cnf transformation 189]
210. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 190]
211. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 190]
212. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 191]
213. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 191]
214. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 192]
215. ~eq(X0,X1) | where_6(X0,X1) [cnf transformation 192]
216. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 196]
217. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 196]
218. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 196]
219. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 197]
220. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 197]
264. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 202]
265. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 202]
302. 0 <=> main_query_1(sK6,sK7) [avatar definition]
303. main_query_1(sK6,sK7) <- {0} [avatar component clause 302]
305. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
306. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 305]
308. 2 <=> main_query_2(sK6,sK7) [avatar definition]
309. main_query_2(sK6,sK7) <- {2} [avatar component clause 308]
311. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
312. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 311]
313. ~1 | ~3 [avatar split clause 265,311,305]
314. 0 | 2 [avatar split clause 264,308,302]
319. select_4(sK6,sK7) <- {2} [resolution 219,309]
321. where_6(sK6,sK7) <- {2} [resolution 319,217]
323. eq(sK6,sK7) <- {2} [resolution 321,214]
328. where_3(sK7,sK6) <- {2} [resolution 323,206]
357. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 207,203]
359. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 216,212]
360. from_5(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 357,213]
363. from_2(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 359,204]
364. ~where_6(X0,X1) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 360,218]
382. ~select_4(X0,X1) | ~where_3(X1,X0) | select_1(X0,X1) [resolution 363,209]
384. ~where_3(sK7,sK6) | select_1(sK6,sK7) <- {2} [resolution 382,319]
385. select_1(sK6,sK7) <- {2} [subsumption resolution 384,328]
386. main_query_1(sK6,sK7) <- {2} [resolution 385,211]
388. $false <- {1, 2} [subsumption resolution 386,306]
389. 1 | ~2 [avatar contradiction clause 388]
390. select_1(sK6,sK7) <- {0} [resolution 303,210]
392. where_3(sK7,sK6) <- {0} [resolution 390,208]
393. eq(sK6,sK7) <- {0} [resolution 392,205]
398. where_6(sK6,sK7) <- {0} [resolution 393,215]
410. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 398,364]
412. select_4(sK6,sK7) <- {0} [subsumption resolution 410,390]
419. main_query_2(sK6,sK7) <- {0} [resolution 412,220]
421. $false <- {0, 3} [subsumption resolution 419,312]
422. ~0 | 3 [avatar contradiction clause 421]
423. $false [avatar sat refutation 313,314,389,422]
% SZS output end Proof for query10_query11_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.016 s
% ------------------------------
% ------------------------------
