% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_07_query_09_tptp4X
% SZS output start Proof for query_07_query_09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (where_6_subquery_7(X5) <=> select_8(X5)) [input]
9. ! [X0] : (where_6(X0) <=> where_6_subquery_7(X0)) [input]
10. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
12. ! [X7] : eq(X7,X7) [input]
18. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
19. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
55. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
56. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 55]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
58. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 7]
59. ! [X0] : (where_6_subquery_7(X0) <=> select_8(X0)) [rectify 8]
60. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
61. ! [X0] : eq(X0,X0) [rectify 12]
67. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 18]
68. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 19]
104. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 56]
176. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 104]
177. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
178. ! [X0] : ((where_3(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
179. ! [X0] : ((where_3(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [flattening 178]
180. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
181. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 180]
182. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
183. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 181,182]
184. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
185. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
186. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
187. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 58]
188. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 187]
189. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
190. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 188,189]
191. ! [X0] : ((where_6_subquery_7(X0) | ~select_8(X0)) & (select_8(X0) | ~where_6_subquery_7(X0))) [nnf transformation 59]
192. ! [X0] : ((where_6(X0) | ~where_6_subquery_7(X0)) & (where_6_subquery_7(X0) | ~where_6(X0))) [nnf transformation 9]
193. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 10]
194. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 193]
195. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
196. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 194,195]
197. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 60]
200. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 67]
201. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 200]
202. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 68]
203. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 176]
204. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
205. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 203,204]
206. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 177]
207. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 177]
210. gt(X0,X0) | where_3(X0) | lt(X0,X0) [cnf transformation 179]
211. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 183]
213. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0,X1) [cnf transformation 183]
214. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 184]
215. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 184]
216. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 185]
217. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 185]
219. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 186]
221. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 190]
223. ~select_8(X0) | where_6_subquery_7(X0) [cnf transformation 191]
225. ~where_6_subquery_7(X0) | where_6(X0) [cnf transformation 192]
226. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 196]
228. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0,X1) [cnf transformation 196]
229. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 197]
230. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 197]
231. eq(X0,X0) [cnf transformation 61]
240. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 201]
242. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 202]
277. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 205]
278. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 205]
280. 0 <=> main_query_1(sK10,sK11) [avatar definition]
281. main_query_1(sK10,sK11) <- {0} [avatar component clause 280]
283. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
284. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 283]
286. 2 <=> main_query_2(sK10,sK11) [avatar definition]
287. main_query_2(sK10,sK11) <- {2} [avatar component clause 286]
289. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
290. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 289]
291. ~1 | ~3 [avatar split clause 278,289,283]
292. 0 | 2 [avatar split clause 277,286,280]
293. select_4(sK10,sK11) <- {2} [resolution 229,287]
303. ~lt(X0,X0) [resolution 240,231]
305. where_3(X1) | lt(X1,X1) | lt(X1,X1) [resolution 210,242]
306. where_3(X1) | lt(X1,X1) [duplicate literal removal 305]
307. where_3(X1) [subsumption resolution 306,303]
325. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 211,206]
329. tab1(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 226,216]
371. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 325,217]
372. from_9(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 325,219]
378. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 329,207]
498. ~select_1(X2,X3) | ~where_6(X2) | select_4(X2,X3) [resolution 371,228]
520. ~select_1(X2,X3) | select_8(X2) [resolution 372,221]
542. ~select_4(X2,X3) | ~where_3(X2) | select_1(X2,X3) [resolution 378,213]
548. ~select_4(X2,X3) | select_1(X2,X3) [subsumption resolution 542,307]
556. select_1(sK10,sK11) <- {2} [resolution 548,293]
567. main_query_1(sK10,sK11) <- {2} [resolution 556,215]
569. $false <- {1, 2} [subsumption resolution 567,284]
570. 1 | ~2 [avatar contradiction clause 569]
571. select_1(sK10,sK11) <- {0} [resolution 281,214]
572. select_8(sK10) <- {0} [resolution 571,520]
573. ~where_6(sK10) | select_4(sK10,sK11) <- {0} [resolution 571,498]
580. 4 <=> select_4(sK10,sK11) [avatar definition]
581. select_4(sK10,sK11) <- {4} [avatar component clause 580]
583. 6 <=> where_6(sK10) [avatar definition]
586. 7 <=> ~where_6(sK10) [avatar definition]
588. 4 | ~7 | ~0 [avatar split clause 573,280,586,580]
589. where_6_subquery_7(sK10) <- {0} [resolution 572,223]
606. where_6(sK10) <- {0} [resolution 589,225]
610. 6 | ~0 [avatar split clause 606,280,583]
629. main_query_2(sK10,sK11) <- {4} [resolution 581,230]
631. $false <- {3, 4} [subsumption resolution 629,290]
632. 3 | ~4 [avatar contradiction clause 631]
633. $false [avatar sat refutation 291,292,570,588,610,632]
% SZS output end Proof for query_07_query_09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
