% Refutation found. Thanks to Tanya!
% SZS status Theorem for query11_query14_tptp4X
% SZS output start Proof for query11_query14_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X1,X0] : (where_3(X1,X0) <=> eq(X0,X1)) [input]
3. ! [X1,X0] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X1,X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8,X9,X10,X11] : (from_9(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
7. ! [X7] : (select_8(X7) <=> ? [X8,X9,X10,X11] : ($true & from_9(X7,X8,X9,X10,X11))) [input]
8. ! [X5] : (query_7(X5) <=> select_8(X5)) [input]
9. ! [X8,X7] : (where_6(X8,X7) <=> (query_7(X7) & eq(X7,X8))) [input]
10. ! [X8,X7] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X8,X7) & from_5(X7,X8,X9,X10,X11))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
62. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
63. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 62]
64. ! [X0,X1] : (where_3(X0,X1) <=> eq(X1,X0)) [rectify 2]
65. ! [X0,X1] : (select_1(X1,X0) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X1,X0,X2,X3,X4))) [rectify 3]
66. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
67. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
68. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 6]
69. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [rectify 7]
70. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 69]
71. ! [X0] : (query_7(X0) <=> select_8(X0)) [rectify 8]
72. ! [X0,X1] : (where_6(X0,X1) <=> (query_7(X1) & eq(X1,X0))) [rectify 9]
73. ! [X0,X1] : (select_4(X1,X0) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X1,X0,X2,X3,X4))) [rectify 10]
74. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
125. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 63]
211. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 125]
212. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
213. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X1,X0)) & (eq(X1,X0) | ~where_3(X0,X1))) [nnf transformation 64]
214. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X1,X0,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X1,X0,X2,X3,X4)) | ~select_1(X1,X0))) [nnf transformation 65]
215. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X1,X0,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X1,X0,X5,X6,X7)) | ~select_1(X1,X0))) [rectify 214]
216. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X1,X0,X5,X6,X7)) => (where_3(X0,X1) & from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
217. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X1,X0,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X1,X0))) [skolemisation 215,216]
218. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 66]
219. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 67]
220. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 68]
221. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 70]
222. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 221]
223. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
224. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 222,223]
225. ! [X0] : ((query_7(X0) | ~select_8(X0)) & (select_8(X0) | ~query_7(X0))) [nnf transformation 71]
226. ! [X0,X1] : ((where_6(X0,X1) | (~query_7(X1) | ~eq(X1,X0))) & ((query_7(X1) & eq(X1,X0)) | ~where_6(X0,X1))) [nnf transformation 72]
227. ! [X0,X1] : ((where_6(X0,X1) | ~query_7(X1) | ~eq(X1,X0)) & ((query_7(X1) & eq(X1,X0)) | ~where_6(X0,X1))) [flattening 226]
228. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X1,X0,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X1,X0,X2,X3,X4)) | ~select_4(X1,X0))) [nnf transformation 73]
229. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X1,X0,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X1,X0,X5,X6,X7)) | ~select_4(X1,X0))) [rectify 228]
230. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X1,X0,X5,X6,X7)) => (where_6(X0,X1) & from_5(X1,X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
231. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X1,X0,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X1,X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X1,X0))) [skolemisation 229,230]
232. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 74]
235. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 211]
236. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
237. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 235,236]
238. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 212]
239. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 212]
240. ~where_3(X0,X1) | eq(X1,X0) [cnf transformation 213]
241. ~eq(X1,X0) | where_3(X0,X1) [cnf transformation 213]
242. from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X1,X0) [cnf transformation 217]
243. ~select_1(X1,X0) | where_3(X0,X1) [cnf transformation 217]
244. ~from_2(X1,X0,X2,X3,X4) | ~where_3(X0,X1) | select_1(X1,X0) [cnf transformation 217]
245. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 218]
246. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 218]
247. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 219]
248. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 219]
250. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 220]
252. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 224]
254. ~select_8(X0) | query_7(X0) [cnf transformation 225]
255. ~where_6(X0,X1) | eq(X1,X0) [cnf transformation 227]
257. ~eq(X1,X0) | ~query_7(X1) | where_6(X0,X1) [cnf transformation 227]
258. from_5(X1,X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X1,X0) [cnf transformation 231]
259. ~select_4(X1,X0) | where_6(X0,X1) [cnf transformation 231]
260. ~from_5(X1,X0,X2,X3,X4) | ~where_6(X0,X1) | select_4(X1,X0) [cnf transformation 231]
261. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 232]
262. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 232]
311. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 237]
312. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 237]
354. 0 <=> main_query_1(sK10,sK11) [avatar definition]
355. main_query_1(sK10,sK11) <- {0} [avatar component clause 354]
357. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
358. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 357]
360. 2 <=> main_query_2(sK10,sK11) [avatar definition]
361. main_query_2(sK10,sK11) <- {2} [avatar component clause 360]
363. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
364. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 363]
365. ~1 | ~3 [avatar split clause 312,363,357]
366. 0 | 2 [avatar split clause 311,360,354]
369. select_4(sK10,sK11) <- {2} [resolution 261,361]
370. where_6(sK11,sK10) <- {2} [resolution 369,259]
372. eq(sK10,sK11) <- {2} [resolution 370,255]
382. where_3(sK11,sK10) <- {2} [resolution 372,241]
430. tab1(X2,X3,sK0(X3,X2),sK1(X3,X2),sK2(X3,X2)) | ~select_1(X2,X3) [resolution 242,238]
434. tab1(X2,X3,sK7(X3,X2),sK8(X3,X2),sK9(X3,X2)) | ~select_4(X2,X3) [resolution 258,247]
435. from_9(X0,X1,sK0(X1,X0),sK1(X1,X0),sK2(X1,X0)) | ~select_1(X0,X1) [resolution 430,250]
436. from_5(X2,X3,sK0(X3,X2),sK1(X3,X2),sK2(X3,X2)) | ~select_1(X2,X3) [resolution 430,248]
442. from_2(X4,X5,sK7(X5,X4),sK8(X5,X4),sK9(X5,X4)) | ~select_4(X4,X5) [resolution 434,239]
444. ~select_1(X2,X3) | select_8(X2) [resolution 435,252]
445. ~where_6(X1,X0) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 436,260]
452. ~select_4(X0,X1) | ~where_3(X1,X0) | select_1(X0,X1) [resolution 442,244]
454. ~where_3(sK11,sK10) | select_1(sK10,sK11) <- {2} [resolution 452,369]
455. select_1(sK10,sK11) <- {2} [subsumption resolution 454,382]
457. main_query_1(sK10,sK11) <- {2} [resolution 455,246]
459. $false <- {1, 2} [subsumption resolution 457,358]
460. 1 | ~2 [avatar contradiction clause 459]
461. select_1(sK10,sK11) <- {0} [resolution 355,245]
462. select_8(sK10) <- {0} [resolution 461,444]
464. where_3(sK11,sK10) <- {0} [resolution 461,243]
465. query_7(sK10) <- {0} [resolution 462,254]
466. eq(sK10,sK11) <- {0} [resolution 464,240]
469. ~query_7(sK10) | where_6(sK11,sK10) <- {0} [resolution 466,257]
474. where_6(sK11,sK10) <- {0} [subsumption resolution 469,465]
485. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {0} [resolution 474,445]
488. select_4(sK10,sK11) <- {0} [subsumption resolution 485,461]
494. main_query_2(sK10,sK11) <- {0} [resolution 488,262]
496. $false <- {0, 3} [subsumption resolution 494,364]
497. ~0 | 3 [avatar contradiction clause 496]
498. $false [avatar sat refutation 365,366,460,497]
% SZS output end Proof for query11_query14_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.012 s
% ------------------------------
% ------------------------------
