% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_05_query_08_tptp4X
% SZS output start Proof for query_05_query_08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_8(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8] : (select_7(X7,X8) <=> ? [X9,X10,X11] : ($true & from_8(X7,X8,X9,X10,X11))) [input]
7. ! [X5,X6] : (subquery_from_6(X5,X6) <=> select_7(X5,X6)) [input]
8. ! [X12,X13] : (from_5(X12,X13) <=> subquery_from_6(X12,X13)) [input]
9. ! [X12,X13] : (select_4(X12,X13) <=> ($true & from_5(X12,X13))) [input]
10. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
11. ! [X14] : eq(X14,X14) [input]
55. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
56. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 55]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
58. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
59. ! [X0,X1] : (select_7(X0,X1) <=> ? [X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
60. ! [X0,X1] : (select_7(X0,X1) <=> ? [X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 59]
61. ! [X0,X1] : (subquery_from_6(X0,X1) <=> select_7(X0,X1)) [rectify 7]
62. ! [X0,X1] : (from_5(X0,X1) <=> subquery_from_6(X0,X1)) [rectify 8]
63. ! [X0,X1] : (select_4(X0,X1) <=> ($true & from_5(X0,X1))) [rectify 9]
64. ! [X0,X1] : (select_4(X0,X1) <=> from_5(X0,X1)) [true and false elimination 63]
65. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 10]
66. ! [X0] : eq(X0,X0) [rectify 11]
110. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 56]
184. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 110]
185. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
186. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
187. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
188. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 187]
189. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
190. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 188,189]
191. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
192. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 58]
193. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0,X1))) [nnf transformation 60]
194. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_8(X0,X1,X5,X6,X7) | ~select_7(X0,X1))) [rectify 193]
195. ! [X1,X0] : (? [X5,X6,X7] : from_8(X0,X1,X5,X6,X7) => from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) [choice axiom]
196. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1))) [skolemisation 194,195]
197. ! [X0,X1] : ((subquery_from_6(X0,X1) | ~select_7(X0,X1)) & (select_7(X0,X1) | ~subquery_from_6(X0,X1))) [nnf transformation 61]
198. ! [X0,X1] : ((from_5(X0,X1) | ~subquery_from_6(X0,X1)) & (subquery_from_6(X0,X1) | ~from_5(X0,X1))) [nnf transformation 62]
199. ! [X0,X1] : ((select_4(X0,X1) | ~from_5(X0,X1)) & (from_5(X0,X1) | ~select_4(X0,X1))) [nnf transformation 64]
200. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 65]
203. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 184]
204. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
205. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 203,204]
206. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 185]
207. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 185]
209. ~eq(X0,X0) | where_3(X0) [cnf transformation 186]
210. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 190]
212. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0,X1) [cnf transformation 190]
213. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 191]
214. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 191]
215. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 192]
216. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 192]
217. from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1) [cnf transformation 196]
218. ~from_8(X0,X1,X2,X3,X4) | select_7(X0,X1) [cnf transformation 196]
219. ~subquery_from_6(X0,X1) | select_7(X0,X1) [cnf transformation 197]
220. ~select_7(X0,X1) | subquery_from_6(X0,X1) [cnf transformation 197]
221. ~from_5(X0,X1) | subquery_from_6(X0,X1) [cnf transformation 198]
222. ~subquery_from_6(X0,X1) | from_5(X0,X1) [cnf transformation 198]
223. ~select_4(X0,X1) | from_5(X0,X1) [cnf transformation 199]
224. ~from_5(X0,X1) | select_4(X0,X1) [cnf transformation 199]
225. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 200]
226. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 200]
227. eq(X0,X0) [cnf transformation 66]
269. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 205]
270. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 205]
272. 0 <=> main_query_1(sK6,sK7) [avatar definition]
273. main_query_1(sK6,sK7) <- {0} [avatar component clause 272]
275. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
276. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 275]
278. 2 <=> main_query_2(sK6,sK7) [avatar definition]
279. main_query_2(sK6,sK7) <- {2} [avatar component clause 278]
281. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
282. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 281]
283. ~1 | ~3 [avatar split clause 270,281,275]
284. 0 | 2 [avatar split clause 269,278,272]
286. where_3(X1) [resolution 209,227]
287. select_4(sK6,sK7) <- {2} [resolution 225,279]
289. from_5(sK6,sK7) <- {2} [resolution 287,223]
293. subquery_from_6(sK6,sK7) <- {2} [resolution 289,221]
297. select_7(sK6,sK7) <- {2} [resolution 293,219]
335. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 210,206]
337. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1) [resolution 217,215]
350. from_8(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 335,216]
356. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1) [resolution 337,207]
364. ~select_1(X2,X3) | select_7(X2,X3) [resolution 350,218]
385. ~select_7(X2,X3) | ~where_3(X2) | select_1(X2,X3) [resolution 356,212]
391. ~select_7(X2,X3) | select_1(X2,X3) [subsumption resolution 385,286]
399. select_1(sK6,sK7) <- {2} [resolution 391,297]
410. main_query_1(sK6,sK7) <- {2} [resolution 399,214]
412. $false <- {1, 2} [subsumption resolution 410,276]
413. 1 | ~2 [avatar contradiction clause 412]
416. select_1(sK6,sK7) <- {0} [resolution 273,213]
417. select_7(sK6,sK7) <- {0} [resolution 416,364]
425. subquery_from_6(sK6,sK7) <- {0} [resolution 417,220]
430. from_5(sK6,sK7) <- {0} [resolution 425,222]
436. select_4(sK6,sK7) <- {0} [resolution 430,224]
447. main_query_2(sK6,sK7) <- {0} [resolution 436,226]
449. $false <- {0, 3} [subsumption resolution 447,282]
450. ~0 | 3 [avatar contradiction clause 449]
451. $false [avatar sat refutation 283,284,413,450]
% SZS output end Proof for query_05_query_08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
