% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_02_query_05_tptp4X
% SZS output start Proof for query_02_query_05_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0,X5] : (main_query_1(X0,X5) <=> select_1(X0,X5)) [input]
4. ! [X6,X1,X2,X3,X4] : (from_4(X6,X1,X2,X3,X4) <=> tab1(X6,X1,X2,X3,X4)) [input]
5. ! [X6] : (where_5(X6) <=> eq(X6,X6)) [input]
6. ! [X6,X1] : (select_3(X6,X1) <=> ? [X2,X3,X4] : (where_5(X6) & from_4(X6,X1,X2,X3,X4))) [input]
7. ! [X7,X5] : (main_query_2(X7,X5) <=> select_3(X7,X5)) [input]
8. ! [X8] : eq(X8,X8) [input]
46. ! [X0,X5] : (main_query_1(X0,X5) <=> main_query_2(X0,X5)) [input]
47. ~! [X0,X5] : (main_query_1(X0,X5) <=> main_query_2(X0,X5)) [negated conjecture 46]
48. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 3]
50. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
51. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [rectify 5]
52. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
53. ! [X0,X1] : (main_query_2(X0,X1) <=> select_3(X0,X1)) [rectify 7]
54. ! [X0] : eq(X0,X0) [rectify 8]
92. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 47]
154. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 92]
155. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
156. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0,X1))) [nnf transformation 48]
157. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) | ~select_1(X0,X1))) [rectify 156]
158. ! [X1,X0] : (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) => from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) [choice axiom]
159. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1))) [skolemisation 157,158]
160. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 49]
161. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 50]
162. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 51]
163. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0,X1))) [nnf transformation 52]
164. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) | ~select_3(X0,X1))) [rectify 163]
165. ! [X1,X0] : (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) => (where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
166. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_3(X0,X1))) [skolemisation 164,165]
167. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 53]
170. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 154]
171. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
172. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 170,171]
173. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 155]
174. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 155]
175. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 159]
176. ~from_2(X0,X1,X2,X3,X4) | select_1(X0,X1) [cnf transformation 159]
177. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 160]
178. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 160]
179. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 161]
180. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 161]
182. ~eq(X0,X0) | where_5(X0) [cnf transformation 162]
183. from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [cnf transformation 166]
185. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0,X1) [cnf transformation 166]
186. ~main_query_2(X0,X1) | select_3(X0,X1) [cnf transformation 167]
187. ~select_3(X0,X1) | main_query_2(X0,X1) [cnf transformation 167]
188. eq(X0,X0) [cnf transformation 54]
224. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 172]
225. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 172]
227. 0 <=> main_query_1(sK6,sK7) [avatar definition]
228. main_query_1(sK6,sK7) <- {0} [avatar component clause 227]
230. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
231. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 230]
233. 2 <=> main_query_2(sK6,sK7) [avatar definition]
234. main_query_2(sK6,sK7) <- {2} [avatar component clause 233]
236. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
237. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 236]
238. ~1 | ~3 [avatar split clause 225,236,230]
239. 0 | 2 [avatar split clause 224,233,227]
241. where_5(X1) [resolution 182,188]
242. select_3(sK6,sK7) <- {2} [resolution 186,234]
275. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 175,173]
277. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 183,179]
290. from_4(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 275,180]
296. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 277,174]
311. ~select_1(X2,X3) | ~where_5(X2) | select_3(X2,X3) [resolution 290,185]
317. ~select_1(X2,X3) | select_3(X2,X3) [subsumption resolution 311,241]
333. ~select_3(X2,X3) | select_1(X2,X3) [resolution 296,176]
346. select_1(sK6,sK7) <- {2} [resolution 333,242]
355. main_query_1(sK6,sK7) <- {2} [resolution 346,178]
356. $false <- {1, 2} [subsumption resolution 355,231]
357. 1 | ~2 [avatar contradiction clause 356]
358. select_1(sK6,sK7) <- {0} [resolution 228,177]
359. select_3(sK6,sK7) <- {0} [resolution 358,317]
364. main_query_2(sK6,sK7) <- {0} [resolution 359,187]
366. $false <- {0, 3} [subsumption resolution 364,237]
367. ~0 | 3 [avatar contradiction clause 366]
368. $false [avatar sat refutation 238,239,357,367]
% SZS output end Proof for query_02_query_05_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
