% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query09_tptp4X
% SZS output start Proof for query00_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_7(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (select_6(X6) <=> ? [X7,X8,X9,X10] : ($true & from_7(X6,X7,X8,X9,X10))) [input]
6. ! [X6] : (subquery_from_5(X6) <=> select_6(X6)) [input]
7. ! [X11] : (from_4(X11) <=> subquery_from_5(X11)) [input]
8. ! [X11] : (select_3(X11) <=> ($true & from_4(X11))) [input]
9. ! [X6] : (main_query_2(X6) <=> select_3(X6)) [input]
41. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
42. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
45. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
46. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [rectify 5]
47. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 46]
48. ! [X0] : (subquery_from_5(X0) <=> select_6(X0)) [rectify 6]
49. ! [X0] : (from_4(X0) <=> subquery_from_5(X0)) [rectify 7]
50. ! [X0] : (select_3(X0) <=> ($true & from_4(X0))) [rectify 8]
51. ! [X0] : (select_3(X0) <=> from_4(X0)) [true and false elimination 50]
52. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 9]
84. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 42]
130. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 84]
131. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
132. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
133. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 132]
134. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 133,134]
136. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 44]
137. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 45]
138. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 47]
139. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 138]
140. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
141. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 139,140]
142. ! [X0] : ((subquery_from_5(X0) | ~select_6(X0)) & (select_6(X0) | ~subquery_from_5(X0))) [nnf transformation 48]
143. ! [X0] : ((from_4(X0) | ~subquery_from_5(X0)) & (subquery_from_5(X0) | ~from_4(X0))) [nnf transformation 49]
144. ! [X0] : ((select_3(X0) | ~from_4(X0)) & (from_4(X0) | ~select_3(X0))) [nnf transformation 51]
145. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 52]
148. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 130]
149. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
150. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 148,149]
151. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 131]
152. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 131]
153. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 135]
154. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 135]
155. ~main_query_1(X0) | select_1(X0) [cnf transformation 136]
156. ~select_1(X0) | main_query_1(X0) [cnf transformation 136]
157. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 137]
158. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 137]
159. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 141]
160. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 141]
161. ~subquery_from_5(X0) | select_6(X0) [cnf transformation 142]
162. ~select_6(X0) | subquery_from_5(X0) [cnf transformation 142]
163. ~from_4(X0) | subquery_from_5(X0) [cnf transformation 143]
164. ~subquery_from_5(X0) | from_4(X0) [cnf transformation 143]
165. ~select_3(X0) | from_4(X0) [cnf transformation 144]
166. ~from_4(X0) | select_3(X0) [cnf transformation 144]
167. ~main_query_2(X0) | select_3(X0) [cnf transformation 145]
168. ~select_3(X0) | main_query_2(X0) [cnf transformation 145]
198. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 150]
199. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 150]
201. 0 <=> main_query_1(sK8) [avatar definition]
202. main_query_1(sK8) <- {0} [avatar component clause 201]
204. 1 <=> ~main_query_1(sK8) [avatar definition]
205. ~main_query_1(sK8) <- {1} [avatar component clause 204]
207. 2 <=> main_query_2(sK8) [avatar definition]
208. main_query_2(sK8) <- {2} [avatar component clause 207]
210. 3 <=> ~main_query_2(sK8) [avatar definition]
211. ~main_query_2(sK8) <- {3} [avatar component clause 210]
212. ~1 | ~3 [avatar split clause 199,210,204]
213. 0 | 2 [avatar split clause 198,207,201]
214. select_3(sK8) <- {2} [resolution 167,208]
216. from_4(sK8) <- {2} [resolution 214,165]
219. subquery_from_5(sK8) <- {2} [resolution 216,163]
222. select_6(sK8) <- {2} [resolution 219,161]
251. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 153,151]
253. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 159,157]
256. from_7(X2,sK0(X2),sK1(X2),sK2(X2),sK3(X2)) | ~select_1(X2) [resolution 251,158]
262. from_2(X5,sK4(X5),sK5(X5),sK6(X5),sK7(X5)) | ~select_6(X5) [resolution 253,152]
266. ~select_1(X1) | select_6(X1) [resolution 256,160]
272. ~select_6(X1) | select_1(X1) [resolution 262,154]
273. select_1(sK8) <- {2} [resolution 272,222]
275. main_query_1(sK8) <- {2} [resolution 273,156]
276. $false <- {1, 2} [subsumption resolution 275,205]
277. 1 | ~2 [avatar contradiction clause 276]
278. select_1(sK8) <- {0} [resolution 202,155]
279. select_6(sK8) <- {0} [resolution 278,266]
282. subquery_from_5(sK8) <- {0} [resolution 279,162]
283. from_4(sK8) <- {0} [resolution 282,164]
287. select_3(sK8) <- {0} [resolution 283,166]
291. main_query_2(sK8) <- {0} [resolution 287,168]
293. $false <- {0, 3} [subsumption resolution 291,211]
294. ~0 | 3 [avatar contradiction clause 293]
295. $false [avatar sat refutation 212,213,277,294]
% SZS output end Proof for query00_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.014 s
% ------------------------------
% ------------------------------
