% Refutation found. Thanks to Tanya!
% SZS status Theorem for query04_query05_tptp4X
% SZS output start Proof for query04_query05_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> eq(X0,X0)) [input]
7. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
48. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
49. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 48]
50. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
51. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
52. ! [X0] : eq(X0,X0) [rectify 9]
91. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 49]
155. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 91]
156. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
157. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
158. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
159. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 158]
160. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
161. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 159,160]
162. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 50]
163. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
164. ! [X0] : ((where_6(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_6(X0))) [nnf transformation 6]
165. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 7]
166. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 165]
167. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
168. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 166,167]
169. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 51]
172. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 155]
173. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
174. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 172,173]
175. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 156]
176. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 156]
177. eq(X0,X0) | ~where_3(X0) [cnf transformation 157]
178. ~eq(X0,X0) | where_3(X0) [cnf transformation 157]
179. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 161]
181. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 161]
182. ~main_query_1(X0) | select_1(X0) [cnf transformation 162]
183. ~select_1(X0) | main_query_1(X0) [cnf transformation 162]
184. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 163]
185. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 163]
187. ~eq(X0,X0) | where_6(X0) [cnf transformation 164]
188. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 168]
190. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 168]
191. ~main_query_2(X0) | select_4(X0) [cnf transformation 169]
192. ~select_4(X0) | main_query_2(X0) [cnf transformation 169]
193. eq(X0,X0) [cnf transformation 52]
230. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 174]
231. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 174]
262. 0 <=> main_query_1(sK8) [avatar definition]
263. main_query_1(sK8) <- {0} [avatar component clause 262]
265. 1 <=> ~main_query_1(sK8) [avatar definition]
266. ~main_query_1(sK8) <- {1} [avatar component clause 265]
268. 2 <=> main_query_2(sK8) [avatar definition]
269. main_query_2(sK8) <- {2} [avatar component clause 268]
271. 3 <=> ~main_query_2(sK8) [avatar definition]
272. ~main_query_2(sK8) <- {3} [avatar component clause 271]
273. ~1 | ~3 [avatar split clause 231,271,265]
274. 0 | 2 [avatar split clause 230,268,262]
275. select_4(sK8) <- {2} [resolution 191,269]
279. where_3(X1) [resolution 178,193]
282. where_6(X1) | ~where_3(X1) [resolution 187,177]
284. where_6(X1) [subsumption resolution 282,279]
299. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 179,175]
301. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 188,184]
303. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 299,185]
306. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 301,176]
308. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 303,190]
309. ~select_1(X1) | select_4(X1) [subsumption resolution 308,284]
311. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 306,181]
312. ~select_4(X1) | select_1(X1) [subsumption resolution 311,279]
313. select_1(sK8) <- {2} [resolution 312,275]
315. main_query_1(sK8) <- {2} [resolution 313,183]
317. $false <- {1, 2} [subsumption resolution 315,266]
318. 1 | ~2 [avatar contradiction clause 317]
319. select_1(sK8) <- {0} [resolution 263,182]
320. select_4(sK8) <- {0} [resolution 319,309]
324. main_query_2(sK8) <- {0} [resolution 320,192]
326. $false <- {0, 3} [subsumption resolution 324,272]
327. ~0 | 3 [avatar contradiction clause 326]
328. $false [avatar sat refutation 273,274,318,327]
% SZS output end Proof for query04_query05_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
