% Refutation found. Thanks to Tanya!
% SZS status Theorem for query03_query09_tptp4X
% SZS output start Proof for query03_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [input]
6. ! [X0] : (subquery_from_5(X0) <=> select_6(X0)) [input]
7. ! [X5] : (from_4(X5) <=> subquery_from_5(X5)) [input]
8. ! [X5] : (select_3(X5) <=> ($true & from_4(X5))) [input]
9. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
44. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
45. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 44]
46. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
47. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 5]
48. ! [X0] : (from_4(X0) <=> subquery_from_5(X0)) [rectify 7]
49. ! [X0] : (select_3(X0) <=> ($true & from_4(X0))) [rectify 8]
50. ! [X0] : (select_3(X0) <=> from_4(X0)) [true and false elimination 49]
136. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 45]
137. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 46]
139. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 138]
140. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
141. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 139,140]
142. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
143. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 4]
144. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 47]
145. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 144]
146. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
147. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 145,146]
148. ! [X0] : ((subquery_from_5(X0) | ~select_6(X0)) & (select_6(X0) | ~subquery_from_5(X0))) [nnf transformation 6]
149. ! [X0] : ((from_4(X0) | ~subquery_from_5(X0)) & (subquery_from_5(X0) | ~from_4(X0))) [nnf transformation 48]
150. ! [X0] : ((select_3(X0) | ~from_4(X0)) & (from_4(X0) | ~select_3(X0))) [nnf transformation 50]
151. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 9]
154. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 136]
155. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
156. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 154,155]
157. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 137]
158. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 137]
159. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 141]
160. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 141]
161. ~main_query_1(X0) | select_1(X0) [cnf transformation 142]
162. ~select_1(X0) | main_query_1(X0) [cnf transformation 142]
163. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 143]
164. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 143]
165. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 147]
166. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 147]
167. ~subquery_from_5(X0) | select_6(X0) [cnf transformation 148]
168. ~select_6(X0) | subquery_from_5(X0) [cnf transformation 148]
169. ~from_4(X0) | subquery_from_5(X0) [cnf transformation 149]
170. ~subquery_from_5(X0) | from_4(X0) [cnf transformation 149]
171. ~select_3(X0) | from_4(X0) [cnf transformation 150]
172. ~from_4(X0) | select_3(X0) [cnf transformation 150]
173. ~main_query_2(X0) | select_3(X0) [cnf transformation 151]
174. ~select_3(X0) | main_query_2(X0) [cnf transformation 151]
207. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 156]
208. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 156]
210. 0 <=> main_query_1(sK8) [avatar definition]
211. main_query_1(sK8) <- {0} [avatar component clause 210]
213. 1 <=> ~main_query_1(sK8) [avatar definition]
214. ~main_query_1(sK8) <- {1} [avatar component clause 213]
216. 2 <=> main_query_2(sK8) [avatar definition]
217. main_query_2(sK8) <- {2} [avatar component clause 216]
219. 3 <=> ~main_query_2(sK8) [avatar definition]
220. ~main_query_2(sK8) <- {3} [avatar component clause 219]
221. ~1 | ~3 [avatar split clause 208,219,213]
222. 0 | 2 [avatar split clause 207,216,210]
223. select_3(sK8) <- {2} [resolution 173,217]
225. from_4(sK8) <- {2} [resolution 223,171]
228. subquery_from_5(sK8) <- {2} [resolution 225,169]
231. select_6(sK8) <- {2} [resolution 228,167]
260. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 159,157]
262. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 165,163]
264. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 260,164]
267. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 262,158]
269. ~select_1(X1) | select_6(X1) [resolution 264,166]
271. ~select_6(X1) | select_1(X1) [resolution 267,160]
272. select_1(sK8) <- {2} [resolution 271,231]
276. main_query_1(sK8) <- {2} [resolution 272,162]
277. $false <- {1, 2} [subsumption resolution 276,214]
278. 1 | ~2 [avatar contradiction clause 277]
279. select_1(sK8) <- {0} [resolution 211,161]
282. select_6(sK8) <- {0} [resolution 279,269]
287. subquery_from_5(sK8) <- {0} [resolution 282,168]
290. from_4(sK8) <- {0} [resolution 287,170]
294. select_3(sK8) <- {0} [resolution 290,172]
298. main_query_2(sK8) <- {0} [resolution 294,174]
300. $false <- {0, 3} [subsumption resolution 298,220]
301. ~0 | 3 [avatar contradiction clause 300]
302. $false [avatar sat refutation 221,222,278,301]
% SZS output end Proof for query03_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
