% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_05_query_09_tptp4X
% SZS output start Proof for query_05_query_09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (where_6_subquery_7(X5) <=> select_8(X5)) [input]
9. ! [X0] : (where_6(X0) <=> where_6_subquery_7(X0)) [input]
10. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
12. ! [X7] : eq(X7,X7) [input]
55. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
56. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 55]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
58. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 7]
59. ! [X0] : (where_6_subquery_7(X0) <=> select_8(X0)) [rectify 8]
60. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
61. ! [X0] : eq(X0,X0) [rectify 12]
104. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 56]
176. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 104]
177. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
178. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
179. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
180. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 179]
181. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
182. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 180,181]
183. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
184. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
185. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
186. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 58]
187. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 186]
188. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
189. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 187,188]
190. ! [X0] : ((where_6_subquery_7(X0) | ~select_8(X0)) & (select_8(X0) | ~where_6_subquery_7(X0))) [nnf transformation 59]
191. ! [X0] : ((where_6(X0) | ~where_6_subquery_7(X0)) & (where_6_subquery_7(X0) | ~where_6(X0))) [nnf transformation 9]
192. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 10]
193. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 192]
194. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
195. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 193,194]
196. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 60]
199. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 176]
200. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
201. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 199,200]
202. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 177]
203. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 177]
205. ~eq(X0,X0) | where_3(X0) [cnf transformation 178]
206. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 182]
208. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0,X1) [cnf transformation 182]
209. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 183]
210. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 183]
211. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 184]
212. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 184]
214. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 185]
216. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 189]
218. ~select_8(X0) | where_6_subquery_7(X0) [cnf transformation 190]
220. ~where_6_subquery_7(X0) | where_6(X0) [cnf transformation 191]
221. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 195]
223. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0,X1) [cnf transformation 195]
224. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 196]
225. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 196]
226. eq(X0,X0) [cnf transformation 61]
267. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 201]
268. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 201]
270. 0 <=> main_query_1(sK10,sK11) [avatar definition]
271. main_query_1(sK10,sK11) <- {0} [avatar component clause 270]
273. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
274. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 273]
276. 2 <=> main_query_2(sK10,sK11) [avatar definition]
277. main_query_2(sK10,sK11) <- {2} [avatar component clause 276]
279. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
280. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 279]
281. ~1 | ~3 [avatar split clause 268,279,273]
282. 0 | 2 [avatar split clause 267,276,270]
284. where_3(X1) [resolution 205,226]
285. select_4(sK10,sK11) <- {2} [resolution 224,277]
316. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 206,202]
321. tab1(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 221,211]
413. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 316,212]
414. from_9(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 316,214]
420. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 321,203]
541. ~select_1(X2,X3) | ~where_6(X2) | select_4(X2,X3) [resolution 413,223]
596. ~select_1(X2,X3) | select_8(X2) [resolution 414,216]
605. ~select_4(X2,X3) | ~where_3(X2) | select_1(X2,X3) [resolution 420,208]
611. ~select_4(X2,X3) | select_1(X2,X3) [subsumption resolution 605,284]
612. select_1(sK10,sK11) <- {2} [resolution 611,285]
627. main_query_1(sK10,sK11) <- {2} [resolution 612,210]
629. $false <- {1, 2} [subsumption resolution 627,274]
630. 1 | ~2 [avatar contradiction clause 629]
633. select_1(sK10,sK11) <- {0} [resolution 271,209]
634. select_8(sK10) <- {0} [resolution 633,596]
635. ~where_6(sK10) | select_4(sK10,sK11) <- {0} [resolution 633,541]
644. 4 <=> select_4(sK10,sK11) [avatar definition]
645. select_4(sK10,sK11) <- {4} [avatar component clause 644]
647. 6 <=> where_6(sK10) [avatar definition]
650. 7 <=> ~where_6(sK10) [avatar definition]
652. 4 | ~7 | ~0 [avatar split clause 635,270,650,644]
653. where_6_subquery_7(sK10) <- {0} [resolution 634,218]
658. where_6(sK10) <- {0} [resolution 653,220]
662. 6 | ~0 [avatar split clause 658,270,647]
671. main_query_2(sK10,sK11) <- {4} [resolution 645,225]
673. $false <- {3, 4} [subsumption resolution 671,280]
674. 3 | ~4 [avatar contradiction clause 673]
675. $false [avatar sat refutation 281,282,630,652,662,674]
% SZS output end Proof for query_05_query_09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
