% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_05_query_07_tptp4X
% SZS output start Proof for query_05_query_07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
7. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
15. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
16. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
47. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
48. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 47]
49. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
50. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
51. ! [X0] : eq(X0,X0) [rectify 9]
57. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 15]
58. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 16]
89. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 48]
151. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 89]
152. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
153. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
154. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
155. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 154]
156. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
157. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 155,156]
158. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 49]
159. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
160. ! [X0] : ((where_6(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [nnf transformation 6]
161. ! [X0] : ((where_6(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [flattening 160]
162. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 7]
163. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 162]
164. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
165. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 163,164]
166. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 50]
169. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 57]
170. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 169]
171. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 58]
172. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 151]
173. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
174. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 172,173]
175. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 152]
176. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 152]
178. ~eq(X0,X0) | where_3(X0) [cnf transformation 153]
179. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 157]
181. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0,X1) [cnf transformation 157]
182. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 158]
183. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 158]
184. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 159]
185. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 159]
188. gt(X0,X0) | where_6(X0) | lt(X0,X0) [cnf transformation 161]
189. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 165]
191. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0,X1) [cnf transformation 165]
192. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 166]
193. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 166]
194. eq(X0,X0) [cnf transformation 51]
203. ~lt(X0,X1) | ~eq(X0,X1) [cnf transformation 170]
205. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 171]
235. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 174]
236. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 174]
238. 0 <=> main_query_1(sK6,sK7) [avatar definition]
239. main_query_1(sK6,sK7) <- {0} [avatar component clause 238]
241. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
242. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 241]
244. 2 <=> main_query_2(sK6,sK7) [avatar definition]
245. main_query_2(sK6,sK7) <- {2} [avatar component clause 244]
247. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
248. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 247]
249. ~1 | ~3 [avatar split clause 236,247,241]
250. 0 | 2 [avatar split clause 235,244,238]
252. where_3(X1) [resolution 178,194]
253. select_4(sK6,sK7) <- {2} [resolution 192,245]
263. where_6(X1) | lt(X1,X1) | lt(X1,X1) [resolution 188,205]
264. lt(X1,X1) | where_6(X1) [duplicate literal removal 263]
269. where_6(X2) | ~eq(X2,X2) [resolution 264,203]
271. where_6(X2) [subsumption resolution 269,194]
289. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 179,175]
291. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 189,184]
304. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 289,185]
310. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 291,176]
318. ~select_1(X2,X3) | ~where_6(X2) | select_4(X2,X3) [resolution 304,191]
324. ~select_1(X2,X3) | select_4(X2,X3) [subsumption resolution 318,271]
340. ~select_4(X2,X3) | ~where_3(X2) | select_1(X2,X3) [resolution 310,181]
346. ~select_4(X2,X3) | select_1(X2,X3) [subsumption resolution 340,252]
354. select_1(sK6,sK7) <- {2} [resolution 346,253]
365. main_query_1(sK6,sK7) <- {2} [resolution 354,183]
367. $false <- {1, 2} [subsumption resolution 365,242]
368. 1 | ~2 [avatar contradiction clause 367]
371. select_1(sK6,sK7) <- {0} [resolution 239,182]
372. select_4(sK6,sK7) <- {0} [resolution 371,324]
380. main_query_2(sK6,sK7) <- {0} [resolution 372,193]
382. $false <- {0, 3} [subsumption resolution 380,248]
383. ~0 | 3 [avatar contradiction clause 382]
384. $false [avatar sat refutation 249,250,368,383]
% SZS output end Proof for query_05_query_07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
