% Refutation found. Thanks to Tanya!
% SZS status Theorem for query06_query09_tptp4X
% SZS output start Proof for query06_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (query_2(X5) <=> select_3(X5)) [input]
4. ! [X6] : (from_5(X6) <=> query_2(X6)) [input]
5. ! [X6] : (select_1(X6) <=> ($true & from_5(X6))) [input]
6. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
7. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
8. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
9. ! [X7] : (query_7(X7) <=> select_8(X7)) [input]
10. ! [X8] : (from_10(X8) <=> query_7(X8)) [input]
11. ! [X8] : (select_6(X8) <=> ($true & from_10(X8))) [input]
12. ! [X7] : (main_query_2(X7) <=> select_6(X7)) [input]
56. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
57. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 56]
58. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 2]
59. ! [X0] : (query_2(X0) <=> select_3(X0)) [rectify 3]
60. ! [X0] : (from_5(X0) <=> query_2(X0)) [rectify 4]
61. ! [X0] : (select_1(X0) <=> ($true & from_5(X0))) [rectify 5]
62. ! [X0] : (select_1(X0) <=> from_5(X0)) [true and false elimination 61]
63. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 6]
64. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 8]
65. ! [X0] : (query_7(X0) <=> select_8(X0)) [rectify 9]
66. ! [X0] : (from_10(X0) <=> query_7(X0)) [rectify 10]
67. ! [X0] : (select_6(X0) <=> ($true & from_10(X0))) [rectify 11]
68. ! [X0] : (select_6(X0) <=> from_10(X0)) [true and false elimination 67]
69. ! [X0] : (main_query_2(X0) <=> select_6(X0)) [rectify 12]
113. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 57]
185. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 113]
186. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 1]
187. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 58]
188. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 187]
189. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
190. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0))) [skolemisation 188,189]
191. ! [X0] : ((query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~query_2(X0))) [nnf transformation 59]
192. ! [X0] : ((from_5(X0) | ~query_2(X0)) & (query_2(X0) | ~from_5(X0))) [nnf transformation 60]
193. ! [X0] : ((select_1(X0) | ~from_5(X0)) & (from_5(X0) | ~select_1(X0))) [nnf transformation 62]
194. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 63]
195. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 7]
196. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 64]
197. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 196]
198. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
199. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_8(X0))) [skolemisation 197,198]
200. ! [X0] : ((query_7(X0) | ~select_8(X0)) & (select_8(X0) | ~query_7(X0))) [nnf transformation 65]
201. ! [X0] : ((from_10(X0) | ~query_7(X0)) & (query_7(X0) | ~from_10(X0))) [nnf transformation 66]
202. ! [X0] : ((select_6(X0) | ~from_10(X0)) & (from_10(X0) | ~select_6(X0))) [nnf transformation 68]
203. ! [X0] : ((main_query_2(X0) | ~select_6(X0)) & (select_6(X0) | ~main_query_2(X0))) [nnf transformation 69]
206. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 185]
207. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
208. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 206,207]
209. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 186]
210. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 186]
211. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [cnf transformation 190]
212. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 190]
213. ~query_2(X0) | select_3(X0) [cnf transformation 191]
214. ~select_3(X0) | query_2(X0) [cnf transformation 191]
215. ~from_5(X0) | query_2(X0) [cnf transformation 192]
216. ~query_2(X0) | from_5(X0) [cnf transformation 192]
217. ~select_1(X0) | from_5(X0) [cnf transformation 193]
218. ~from_5(X0) | select_1(X0) [cnf transformation 193]
219. ~main_query_1(X0) | select_1(X0) [cnf transformation 194]
220. ~select_1(X0) | main_query_1(X0) [cnf transformation 194]
221. ~from_9(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 195]
222. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 195]
223. from_9(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_8(X0) [cnf transformation 199]
224. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 199]
225. ~query_7(X0) | select_8(X0) [cnf transformation 200]
226. ~select_8(X0) | query_7(X0) [cnf transformation 200]
227. ~from_10(X0) | query_7(X0) [cnf transformation 201]
228. ~query_7(X0) | from_10(X0) [cnf transformation 201]
229. ~select_6(X0) | from_10(X0) [cnf transformation 202]
230. ~from_10(X0) | select_6(X0) [cnf transformation 202]
231. ~main_query_2(X0) | select_6(X0) [cnf transformation 203]
232. ~select_6(X0) | main_query_2(X0) [cnf transformation 203]
274. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 208]
275. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 208]
310. 0 <=> main_query_1(sK8) [avatar definition]
311. main_query_1(sK8) <- {0} [avatar component clause 310]
313. 1 <=> ~main_query_1(sK8) [avatar definition]
314. ~main_query_1(sK8) <- {1} [avatar component clause 313]
316. 2 <=> main_query_2(sK8) [avatar definition]
317. main_query_2(sK8) <- {2} [avatar component clause 316]
319. 3 <=> ~main_query_2(sK8) [avatar definition]
320. ~main_query_2(sK8) <- {3} [avatar component clause 319]
321. ~1 | ~3 [avatar split clause 275,319,313]
322. 0 | 2 [avatar split clause 274,316,310]
323. select_6(sK8) <- {2} [resolution 231,317]
325. from_10(sK8) <- {2} [resolution 323,229]
328. query_7(sK8) <- {2} [resolution 325,227]
331. select_8(sK8) <- {2} [resolution 328,225]
337. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [resolution 211,209]
339. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_8(X0) [resolution 223,221]
341. from_9(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [resolution 337,222]
344. from_4(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_8(X1) [resolution 339,210]
346. ~select_3(X1) | select_8(X1) [resolution 341,224]
348. ~select_8(X1) | select_3(X1) [resolution 344,212]
349. select_3(sK8) <- {2} [resolution 348,331]
351. query_2(sK8) <- {2} [resolution 349,214]
352. from_5(sK8) <- {2} [resolution 351,216]
354. select_1(sK8) <- {2} [resolution 352,218]
356. main_query_1(sK8) <- {2} [resolution 354,220]
358. $false <- {1, 2} [subsumption resolution 356,314]
359. 1 | ~2 [avatar contradiction clause 358]
360. select_1(sK8) <- {0} [resolution 311,219]
362. from_5(sK8) <- {0} [resolution 360,217]
364. query_2(sK8) <- {0} [resolution 362,215]
366. select_3(sK8) <- {0} [resolution 364,213]
367. select_8(sK8) <- {0} [resolution 366,346]
370. query_7(sK8) <- {0} [resolution 367,226]
371. from_10(sK8) <- {0} [resolution 370,228]
373. select_6(sK8) <- {0} [resolution 371,230]
375. main_query_2(sK8) <- {0} [resolution 373,232]
377. $false <- {0, 3} [subsumption resolution 375,320]
378. ~0 | 3 [avatar contradiction clause 377]
379. $false [avatar sat refutation 321,322,359,378]
% SZS output end Proof for query06_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
