% Refutation found. Thanks to Tanya!
% SZS status Theorem for query03_query05_tptp4X
% SZS output start Proof for query03_query05_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_4(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (where_5(X6) <=> eq(X6,X6)) [input]
6. ! [X6] : (select_3(X6) <=> ? [X7,X8,X9,X10] : (where_5(X6) & from_4(X6,X7,X8,X9,X10))) [input]
7. ! [X11] : (main_query_2(X11) <=> select_3(X11)) [input]
8. ! [X12] : eq(X12,X12) [input]
46. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
47. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 46]
48. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
50. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
51. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [rectify 5]
52. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
53. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
54. ! [X0] : eq(X0,X0) [rectify 8]
92. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 47]
154. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 92]
155. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
156. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 48]
157. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 156]
158. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
159. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 157,158]
160. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 49]
161. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 50]
162. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 51]
163. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 52]
164. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 163]
165. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
166. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 164,165]
167. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 53]
170. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 154]
171. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
172. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 170,171]
173. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 155]
174. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 155]
175. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 159]
176. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 159]
177. ~main_query_1(X0) | select_1(X0) [cnf transformation 160]
178. ~select_1(X0) | main_query_1(X0) [cnf transformation 160]
179. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 161]
180. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 161]
182. ~eq(X0,X0) | where_5(X0) [cnf transformation 162]
183. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 166]
185. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 166]
186. ~main_query_2(X0) | select_3(X0) [cnf transformation 167]
187. ~select_3(X0) | main_query_2(X0) [cnf transformation 167]
188. eq(X0,X0) [cnf transformation 54]
224. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 172]
225. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 172]
255. 0 <=> main_query_1(sK8) [avatar definition]
256. main_query_1(sK8) <- {0} [avatar component clause 255]
258. 1 <=> ~main_query_1(sK8) [avatar definition]
259. ~main_query_1(sK8) <- {1} [avatar component clause 258]
261. 2 <=> main_query_2(sK8) [avatar definition]
262. main_query_2(sK8) <- {2} [avatar component clause 261]
264. 3 <=> ~main_query_2(sK8) [avatar definition]
265. ~main_query_2(sK8) <- {3} [avatar component clause 264]
266. ~1 | ~3 [avatar split clause 225,264,258]
267. 0 | 2 [avatar split clause 224,261,255]
268. select_3(sK8) <- {2} [resolution 186,262]
272. where_5(X1) [resolution 182,188]
283. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 175,173]
285. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 183,179]
287. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 283,180]
290. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 285,174]
292. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 287,185]
293. ~select_1(X1) | select_3(X1) [subsumption resolution 292,272]
295. ~select_3(X1) | select_1(X1) [resolution 290,176]
296. select_1(sK8) <- {2} [resolution 295,268]
298. main_query_1(sK8) <- {2} [resolution 296,178]
299. $false <- {1, 2} [subsumption resolution 298,259]
300. 1 | ~2 [avatar contradiction clause 299]
301. select_1(sK8) <- {0} [resolution 256,177]
302. select_3(sK8) <- {0} [resolution 301,293]
305. main_query_2(sK8) <- {0} [resolution 302,187]
307. $false <- {0, 3} [subsumption resolution 305,265]
308. ~0 | 3 [avatar contradiction clause 307]
309. $false [avatar sat refutation 266,267,300,308]
% SZS output end Proof for query03_query05_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.013 s
% ------------------------------
% ------------------------------
