% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_01_query_06_tptp4X
% SZS output start Proof for query_01_query_06_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
6. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X5,X6] : (main_query_2(X5,X6) <=> select_3(X5,X6)) [input]
11. ! [X7] : lte(X7,X7) [input]
17. ! [X7,X8] : (gte(X7,X8) <=> lte(X8,X7)) [input]
46. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
47. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 46]
48. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 3]
50. ! [X0,X1] : (main_query_2(X0,X1) <=> select_3(X0,X1)) [rectify 7]
54. ! [X0] : lte(X0,X0) [rectify 11]
60. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 17]
89. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 47]
151. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 89]
152. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
153. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0,X1))) [nnf transformation 48]
154. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) | ~select_1(X0,X1))) [rectify 153]
155. ! [X1,X0] : (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) => from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) [choice axiom]
156. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1))) [skolemisation 154,155]
157. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 49]
158. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
159. ! [X0] : ((where_5(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [nnf transformation 5]
160. ! [X0] : ((where_5(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [flattening 159]
161. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0,X1))) [nnf transformation 6]
162. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) | ~select_3(X0,X1))) [rectify 161]
163. ! [X1,X0] : (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) => (where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
164. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_3(X0,X1))) [skolemisation 162,163]
165. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 50]
168. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 60]
169. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 151]
170. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
171. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 169,170]
172. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 152]
173. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 152]
174. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 156]
175. ~from_2(X0,X1,X2,X3,X4) | select_1(X0,X1) [cnf transformation 156]
176. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 157]
177. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 157]
178. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 158]
179. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 158]
182. ~gte(X0,X0) | where_5(X0) | ~lte(X0,X0) [cnf transformation 160]
183. from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [cnf transformation 164]
185. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0,X1) [cnf transformation 164]
186. ~main_query_2(X0,X1) | select_3(X0,X1) [cnf transformation 165]
187. ~select_3(X0,X1) | main_query_2(X0,X1) [cnf transformation 165]
191. lte(X0,X0) [cnf transformation 54]
197. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 168]
226. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 171]
227. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 171]
229. 0 <=> main_query_1(sK6,sK7) [avatar definition]
230. main_query_1(sK6,sK7) <- {0} [avatar component clause 229]
232. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
233. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 232]
235. 2 <=> main_query_2(sK6,sK7) [avatar definition]
236. main_query_2(sK6,sK7) <- {2} [avatar component clause 235]
238. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
239. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 238]
240. ~1 | ~3 [avatar split clause 227,238,232]
241. 0 | 2 [avatar split clause 226,235,229]
242. select_3(sK6,sK7) <- {2} [resolution 186,236]
250. gte(X1,X1) [resolution 197,191]
252. where_5(X0) | ~lte(X0,X0) [resolution 182,250]
254. where_5(X0) [subsumption resolution 252,191]
268. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 174,172]
270. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 183,178]
283. from_4(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 268,179]
289. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 270,173]
297. ~select_1(X2,X3) | ~where_5(X2) | select_3(X2,X3) [resolution 283,185]
303. ~select_1(X2,X3) | select_3(X2,X3) [subsumption resolution 297,254]
319. ~select_3(X2,X3) | select_1(X2,X3) [resolution 289,175]
332. select_1(sK6,sK7) <- {2} [resolution 319,242]
341. main_query_1(sK6,sK7) <- {2} [resolution 332,177]
342. $false <- {1, 2} [subsumption resolution 341,233]
343. 1 | ~2 [avatar contradiction clause 342]
344. select_1(sK6,sK7) <- {0} [resolution 230,176]
345. select_3(sK6,sK7) <- {0} [resolution 344,303]
348. main_query_2(sK6,sK7) <- {0} [resolution 345,187]
350. $false <- {0, 3} [subsumption resolution 348,239]
351. ~0 | 3 [avatar contradiction clause 350]
352. $false [avatar sat refutation 240,241,343,351]
% SZS output end Proof for query_01_query_06_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.011 s
% ------------------------------
% ------------------------------
