% Refutation found. Thanks to Tanya!
% SZS status Theorem for query02_query07_tptp4X
% SZS output start Proof for query02_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_4(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (where_5(X6) <=> (~gt(X6,X6) & ~lt(X6,X6))) [input]
6. ! [X6] : (select_3(X6) <=> ? [X7,X8,X9,X10] : (where_5(X6) & from_4(X6,X7,X8,X9,X10))) [input]
7. ! [X6] : (main_query_2(X6) <=> select_3(X6)) [input]
8. ! [X11] : eq(X11,X11) [input]
15. ! [X11,X12] : (lt(X11,X12) <=> (~eq(X11,X12) & lte(X11,X12))) [input]
16. ! [X11,X12] : (gt(X11,X12) <=> lt(X12,X11)) [input]
42. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
43. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 42]
44. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
45. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
46. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
47. ! [X0] : (where_5(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [rectify 5]
48. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
49. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
50. ! [X0] : eq(X0,X0) [rectify 8]
57. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 15]
58. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 16]
84. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 43]
136. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 84]
137. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 44]
139. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 138]
140. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
141. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 139,140]
142. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 45]
143. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 46]
144. ! [X0] : ((where_5(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [nnf transformation 47]
145. ! [X0] : ((where_5(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [flattening 144]
146. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 48]
147. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 146]
148. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
149. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 147,148]
150. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 49]
153. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 57]
154. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 153]
155. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 58]
156. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 136]
157. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
158. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 156,157]
159. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 137]
160. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 137]
161. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 141]
162. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 141]
163. ~main_query_1(X0) | select_1(X0) [cnf transformation 142]
164. ~select_1(X0) | main_query_1(X0) [cnf transformation 142]
165. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 143]
166. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 143]
169. gt(X0,X0) | where_5(X0) | lt(X0,X0) [cnf transformation 145]
170. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 149]
172. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 149]
173. ~main_query_2(X0) | select_3(X0) [cnf transformation 150]
174. ~select_3(X0) | main_query_2(X0) [cnf transformation 150]
175. eq(X0,X0) [cnf transformation 50]
185. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 154]
187. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 155]
212. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 158]
213. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 158]
215. 0 <=> main_query_1(sK8) [avatar definition]
216. main_query_1(sK8) <- {0} [avatar component clause 215]
218. 1 <=> ~main_query_1(sK8) [avatar definition]
219. ~main_query_1(sK8) <- {1} [avatar component clause 218]
221. 2 <=> main_query_2(sK8) [avatar definition]
222. main_query_2(sK8) <- {2} [avatar component clause 221]
224. 3 <=> ~main_query_2(sK8) [avatar definition]
225. ~main_query_2(sK8) <- {3} [avatar component clause 224]
226. ~1 | ~3 [avatar split clause 213,224,218]
227. 0 | 2 [avatar split clause 212,221,215]
228. select_3(sK8) <- {2} [resolution 173,222]
235. ~lt(X0,X0) [resolution 185,175]
237. where_5(X1) | lt(X1,X1) | lt(X1,X1) [resolution 169,187]
238. where_5(X1) | lt(X1,X1) [duplicate literal removal 237]
239. where_5(X1) [subsumption resolution 238,235]
303. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 161,159]
343. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 170,165]
426. from_4(X6,sK0(X6),sK1(X6),sK2(X6),sK3(X6)) | ~select_1(X6) [resolution 303,166]
434. from_2(X9,sK4(X9),sK5(X9),sK6(X9),sK7(X9)) | ~select_3(X9) [resolution 343,160]
438. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 426,172]
439. ~select_1(X1) | select_3(X1) [subsumption resolution 438,239]
441. ~select_3(X1) | select_1(X1) [resolution 434,162]
447. select_1(sK8) <- {2} [resolution 441,228]
449. main_query_1(sK8) <- {2} [resolution 447,164]
450. $false <- {1, 2} [subsumption resolution 449,219]
451. 1 | ~2 [avatar contradiction clause 450]
452. select_1(sK8) <- {0} [resolution 216,163]
453. select_3(sK8) <- {0} [resolution 452,439]
458. main_query_2(sK8) <- {0} [resolution 453,174]
460. $false <- {0, 3} [subsumption resolution 458,225]
461. ~0 | 3 [avatar contradiction clause 460]
462. $false [avatar sat refutation 226,227,451,461]
% SZS output end Proof for query02_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.022 s
% ------------------------------
% ------------------------------
