% Refutation found. Thanks to Tanya!
% SZS status Theorem for query05_query06_tptp4X
% SZS output start Proof for query05_query06_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X6,X7,X8,X9,X10] : (from_7(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
6. ! [X6] : (select_6(X6) <=> ? [X7,X8,X9,X10] : ($true & from_7(X6,X7,X8,X9,X10))) [input]
7. ! [X5] : (query_5(X5) <=> select_6(X5)) [input]
8. ! [X11] : (from_8(X11) <=> query_5(X11)) [input]
9. ! [X11] : (select_4(X11) <=> ($true & from_8(X11))) [input]
10. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
11. ! [X12] : eq(X12,X12) [input]
44. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
45. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 44]
46. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
47. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
48. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [rectify 6]
49. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 48]
50. ! [X0] : (query_5(X0) <=> select_6(X0)) [rectify 7]
51. ! [X0] : (from_8(X0) <=> query_5(X0)) [rectify 8]
52. ! [X0] : (select_4(X0) <=> ($true & from_8(X0))) [rectify 9]
53. ! [X0] : (select_4(X0) <=> from_8(X0)) [true and false elimination 52]
54. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 10]
55. ! [X0] : eq(X0,X0) [rectify 11]
88. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 45]
140. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 88]
141. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
142. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
143. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
144. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 143]
145. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
146. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 144,145]
147. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 46]
148. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 47]
149. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 49]
150. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 149]
151. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
152. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 150,151]
153. ! [X0] : ((query_5(X0) | ~select_6(X0)) & (select_6(X0) | ~query_5(X0))) [nnf transformation 50]
154. ! [X0] : ((from_8(X0) | ~query_5(X0)) & (query_5(X0) | ~from_8(X0))) [nnf transformation 51]
155. ! [X0] : ((select_4(X0) | ~from_8(X0)) & (from_8(X0) | ~select_4(X0))) [nnf transformation 53]
156. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 54]
159. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 140]
160. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
161. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 159,160]
162. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 141]
163. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 141]
165. ~eq(X0,X0) | where_3(X0) [cnf transformation 142]
166. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 146]
168. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 146]
169. ~main_query_1(X0) | select_1(X0) [cnf transformation 147]
170. ~select_1(X0) | main_query_1(X0) [cnf transformation 147]
171. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 148]
172. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 148]
173. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 152]
174. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 152]
175. ~query_5(X0) | select_6(X0) [cnf transformation 153]
176. ~select_6(X0) | query_5(X0) [cnf transformation 153]
177. ~from_8(X0) | query_5(X0) [cnf transformation 154]
178. ~query_5(X0) | from_8(X0) [cnf transformation 154]
179. ~select_4(X0) | from_8(X0) [cnf transformation 155]
180. ~from_8(X0) | select_4(X0) [cnf transformation 155]
181. ~main_query_2(X0) | select_4(X0) [cnf transformation 156]
182. ~select_4(X0) | main_query_2(X0) [cnf transformation 156]
183. eq(X0,X0) [cnf transformation 55]
214. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 161]
215. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 161]
240. 0 <=> main_query_1(sK8) [avatar definition]
241. main_query_1(sK8) <- {0} [avatar component clause 240]
243. 1 <=> ~main_query_1(sK8) [avatar definition]
244. ~main_query_1(sK8) <- {1} [avatar component clause 243]
246. 2 <=> main_query_2(sK8) [avatar definition]
247. main_query_2(sK8) <- {2} [avatar component clause 246]
249. 3 <=> ~main_query_2(sK8) [avatar definition]
250. ~main_query_2(sK8) <- {3} [avatar component clause 249]
251. ~1 | ~3 [avatar split clause 215,249,243]
252. 0 | 2 [avatar split clause 214,246,240]
253. select_4(sK8) <- {2} [resolution 181,247]
254. from_8(sK8) <- {2} [resolution 253,179]
257. query_5(sK8) <- {2} [resolution 254,177]
259. select_6(sK8) <- {2} [resolution 257,175]
261. where_3(X1) [resolution 165,183]
273. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 166,162]
275. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 173,171]
277. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 273,172]
280. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 275,163]
282. ~select_1(X1) | select_6(X1) [resolution 277,174]
284. ~select_6(X1) | ~where_3(X1) | select_1(X1) [resolution 280,168]
285. ~select_6(X1) | select_1(X1) [subsumption resolution 284,261]
286. select_1(sK8) <- {2} [resolution 285,259]
288. main_query_1(sK8) <- {2} [resolution 286,170]
290. $false <- {1, 2} [subsumption resolution 288,244]
291. 1 | ~2 [avatar contradiction clause 290]
292. select_1(sK8) <- {0} [resolution 241,169]
293. select_6(sK8) <- {0} [resolution 292,282]
297. query_5(sK8) <- {0} [resolution 293,176]
298. from_8(sK8) <- {0} [resolution 297,178]
300. select_4(sK8) <- {0} [resolution 298,180]
302. main_query_2(sK8) <- {0} [resolution 300,182]
304. $false <- {0, 3} [subsumption resolution 302,250]
305. ~0 | 3 [avatar contradiction clause 304]
306. $false [avatar sat refutation 251,252,291,305]
% SZS output end Proof for query05_query06_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.014 s
% ------------------------------
% ------------------------------
