% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_18_query_19_tptp4X
% SZS output start Proof for query_18_query_19_tptp4X
1. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) <=> (tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4))) [input]
2. ! [X0,X5] : (where_3(X0,X5) <=> lt(X0,X5)) [input]
3. ! [X1,X6] : (select_1(X1,X6) <=> ? [X0,X2,X3,X4,X5,X7,X8,X9] : (where_3(X0,X5) & from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [input]
4. ! [X10,X11] : (main_query_1(X10,X11) <=> select_1(X10,X11)) [input]
5. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) <=> (tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4))) [input]
6. ! [X0,X5] : (where_6(X0,X5) <=> (neq(X0,X5) & lte(X0,X5))) [input]
7. ! [X1,X6] : (select_4(X1,X6) <=> ? [X0,X2,X3,X4,X5,X7,X8,X9] : (where_6(X0,X5) & from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [input]
8. ! [X10,X11] : (main_query_2(X10,X11) <=> select_4(X10,X11)) [input]
15. ! [X12,X13] : (lt(X12,X13) <=> (~eq(X12,X13) & lte(X12,X13))) [input]
17. ! [X12,X13] : (neq(X12,X13) <=> ~eq(X12,X13)) [input]
61. ! [X10,X11] : (main_query_1(X10,X11) <=> main_query_2(X10,X11)) [input]
62. ~! [X10,X11] : (main_query_1(X10,X11) <=> main_query_2(X10,X11)) [negated conjecture 61]
63. ! [X0,X1] : (where_3(X0,X1) <=> lt(X0,X1)) [rectify 2]
64. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_3(X2,X6) & from_2(X2,X0,X3,X4,X5,X6,X1,X7,X8,X9))) [rectify 3]
65. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
66. ! [X0,X1] : (where_6(X0,X1) <=> (neq(X0,X1) & lte(X0,X1))) [rectify 6]
67. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_6(X2,X6) & from_5(X2,X0,X3,X4,X5,X6,X1,X7,X8,X9))) [rectify 7]
68. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
75. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 15]
77. ! [X0,X1] : (neq(X0,X1) <=> ~eq(X0,X1)) [rectify 17]
121. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 62]
211. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 121]
212. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | (~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4))) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [nnf transformation 1]
213. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4)) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [flattening 212]
214. ! [X0,X1] : ((where_3(X0,X1) | ~lt(X0,X1)) & (lt(X0,X1) | ~where_3(X0,X1))) [nnf transformation 63]
215. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_3(X2,X6) | ~from_2(X2,X0,X3,X4,X5,X6,X1,X7,X8,X9))) & (? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_3(X2,X6) & from_2(X2,X0,X3,X4,X5,X6,X1,X7,X8,X9)) | ~select_1(X0,X1))) [nnf transformation 64]
216. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_3(X2,X6) | ~from_2(X2,X0,X3,X4,X5,X6,X1,X7,X8,X9))) & (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_3(X10,X14) & from_2(X10,X0,X11,X12,X13,X14,X1,X15,X16,X17)) | ~select_1(X0,X1))) [rectify 215]
217. ! [X1,X0] : (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_3(X10,X14) & from_2(X10,X0,X11,X12,X13,X14,X1,X15,X16,X17)) => (where_3(sK0(X0,X1),sK4(X0,X1)) & from_2(sK0(X0,X1),X0,sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1),X1,sK5(X0,X1),sK6(X0,X1),sK7(X0,X1)))) [choice axiom]
218. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_3(X2,X6) | ~from_2(X2,X0,X3,X4,X5,X6,X1,X7,X8,X9))) & ((where_3(sK0(X0,X1),sK4(X0,X1)) & from_2(sK0(X0,X1),X0,sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1),X1,sK5(X0,X1),sK6(X0,X1),sK7(X0,X1))) | ~select_1(X0,X1))) [skolemisation 216,217]
219. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 65]
220. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | (~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4))) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [nnf transformation 5]
221. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4)) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [flattening 220]
222. ! [X0,X1] : ((where_6(X0,X1) | (~neq(X0,X1) | ~lte(X0,X1))) & ((neq(X0,X1) & lte(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 66]
223. ! [X0,X1] : ((where_6(X0,X1) | ~neq(X0,X1) | ~lte(X0,X1)) & ((neq(X0,X1) & lte(X0,X1)) | ~where_6(X0,X1))) [flattening 222]
224. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_6(X2,X6) | ~from_5(X2,X0,X3,X4,X5,X6,X1,X7,X8,X9))) & (? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_6(X2,X6) & from_5(X2,X0,X3,X4,X5,X6,X1,X7,X8,X9)) | ~select_4(X0,X1))) [nnf transformation 67]
225. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_6(X2,X6) | ~from_5(X2,X0,X3,X4,X5,X6,X1,X7,X8,X9))) & (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_6(X10,X14) & from_5(X10,X0,X11,X12,X13,X14,X1,X15,X16,X17)) | ~select_4(X0,X1))) [rectify 224]
226. ! [X1,X0] : (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_6(X10,X14) & from_5(X10,X0,X11,X12,X13,X14,X1,X15,X16,X17)) => (where_6(sK8(X0,X1),sK12(X0,X1)) & from_5(sK8(X0,X1),X0,sK9(X0,X1),sK10(X0,X1),sK11(X0,X1),sK12(X0,X1),X1,sK13(X0,X1),sK14(X0,X1),sK15(X0,X1)))) [choice axiom]
227. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_6(X2,X6) | ~from_5(X2,X0,X3,X4,X5,X6,X1,X7,X8,X9))) & ((where_6(sK8(X0,X1),sK12(X0,X1)) & from_5(sK8(X0,X1),X0,sK9(X0,X1),sK10(X0,X1),sK11(X0,X1),sK12(X0,X1),X1,sK13(X0,X1),sK14(X0,X1),sK15(X0,X1))) | ~select_4(X0,X1))) [skolemisation 225,226]
228. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 68]
231. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 75]
232. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 231]
233. ! [X0,X1] : ((neq(X0,X1) | eq(X0,X1)) & (~eq(X0,X1) | ~neq(X0,X1))) [nnf transformation 77]
234. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 211]
235. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK16,sK17) | ~main_query_1(sK16,sK17)) & (main_query_2(sK16,sK17) | main_query_1(sK16,sK17))) [choice axiom]
236. (~main_query_2(sK16,sK17) | ~main_query_1(sK16,sK17)) & (main_query_2(sK16,sK17) | main_query_1(sK16,sK17)) [skolemisation 234,235]
237. ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | tab1(X0,X1,X2,X3,X4) [cnf transformation 213]
238. ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | tab2(X5,X6,X7,X8,X9) [cnf transformation 213]
239. from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4) [cnf transformation 213]
240. ~where_3(X0,X1) | lt(X0,X1) [cnf transformation 214]
241. ~lt(X0,X1) | where_3(X0,X1) [cnf transformation 214]
242. from_2(sK0(X0,X1),X0,sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1),X1,sK5(X0,X1),sK6(X0,X1),sK7(X0,X1)) | ~select_1(X0,X1) [cnf transformation 218]
243. where_3(sK0(X0,X1),sK4(X0,X1)) | ~select_1(X0,X1) [cnf transformation 218]
244. ~from_2(X2,X0,X3,X4,X5,X6,X1,X7,X8,X9) | ~where_3(X2,X6) | select_1(X0,X1) [cnf transformation 218]
245. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 219]
246. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 219]
247. ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | tab1(X0,X1,X2,X3,X4) [cnf transformation 221]
248. ~from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | tab2(X5,X6,X7,X8,X9) [cnf transformation 221]
249. from_5(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4) [cnf transformation 221]
250. ~where_6(X0,X1) | lte(X0,X1) [cnf transformation 223]
251. ~where_6(X0,X1) | neq(X0,X1) [cnf transformation 223]
252. ~neq(X0,X1) | where_6(X0,X1) | ~lte(X0,X1) [cnf transformation 223]
253. from_5(sK8(X0,X1),X0,sK9(X0,X1),sK10(X0,X1),sK11(X0,X1),sK12(X0,X1),X1,sK13(X0,X1),sK14(X0,X1),sK15(X0,X1)) | ~select_4(X0,X1) [cnf transformation 227]
254. where_6(sK8(X0,X1),sK12(X0,X1)) | ~select_4(X0,X1) [cnf transformation 227]
255. ~from_5(X2,X0,X3,X4,X5,X6,X1,X7,X8,X9) | ~where_6(X2,X6) | select_4(X0,X1) [cnf transformation 227]
256. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 228]
257. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 228]
266. ~lt(X0,X1) | lte(X0,X1) [cnf transformation 232]
267. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 232]
268. ~lte(X0,X1) | eq(X0,X1) | lt(X0,X1) [cnf transformation 232]
269. ~eq(X0,X1) | ~neq(X0,X1) [cnf transformation 233]
270. eq(X0,X1) | neq(X0,X1) [cnf transformation 233]
313. main_query_2(sK16,sK17) | main_query_1(sK16,sK17) [cnf transformation 236]
314. ~main_query_2(sK16,sK17) | ~main_query_1(sK16,sK17) [cnf transformation 236]
316. 0 <=> main_query_1(sK16,sK17) [avatar definition]
317. main_query_1(sK16,sK17) <- {0} [avatar component clause 316]
319. 1 <=> ~main_query_1(sK16,sK17) [avatar definition]
320. ~main_query_1(sK16,sK17) <- {1} [avatar component clause 319]
322. 2 <=> main_query_2(sK16,sK17) [avatar definition]
323. main_query_2(sK16,sK17) <- {2} [avatar component clause 322]
325. 3 <=> ~main_query_2(sK16,sK17) [avatar definition]
326. ~main_query_2(sK16,sK17) <- {3} [avatar component clause 325]
327. ~1 | ~3 [avatar split clause 314,325,319]
328. 0 | 2 [avatar split clause 313,322,316]
329. select_4(sK16,sK17) <- {2} [resolution 256,323]
337. ~lt(X2,X3) | neq(X2,X3) [resolution 270,267]
394. lt(sK0(X0,X1),sK4(X0,X1)) | ~select_1(X0,X1) [resolution 243,240]
398. neq(sK8(X0,X1),sK12(X0,X1)) | ~select_4(X0,X1) [resolution 254,251]
399. lte(sK8(X2,X3),sK12(X2,X3)) | ~select_4(X2,X3) [resolution 254,250]
421. ~tab2(X0,X1,X2,X3,X4) | ~tab1(X5,X6,X7,X8,X9) | ~where_3(X5,X0) | select_1(X6,X1) [resolution 239,244]
427. ~tab2(X0,X1,X2,X3,X4) | ~tab1(X5,X6,X7,X8,X9) | ~where_6(X5,X0) | select_4(X6,X1) [resolution 249,255]
569. tab1(sK0(X0,X1),X0,sK1(X0,X1),sK2(X0,X1),sK3(X0,X1)) | ~select_1(X0,X1) [resolution 242,237]
570. tab2(sK4(X2,X3),X3,sK5(X2,X3),sK6(X2,X3),sK7(X2,X3)) | ~select_1(X2,X3) [resolution 242,238]
594. tab1(sK8(X0,X1),X0,sK9(X0,X1),sK10(X0,X1),sK11(X0,X1)) | ~select_4(X0,X1) [resolution 253,247]
595. tab2(sK12(X2,X3),X3,sK13(X2,X3),sK14(X2,X3),sK15(X2,X3)) | ~select_4(X2,X3) [resolution 253,248]
613. neq(sK0(X2,X3),sK4(X2,X3)) | ~select_1(X2,X3) [resolution 394,337]
614. lte(sK0(X4,X5),sK4(X4,X5)) | ~select_1(X4,X5) [resolution 394,266]
636. eq(sK8(X2,X3),sK12(X2,X3)) | ~select_4(X2,X3) | lt(sK8(X2,X3),sK12(X2,X3)) [resolution 399,268]
747. ~tab1(X7,X8,X9,X10,X11) | ~where_3(X7,sK12(X12,X13)) | select_1(X8,X13) | ~select_4(X12,X13) [resolution 421,595]
761. ~tab1(X0,X1,X2,X3,X4) | ~where_6(X0,sK4(X5,X6)) | select_4(X1,X6) | ~select_1(X5,X6) [resolution 427,570]
1248. ~select_1(X14,X15) | where_6(sK0(X14,X15),sK4(X14,X15)) | ~lte(sK0(X14,X15),sK4(X14,X15)) [resolution 613,252]
1249. where_6(sK0(X14,X15),sK4(X14,X15)) | ~select_1(X14,X15) [subsumption resolution 1248,614]
1684. ~select_4(X11,X12) | lt(sK8(X11,X12),sK12(X11,X12)) | ~neq(sK8(X11,X12),sK12(X11,X12)) [resolution 636,269]
1700. lt(sK8(X11,X12),sK12(X11,X12)) | ~select_4(X11,X12) [subsumption resolution 1684,398]
1701. where_3(sK8(X0,X1),sK12(X0,X1)) | ~select_4(X0,X1) [resolution 1700,241]
1730. ~where_3(sK8(X4,X5),sK12(X6,X7)) | select_1(X4,X7) | ~select_4(X6,X7) | ~select_4(X4,X5) [resolution 747,594]
1741. ~where_6(sK0(X0,X1),sK4(X2,X3)) | select_4(X0,X3) | ~select_1(X2,X3) | ~select_1(X0,X1) [resolution 761,569]
8743. select_1(X0,X1) | ~select_4(X0,X1) | ~select_4(X0,X1) | ~select_4(X0,X1) [resolution 1730,1701]
8744. ~select_4(X0,X1) | select_1(X0,X1) [duplicate literal removal 8743]
8745. select_1(sK16,sK17) <- {2} [resolution 8744,329]
8795. select_4(X0,X1) | ~select_1(X0,X1) | ~select_1(X0,X1) | ~select_1(X0,X1) [resolution 1741,1249]
8796. ~select_1(X0,X1) | select_4(X0,X1) [duplicate literal removal 8795]
8797. main_query_1(sK16,sK17) <- {2} [resolution 8745,246]
8810. $false <- {1, 2} [subsumption resolution 8797,320]
8811. 1 | ~2 [avatar contradiction clause 8810]
8812. select_1(sK16,sK17) <- {0} [resolution 317,245]
8926. select_4(sK16,sK17) <- {0} [resolution 8796,8812]
8938. main_query_2(sK16,sK17) <- {0} [resolution 8926,257]
8952. $false <- {0, 3} [subsumption resolution 8938,326]
8953. ~0 | 3 [avatar contradiction clause 8952]
8954. $false [avatar sat refutation 327,328,8811,8953]
% SZS output end Proof for query_18_query_19_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 9466
% Time elapsed: 1.063 s
% ------------------------------
% ------------------------------
