% Refutation found. Thanks to Tanya!
% SZS status Theorem for query08_query09_tptp4X
% SZS output start Proof for query08_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X5,X6,X7,X8,X9] : (from_8(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
6. ! [X5] : (select_7(X5) <=> ? [X6,X7,X8,X9] : ($true & from_8(X5,X6,X7,X8,X9))) [input]
7. ! [X5] : (subquery_from_6(X5) <=> select_7(X5)) [input]
8. ! [X5] : (from_5(X5) <=> subquery_from_6(X5)) [input]
9. ! [X5] : (select_4(X5) <=> ($true & from_5(X5))) [input]
10. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
14. ! [X10] : lte(X10,X10) [input]
20. ! [X10,X11] : (gte(X10,X11) <=> lte(X11,X10)) [input]
44. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
45. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 44]
46. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
47. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
48. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 47]
49. ! [X0] : (subquery_from_6(X0) <=> select_7(X0)) [rectify 7]
50. ! [X0] : (from_5(X0) <=> subquery_from_6(X0)) [rectify 8]
51. ! [X0] : (select_4(X0) <=> ($true & from_5(X0))) [rectify 9]
52. ! [X0] : (select_4(X0) <=> from_5(X0)) [true and false elimination 51]
53. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 10]
57. ! [X0] : lte(X0,X0) [rectify 14]
63. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 20]
138. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 45]
139. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
140. ! [X0] : ((where_3(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
141. ! [X0] : ((where_3(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_3(X0))) [flattening 140]
142. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
143. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 142]
144. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
145. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 143,144]
146. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
147. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 46]
148. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 48]
149. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 148]
150. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
151. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0))) [skolemisation 149,150]
152. ! [X0] : ((subquery_from_6(X0) | ~select_7(X0)) & (select_7(X0) | ~subquery_from_6(X0))) [nnf transformation 49]
153. ! [X0] : ((from_5(X0) | ~subquery_from_6(X0)) & (subquery_from_6(X0) | ~from_5(X0))) [nnf transformation 50]
154. ! [X0] : ((select_4(X0) | ~from_5(X0)) & (from_5(X0) | ~select_4(X0))) [nnf transformation 52]
155. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 53]
158. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 63]
159. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 138]
160. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
161. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 159,160]
162. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 139]
163. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 139]
166. ~gte(X0,X0) | where_3(X0) | ~lte(X0,X0) [cnf transformation 141]
167. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 145]
169. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 145]
170. ~main_query_1(X0) | select_1(X0) [cnf transformation 146]
171. ~select_1(X0) | main_query_1(X0) [cnf transformation 146]
172. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 147]
173. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 147]
174. from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [cnf transformation 151]
175. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 151]
176. ~subquery_from_6(X0) | select_7(X0) [cnf transformation 152]
177. ~select_7(X0) | subquery_from_6(X0) [cnf transformation 152]
178. ~from_5(X0) | subquery_from_6(X0) [cnf transformation 153]
179. ~subquery_from_6(X0) | from_5(X0) [cnf transformation 153]
180. ~select_4(X0) | from_5(X0) [cnf transformation 154]
181. ~from_5(X0) | select_4(X0) [cnf transformation 154]
182. ~main_query_2(X0) | select_4(X0) [cnf transformation 155]
183. ~select_4(X0) | main_query_2(X0) [cnf transformation 155]
187. lte(X0,X0) [cnf transformation 57]
193. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 158]
217. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 161]
218. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 161]
243. 0 <=> main_query_1(sK8) [avatar definition]
244. main_query_1(sK8) <- {0} [avatar component clause 243]
246. 1 <=> ~main_query_1(sK8) [avatar definition]
247. ~main_query_1(sK8) <- {1} [avatar component clause 246]
249. 2 <=> main_query_2(sK8) [avatar definition]
250. main_query_2(sK8) <- {2} [avatar component clause 249]
252. 3 <=> ~main_query_2(sK8) [avatar definition]
253. ~main_query_2(sK8) <- {3} [avatar component clause 252]
254. ~1 | ~3 [avatar split clause 218,252,246]
255. 0 | 2 [avatar split clause 217,249,243]
256. select_4(sK8) <- {2} [resolution 182,250]
257. from_5(sK8) <- {2} [resolution 256,180]
260. subquery_from_6(sK8) <- {2} [resolution 257,178]
262. select_7(sK8) <- {2} [resolution 260,176]
269. gte(X1,X1) [resolution 193,187]
271. where_3(X0) | ~lte(X0,X0) [resolution 166,269]
273. where_3(X0) [subsumption resolution 271,187]
279. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 167,162]
281. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [resolution 174,172]
283. from_8(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 279,173]
286. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_7(X1) [resolution 281,163]
288. ~select_1(X1) | select_7(X1) [resolution 283,175]
290. ~select_7(X1) | ~where_3(X1) | select_1(X1) [resolution 286,169]
291. ~select_7(X1) | select_1(X1) [subsumption resolution 290,273]
292. select_1(sK8) <- {2} [resolution 291,262]
294. main_query_1(sK8) <- {2} [resolution 292,171]
296. $false <- {1, 2} [subsumption resolution 294,247]
297. 1 | ~2 [avatar contradiction clause 296]
298. select_1(sK8) <- {0} [resolution 244,170]
299. select_7(sK8) <- {0} [resolution 298,288]
303. subquery_from_6(sK8) <- {0} [resolution 299,177]
304. from_5(sK8) <- {0} [resolution 303,179]
306. select_4(sK8) <- {0} [resolution 304,181]
308. main_query_2(sK8) <- {0} [resolution 306,183]
310. $false <- {0, 3} [subsumption resolution 308,253]
311. ~0 | 3 [avatar contradiction clause 310]
312. $false [avatar sat refutation 254,255,297,311]
% SZS output end Proof for query08_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
