% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query06_tptp4X
% SZS output start Proof for query00_query06_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X5,X6,X7,X8] : (from_7(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
5. ! [X0] : (select_6(X0) <=> ? [X5,X6,X7,X8] : ($true & from_7(X0,X5,X6,X7,X8))) [input]
6. ! [X0] : (subquery_from_5(X0) <=> select_6(X0)) [input]
7. ! [X0] : (from_4(X0) <=> subquery_from_5(X0)) [input]
8. ! [X0] : (select_3(X0) <=> ($true & from_4(X0))) [input]
9. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
43. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
44. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 43]
45. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
46. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
47. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [rectify 5]
48. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 47]
49. ! [X0] : (select_3(X0) <=> from_4(X0)) [true and false elimination 8]
134. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 44]
135. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 45]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 136]
138. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
139. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 137,138]
140. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
141. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 46]
142. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 48]
143. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 142]
144. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
145. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 143,144]
146. ! [X0] : ((subquery_from_5(X0) | ~select_6(X0)) & (select_6(X0) | ~subquery_from_5(X0))) [nnf transformation 6]
147. ! [X0] : ((from_4(X0) | ~subquery_from_5(X0)) & (subquery_from_5(X0) | ~from_4(X0))) [nnf transformation 7]
148. ! [X0] : ((select_3(X0) | ~from_4(X0)) & (from_4(X0) | ~select_3(X0))) [nnf transformation 49]
149. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 9]
152. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 134]
153. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
154. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 152,153]
155. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 135]
156. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 135]
157. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 139]
158. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 139]
159. ~main_query_1(X0) | select_1(X0) [cnf transformation 140]
160. ~select_1(X0) | main_query_1(X0) [cnf transformation 140]
161. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 141]
162. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 141]
163. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 145]
164. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 145]
165. ~subquery_from_5(X0) | select_6(X0) [cnf transformation 146]
166. ~select_6(X0) | subquery_from_5(X0) [cnf transformation 146]
167. ~from_4(X0) | subquery_from_5(X0) [cnf transformation 147]
168. ~subquery_from_5(X0) | from_4(X0) [cnf transformation 147]
169. ~select_3(X0) | from_4(X0) [cnf transformation 148]
170. ~from_4(X0) | select_3(X0) [cnf transformation 148]
171. ~main_query_2(X0) | select_3(X0) [cnf transformation 149]
172. ~select_3(X0) | main_query_2(X0) [cnf transformation 149]
204. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 154]
205. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 154]
230. 0 <=> main_query_1(sK8) [avatar definition]
231. main_query_1(sK8) <- {0} [avatar component clause 230]
233. 1 <=> ~main_query_1(sK8) [avatar definition]
234. ~main_query_1(sK8) <- {1} [avatar component clause 233]
236. 2 <=> main_query_2(sK8) [avatar definition]
237. main_query_2(sK8) <- {2} [avatar component clause 236]
239. 3 <=> ~main_query_2(sK8) [avatar definition]
240. ~main_query_2(sK8) <- {3} [avatar component clause 239]
241. ~1 | ~3 [avatar split clause 205,239,233]
242. 0 | 2 [avatar split clause 204,236,230]
243. select_3(sK8) <- {2} [resolution 171,237]
245. from_4(sK8) <- {2} [resolution 243,169]
248. subquery_from_5(sK8) <- {2} [resolution 245,167]
251. select_6(sK8) <- {2} [resolution 248,165]
257. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 157,155]
259. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 163,161]
261. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 257,162]
264. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 259,156]
266. ~select_1(X1) | select_6(X1) [resolution 261,164]
268. ~select_6(X1) | select_1(X1) [resolution 264,158]
269. select_1(sK8) <- {2} [resolution 268,251]
271. main_query_1(sK8) <- {2} [resolution 269,160]
272. $false <- {1, 2} [subsumption resolution 271,234]
273. 1 | ~2 [avatar contradiction clause 272]
274. select_1(sK8) <- {0} [resolution 231,159]
275. select_6(sK8) <- {0} [resolution 274,266]
278. subquery_from_5(sK8) <- {0} [resolution 275,166]
279. from_4(sK8) <- {0} [resolution 278,168]
281. select_3(sK8) <- {0} [resolution 279,170]
283. main_query_2(sK8) <- {0} [resolution 281,172]
285. $false <- {0, 3} [subsumption resolution 283,240]
286. ~0 | 3 [avatar contradiction clause 285]
287. $false [avatar sat refutation 241,242,273,286]
% SZS output end Proof for query00_query06_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
