% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_07_query_08_tptp4X
% SZS output start Proof for query_07_query_08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : ($true & from_5(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (subquery_from_3(X5) <=> select_4(X5)) [input]
4. ! [X6] : (from_2(X6) <=> subquery_from_3(X6)) [input]
5. ! [X6] : (select_1(X6) <=> ($true & from_2(X6))) [input]
6. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
7. ! [X7,X8,X9,X10,X11] : (from_7(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
8. ! [X7] : (where_8(X7) <=> (~gt(X7,X7) & ~lt(X7,X7))) [input]
9. ! [X7] : (select_6(X7) <=> ? [X8,X9,X10,X11] : (where_8(X7) & from_7(X7,X8,X9,X10,X11))) [input]
10. ! [X7] : (main_query_2(X7) <=> select_6(X7)) [input]
11. ! [X12] : eq(X12,X12) [input]
18. ! [X12,X13] : (lt(X12,X13) <=> (~eq(X12,X13) & lte(X12,X13))) [input]
19. ! [X12,X13] : (gt(X12,X13) <=> lt(X13,X12)) [input]
42. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
43. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 42]
44. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4)) [true and false elimination 2]
45. ! [X0] : (subquery_from_3(X0) <=> select_4(X0)) [rectify 3]
46. ! [X0] : (from_2(X0) <=> subquery_from_3(X0)) [rectify 4]
47. ! [X0] : (select_1(X0) <=> ($true & from_2(X0))) [rectify 5]
48. ! [X0] : (select_1(X0) <=> from_2(X0)) [true and false elimination 47]
49. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 6]
50. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 7]
51. ! [X0] : (where_8(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [rectify 8]
52. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4))) [rectify 9]
53. ! [X0] : (main_query_2(X0) <=> select_6(X0)) [rectify 10]
54. ! [X0] : eq(X0,X0) [rectify 11]
61. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 18]
62. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 19]
85. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 43]
131. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 85]
132. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 1]
133. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4) | ~select_4(X0))) [nnf transformation 44]
134. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) | ~select_4(X0))) [rectify 133]
135. ! [X0] : (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) => from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
136. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0))) [skolemisation 134,135]
137. ! [X0] : ((subquery_from_3(X0) | ~select_4(X0)) & (select_4(X0) | ~subquery_from_3(X0))) [nnf transformation 45]
138. ! [X0] : ((from_2(X0) | ~subquery_from_3(X0)) & (subquery_from_3(X0) | ~from_2(X0))) [nnf transformation 46]
139. ! [X0] : ((select_1(X0) | ~from_2(X0)) & (from_2(X0) | ~select_1(X0))) [nnf transformation 48]
140. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 49]
141. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 50]
142. ! [X0] : ((where_8(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_8(X0))) [nnf transformation 51]
143. ! [X0] : ((where_8(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_8(X0))) [flattening 142]
144. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4)) | ~select_6(X0))) [nnf transformation 52]
145. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) | ~select_6(X0))) [rectify 144]
146. ! [X0] : (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) => (where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
147. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & ((where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_6(X0))) [skolemisation 145,146]
148. ! [X0] : ((main_query_2(X0) | ~select_6(X0)) & (select_6(X0) | ~main_query_2(X0))) [nnf transformation 53]
151. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 61]
152. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 151]
153. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 62]
154. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 131]
155. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
156. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 154,155]
157. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 132]
158. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 132]
159. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [cnf transformation 136]
160. ~from_5(X0,X1,X2,X3,X4) | select_4(X0) [cnf transformation 136]
161. ~subquery_from_3(X0) | select_4(X0) [cnf transformation 137]
162. ~select_4(X0) | subquery_from_3(X0) [cnf transformation 137]
163. ~from_2(X0) | subquery_from_3(X0) [cnf transformation 138]
164. ~subquery_from_3(X0) | from_2(X0) [cnf transformation 138]
165. ~select_1(X0) | from_2(X0) [cnf transformation 139]
166. ~from_2(X0) | select_1(X0) [cnf transformation 139]
167. ~main_query_1(X0) | select_1(X0) [cnf transformation 140]
168. ~select_1(X0) | main_query_1(X0) [cnf transformation 140]
169. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 141]
170. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 141]
173. gt(X0,X0) | where_8(X0) | lt(X0,X0) [cnf transformation 143]
174. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 147]
176. ~from_7(X0,X1,X2,X3,X4) | ~where_8(X0) | select_6(X0) [cnf transformation 147]
177. ~main_query_2(X0) | select_6(X0) [cnf transformation 148]
178. ~select_6(X0) | main_query_2(X0) [cnf transformation 148]
179. eq(X0,X0) [cnf transformation 54]
189. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 152]
191. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 153]
213. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 156]
214. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 156]
216. 0 <=> main_query_1(sK8) [avatar definition]
217. main_query_1(sK8) <- {0} [avatar component clause 216]
219. 1 <=> ~main_query_1(sK8) [avatar definition]
220. ~main_query_1(sK8) <- {1} [avatar component clause 219]
222. 2 <=> main_query_2(sK8) [avatar definition]
223. main_query_2(sK8) <- {2} [avatar component clause 222]
225. 3 <=> ~main_query_2(sK8) [avatar definition]
226. ~main_query_2(sK8) <- {3} [avatar component clause 225]
227. ~1 | ~3 [avatar split clause 214,225,219]
228. 0 | 2 [avatar split clause 213,222,216]
229. select_6(sK8) <- {2} [resolution 177,223]
236. ~lt(X0,X0) [resolution 189,179]
238. where_8(X1) | lt(X1,X1) | lt(X1,X1) [resolution 173,191]
239. where_8(X1) | lt(X1,X1) [duplicate literal removal 238]
240. where_8(X1) [subsumption resolution 239,236]
264. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [resolution 159,157]
306. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 174,169]
426. from_7(X1,sK0(X1),sK1(X1),sK2(X1),sK3(X1)) | ~select_4(X1) [resolution 264,170]
433. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 306,158]
445. ~select_4(X7) | ~where_8(X7) | select_6(X7) [resolution 426,176]
446. ~select_4(X7) | select_6(X7) [subsumption resolution 445,240]
452. ~select_6(X1) | select_4(X1) [resolution 433,160]
458. select_4(sK8) <- {2} [resolution 452,229]
460. subquery_from_3(sK8) <- {2} [resolution 458,162]
461. from_2(sK8) <- {2} [resolution 460,164]
463. select_1(sK8) <- {2} [resolution 461,166]
465. main_query_1(sK8) <- {2} [resolution 463,168]
467. $false <- {1, 2} [subsumption resolution 465,220]
468. 1 | ~2 [avatar contradiction clause 467]
469. select_1(sK8) <- {0} [resolution 217,167]
471. from_2(sK8) <- {0} [resolution 469,165]
473. subquery_from_3(sK8) <- {0} [resolution 471,163]
482. select_4(sK8) <- {0} [resolution 473,161]
490. select_6(sK8) <- {0} [resolution 482,446]
500. main_query_2(sK8) <- {0} [resolution 490,178]
502. $false <- {0, 3} [subsumption resolution 500,226]
503. ~0 | 3 [avatar contradiction clause 502]
504. $false [avatar sat refutation 227,228,468,503]
% SZS output end Proof for query_07_query_08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.047 s
% ------------------------------
% ------------------------------
