% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_12_query_13_tptp4X
% SZS output start Proof for query_12_query_13_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8] : (where_6(X7,X8) <=> (lte(X8,X7) & lte(X7,X8))) [input]
7. ! [X7,X8] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X7,X8) & from_5(X7,X8,X9,X10,X11))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
12. ! [X12] : lte(X12,X12) [input]
13. ! [X12,X13] : ((lte(X13,X12) & lte(X12,X13)) <=> eq(X12,X13)) [input]
29. ! [X12,X7,X8] : ((where_6(X7,X8) & eq(X8,X12)) => where_6(X7,X12)) [input]
51. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
52. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 51]
53. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
54. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
55. ! [X0,X1] : (where_6(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [rectify 6]
56. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
57. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
61. ! [X0] : lte(X0,X0) [rectify 12]
62. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
78. ! [X0,X1,X2] : ((where_6(X1,X2) & eq(X2,X0)) => where_6(X1,X0)) [rectify 29]
100. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 52]
126. ! [X0,X1,X2] : (where_6(X1,X0) | (~where_6(X1,X2) | ~eq(X2,X0))) [ennf transformation 78]
127. ! [X0,X1,X2] : (where_6(X1,X0) | ~where_6(X1,X2) | ~eq(X2,X0)) [flattening 126]
170. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 100]
171. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
172. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
173. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
174. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 173]
175. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
176. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 174,175]
177. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 53]
178. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 54]
179. ! [X0,X1] : ((where_6(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 55]
180. ! [X0,X1] : ((where_6(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [flattening 179]
181. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 56]
182. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 181]
183. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
184. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 182,183]
185. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 57]
186. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 62]
187. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 186]
188. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 170]
189. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
190. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 188,189]
191. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 171]
192. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 171]
193. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 172]
194. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 172]
195. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 176]
196. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 176]
197. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 176]
198. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 177]
199. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 177]
200. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 178]
201. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 178]
202. ~where_6(X0,X1) | lte(X0,X1) [cnf transformation 180]
203. ~where_6(X0,X1) | lte(X1,X0) [cnf transformation 180]
204. ~lte(X1,X0) | where_6(X0,X1) | ~lte(X0,X1) [cnf transformation 180]
205. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 184]
206. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 184]
207. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 184]
208. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 185]
209. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 185]
213. lte(X0,X0) [cnf transformation 61]
214. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 187]
228. ~where_6(X1,X2) | where_6(X1,X0) | ~eq(X2,X0) [cnf transformation 127]
250. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 190]
251. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 190]
253. 0 <=> main_query_1(sK6,sK7) [avatar definition]
254. main_query_1(sK6,sK7) <- {0} [avatar component clause 253]
256. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
257. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 256]
259. 2 <=> main_query_2(sK6,sK7) [avatar definition]
260. main_query_2(sK6,sK7) <- {2} [avatar component clause 259]
262. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
263. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 262]
264. ~1 | ~3 [avatar split clause 251,262,256]
265. 0 | 2 [avatar split clause 250,259,253]
268. select_4(sK6,sK7) <- {2} [resolution 208,260]
269. where_6(sK6,sK7) <- {2} [resolution 268,206]
271. lte(sK7,sK6) <- {2} [resolution 269,203]
272. lte(sK6,sK7) <- {2} [resolution 269,202]
276. where_6(X0,X0) | ~lte(X0,X0) [resolution 204,213]
279. where_6(X0,X0) [subsumption resolution 276,213]
285. eq(sK6,sK7) | ~lte(sK6,sK7) <- {2} [resolution 214,271]
287. eq(sK6,sK7) <- {2} [subsumption resolution 285,272]
298. where_3(sK6,sK7) <- {2} [resolution 287,194]
323. ~eq(X1,X2) | where_6(X1,X2) [resolution 228,279]
361. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 195,191]
365. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 205,200]
420. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 361,201]
426. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 365,192]
434. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 420,207]
465. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 426,197]
478. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 465,268]
482. select_1(sK6,sK7) <- {2} [subsumption resolution 478,298]
495. main_query_1(sK6,sK7) <- {2} [resolution 482,199]
497. $false <- {1, 2} [subsumption resolution 495,257]
498. 1 | ~2 [avatar contradiction clause 497]
501. select_1(sK6,sK7) <- {0} [resolution 254,198]
505. where_3(sK6,sK7) <- {0} [resolution 501,196]
506. eq(sK6,sK7) <- {0} [resolution 505,193]
516. where_6(sK6,sK7) <- {0} [resolution 506,323]
581. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 516,434]
586. select_4(sK6,sK7) <- {0} [subsumption resolution 581,501]
622. main_query_2(sK6,sK7) <- {0} [resolution 586,209]
624. $false <- {0, 3} [subsumption resolution 622,263]
625. ~0 | 3 [avatar contradiction clause 624]
626. $false [avatar sat refutation 264,265,498,625]
% SZS output end Proof for query_12_query_13_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
