% Refutation found. Thanks to Tanya!
% SZS status Theorem for query11_query14_tptp4X
% SZS output start Proof for query11_query14_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8,X9,X10,X11] : (from_9(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
7. ! [X7] : (select_8(X7) <=> ? [X8,X9,X10,X11] : ($true & from_9(X7,X8,X9,X10,X11))) [input]
8. ! [X5] : (query_7(X5) <=> select_8(X5)) [input]
9. ! [X7,X8] : (where_6(X7,X8) <=> (query_7(X7) & eq(X7,X8))) [input]
10. ! [X7,X8] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X7,X8) & from_5(X7,X8,X9,X10,X11))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
62. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
63. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 62]
64. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
65. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
66. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 6]
67. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [rectify 7]
68. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 67]
69. ! [X0] : (query_7(X0) <=> select_8(X0)) [rectify 8]
70. ! [X0,X1] : (where_6(X0,X1) <=> (query_7(X0) & eq(X0,X1))) [rectify 9]
71. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 10]
72. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
123. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 63]
209. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 123]
210. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
211. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
212. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
213. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 212]
214. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
215. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 213,214]
216. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 64]
217. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 65]
218. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 66]
219. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 68]
220. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 219]
221. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
222. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 220,221]
223. ! [X0] : ((query_7(X0) | ~select_8(X0)) & (select_8(X0) | ~query_7(X0))) [nnf transformation 69]
224. ! [X0,X1] : ((where_6(X0,X1) | (~query_7(X0) | ~eq(X0,X1))) & ((query_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 70]
225. ! [X0,X1] : ((where_6(X0,X1) | ~query_7(X0) | ~eq(X0,X1)) & ((query_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 224]
226. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 71]
227. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 226]
228. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
229. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 227,228]
230. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 72]
233. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 209]
234. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
235. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 233,234]
236. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 210]
237. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 210]
238. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 211]
239. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 211]
240. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 215]
241. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 215]
242. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 215]
243. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 216]
244. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 216]
245. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 217]
246. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 217]
248. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 218]
250. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 222]
252. ~select_8(X0) | query_7(X0) [cnf transformation 223]
253. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 225]
255. ~eq(X0,X1) | ~query_7(X0) | where_6(X0,X1) [cnf transformation 225]
256. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 229]
257. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 229]
258. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 229]
259. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 230]
260. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 230]
309. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 235]
310. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 235]
352. 0 <=> main_query_1(sK10,sK11) [avatar definition]
353. main_query_1(sK10,sK11) <- {0} [avatar component clause 352]
355. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
356. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 355]
358. 2 <=> main_query_2(sK10,sK11) [avatar definition]
359. main_query_2(sK10,sK11) <- {2} [avatar component clause 358]
361. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
362. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 361]
363. ~1 | ~3 [avatar split clause 310,361,355]
364. 0 | 2 [avatar split clause 309,358,352]
367. select_4(sK10,sK11) <- {2} [resolution 259,359]
368. where_6(sK10,sK11) <- {2} [resolution 367,257]
370. eq(sK10,sK11) <- {2} [resolution 368,253]
380. where_3(sK10,sK11) <- {2} [resolution 370,239]
428. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 240,236]
432. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 256,245]
433. from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 428,248]
434. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 428,246]
440. from_2(X4,X5,sK7(X4,X5),sK8(X4,X5),sK9(X4,X5)) | ~select_4(X4,X5) [resolution 432,237]
442. ~select_1(X2,X3) | select_8(X2) [resolution 433,250]
443. ~where_6(X0,X1) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 434,258]
450. ~select_4(X0,X1) | ~where_3(X0,X1) | select_1(X0,X1) [resolution 440,242]
452. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 450,367]
453. select_1(sK10,sK11) <- {2} [subsumption resolution 452,380]
455. main_query_1(sK10,sK11) <- {2} [resolution 453,244]
457. $false <- {1, 2} [subsumption resolution 455,356]
458. 1 | ~2 [avatar contradiction clause 457]
459. select_1(sK10,sK11) <- {0} [resolution 353,243]
460. select_8(sK10) <- {0} [resolution 459,442]
462. where_3(sK10,sK11) <- {0} [resolution 459,241]
463. query_7(sK10) <- {0} [resolution 460,252]
464. eq(sK10,sK11) <- {0} [resolution 462,238]
467. ~query_7(sK10) | where_6(sK10,sK11) <- {0} [resolution 464,255]
472. where_6(sK10,sK11) <- {0} [subsumption resolution 467,463]
483. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {0} [resolution 472,443]
486. select_4(sK10,sK11) <- {0} [subsumption resolution 483,459]
492. main_query_2(sK10,sK11) <- {0} [resolution 486,260]
494. $false <- {0, 3} [subsumption resolution 492,362]
495. ~0 | 3 [avatar contradiction clause 494]
496. $false [avatar sat refutation 363,364,458,495]
% SZS output end Proof for query11_query14_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.026 s
% ------------------------------
% ------------------------------
