% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_03_query_06_tptp4X
% SZS output start Proof for query_03_query_06_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_4(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (where_5(X6) <=> eq(X6,X6)) [input]
6. ! [X6] : (select_3(X6) <=> ? [X7,X8,X9,X10] : (where_5(X6) & from_4(X6,X7,X8,X9,X10))) [input]
7. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
8. ! [X11] : eq(X11,X11) [input]
39. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
40. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 39]
41. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
42. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
43. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
44. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [rectify 5]
45. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
46. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
47. ! [X0] : eq(X0,X0) [rectify 8]
78. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 40]
124. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 78]
125. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
126. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 41]
127. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 126]
128. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
129. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 127,128]
130. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 42]
131. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 43]
132. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 44]
133. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 45]
134. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 133]
135. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
136. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 134,135]
137. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 46]
140. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 124]
141. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
142. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 140,141]
143. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 125]
144. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 125]
145. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 129]
146. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 129]
147. ~main_query_1(X0) | select_1(X0) [cnf transformation 130]
148. ~select_1(X0) | main_query_1(X0) [cnf transformation 130]
149. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 131]
150. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 131]
152. ~eq(X0,X0) | where_5(X0) [cnf transformation 132]
153. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 136]
155. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 136]
156. ~main_query_2(X0) | select_3(X0) [cnf transformation 137]
157. ~select_3(X0) | main_query_2(X0) [cnf transformation 137]
158. eq(X0,X0) [cnf transformation 47]
187. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 142]
188. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 142]
190. 0 <=> main_query_1(sK8) [avatar definition]
191. main_query_1(sK8) <- {0} [avatar component clause 190]
193. 1 <=> ~main_query_1(sK8) [avatar definition]
194. ~main_query_1(sK8) <- {1} [avatar component clause 193]
196. 2 <=> main_query_2(sK8) [avatar definition]
197. main_query_2(sK8) <- {2} [avatar component clause 196]
199. 3 <=> ~main_query_2(sK8) [avatar definition]
200. ~main_query_2(sK8) <- {3} [avatar component clause 199]
201. ~1 | ~3 [avatar split clause 188,199,193]
202. 0 | 2 [avatar split clause 187,196,190]
203. select_3(sK8) <- {2} [resolution 156,197]
207. where_5(X1) [resolution 152,158]
241. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 145,143]
243. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 153,149]
245. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 241,150]
248. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 243,144]
250. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 245,155]
251. ~select_1(X1) | select_3(X1) [subsumption resolution 250,207]
253. ~select_3(X1) | select_1(X1) [resolution 248,146]
256. select_1(sK8) <- {2} [resolution 253,203]
260. main_query_1(sK8) <- {2} [resolution 256,148]
261. $false <- {1, 2} [subsumption resolution 260,194]
262. 1 | ~2 [avatar contradiction clause 261]
263. select_1(sK8) <- {0} [resolution 191,147]
266. select_3(sK8) <- {0} [resolution 263,251]
271. main_query_2(sK8) <- {0} [resolution 266,157]
273. $false <- {0, 3} [subsumption resolution 271,200]
274. ~0 | 3 [avatar contradiction clause 273]
275. $false [avatar sat refutation 201,202,262,274]
% SZS output end Proof for query_03_query_06_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.043 s
% ------------------------------
% ------------------------------
