% Refutation found. Thanks to Tanya!
% SZS status Theorem for query12_query14_tptp4X
% SZS output start Proof for query12_query14_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (where_6_subquery_7(X5) <=> select_8(X5)) [input]
9. ! [X0,X1] : (where_6(X0,X1) <=> (where_6_subquery_7(X0) & eq(X0,X1))) [input]
10. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
13. ! [X7,X8] : (eq(X7,X8) => eq(X8,X7)) [input]
16. ! [X7,X8] : ((lte(X8,X7) & lte(X7,X8)) <=> eq(X7,X8)) [input]
26. ! [X7,X5,X6] : ((main_query_2(X5,X6) & eq(X5,X7)) => main_query_2(X7,X6)) [input]
63. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
64. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 63]
65. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
66. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 7]
67. ! [X0] : (where_6_subquery_7(X0) <=> select_8(X0)) [rectify 8]
68. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
70. ! [X0,X1] : (eq(X0,X1) => eq(X1,X0)) [rectify 13]
73. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 16]
83. ! [X0,X1,X2] : ((main_query_2(X1,X2) & eq(X1,X0)) => main_query_2(X0,X2)) [rectify 26]
120. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 64]
121. ! [X0,X1] : (eq(X1,X0) | ~eq(X0,X1)) [ennf transformation 70]
132. ! [X0,X1,X2] : (main_query_2(X0,X2) | (~main_query_2(X1,X2) | ~eq(X1,X0))) [ennf transformation 83]
133. ! [X0,X1,X2] : (main_query_2(X0,X2) | ~main_query_2(X1,X2) | ~eq(X1,X0)) [flattening 132]
206. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 120]
207. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
208. ! [X0,X1] : ((where_3(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [nnf transformation 2]
209. ! [X0,X1] : ((where_3(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [flattening 208]
210. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
211. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 210]
212. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
213. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 211,212]
214. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 65]
215. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
216. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
217. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 66]
218. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 217]
219. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
220. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 218,219]
221. ! [X0] : ((where_6_subquery_7(X0) | ~select_8(X0)) & (select_8(X0) | ~where_6_subquery_7(X0))) [nnf transformation 67]
222. ! [X0,X1] : ((where_6(X0,X1) | (~where_6_subquery_7(X0) | ~eq(X0,X1))) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 9]
223. ! [X0,X1] : ((where_6(X0,X1) | ~where_6_subquery_7(X0) | ~eq(X0,X1)) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 222]
224. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 10]
225. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 224]
226. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
227. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 225,226]
228. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 68]
229. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 73]
230. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 229]
231. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 206]
232. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
233. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 231,232]
234. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 207]
235. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 207]
236. ~where_3(X0,X1) | lte(X0,X1) [cnf transformation 209]
237. ~where_3(X0,X1) | lte(X1,X0) [cnf transformation 209]
238. ~lte(X1,X0) | where_3(X0,X1) | ~lte(X0,X1) [cnf transformation 209]
239. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 213]
240. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 213]
241. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 213]
242. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 214]
243. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 214]
244. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 215]
245. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 215]
247. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 216]
249. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 220]
251. ~select_8(X0) | where_6_subquery_7(X0) [cnf transformation 221]
252. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 223]
253. ~where_6(X0,X1) | where_6_subquery_7(X0) [cnf transformation 223]
254. ~eq(X0,X1) | ~where_6_subquery_7(X0) | where_6(X0,X1) [cnf transformation 223]
255. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 227]
256. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 227]
257. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 227]
258. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 228]
259. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 228]
261. ~eq(X0,X1) | eq(X1,X0) [cnf transformation 121]
264. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 230]
265. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 230]
266. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 230]
272. ~main_query_2(X1,X2) | main_query_2(X0,X2) | ~eq(X1,X0) [cnf transformation 133]
309. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 233]
310. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 233]
312. 0 <=> main_query_1(sK10,sK11) [avatar definition]
313. main_query_1(sK10,sK11) <- {0} [avatar component clause 312]
315. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
316. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 315]
318. 2 <=> main_query_2(sK10,sK11) [avatar definition]
319. main_query_2(sK10,sK11) <- {2} [avatar component clause 318]
321. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
322. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 321]
323. ~1 | ~3 [avatar split clause 310,321,315]
324. 0 | 2 [avatar split clause 309,318,312]
325. select_4(sK10,sK11) <- {2} [resolution 258,319]
327. where_6(sK10,sK11) <- {2} [resolution 325,256]
329. eq(sK10,sK11) <- {2} [resolution 327,252]
338. eq(sK11,sK10) <- {2} [resolution 329,261]
354. ~where_6_subquery_7(sK11) | where_6(sK11,sK10) <- {2} [resolution 338,254]
362. 4 <=> where_6(sK11,sK10) [avatar definition]
363. where_6(sK11,sK10) <- {4} [avatar component clause 362]
368. 7 <=> ~where_6_subquery_7(sK11) [avatar definition]
369. ~where_6_subquery_7(sK11) <- {7} [avatar component clause 368]
370. 4 | ~7 | ~2 [avatar split clause 354,318,368,362]
383. ~eq(sK10,X0) | main_query_2(X0,sK11) <- {2} [resolution 272,319]
392. main_query_2(sK11,sK11) <- {2} [resolution 383,329]
399. select_4(sK11,sK11) <- {2} [resolution 392,258]
403. where_6(sK11,sK11) <- {2} [resolution 399,256]
407. where_6_subquery_7(sK11) <- {2} [resolution 403,253]
408. $false <- {2, 7} [subsumption resolution 407,369]
409. ~2 | 7 [avatar contradiction clause 408]
410. select_1(sK10,sK11) <- {0} [resolution 313,242]
413. where_3(sK10,sK11) <- {0} [resolution 410,240]
415. lte(sK11,sK10) <- {0} [resolution 413,237]
416. lte(sK10,sK11) <- {0} [resolution 413,236]
419. eq(sK10,sK11) | ~lte(sK10,sK11) <- {0} [resolution 415,264]
421. eq(sK10,sK11) <- {0} [subsumption resolution 419,416]
429. ~where_6_subquery_7(sK10) | where_6(sK10,sK11) <- {0} [resolution 421,254]
445. 8 <=> where_6(sK10,sK11) [avatar definition]
446. where_6(sK10,sK11) <- {8} [avatar component clause 445]
448. 10 <=> where_6_subquery_7(sK10) [avatar definition]
451. 11 <=> ~where_6_subquery_7(sK10) [avatar definition]
453. 8 | ~11 | ~0 [avatar split clause 429,312,451,445]
490. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 239,234]
493. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 255,244]
684. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 490,245]
685. from_9(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 490,247]
701. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 493,235]
767. 14 <=> select_8(sK10) [avatar definition]
768. select_8(sK10) <- {14} [avatar component clause 767]
1402. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 684,257]
1410. ~select_1(X2,X3) | select_8(X2) [resolution 685,249]
1416. select_8(sK10) <- {0} [resolution 1410,410]
1431. 14 | ~0 [avatar split clause 1416,312,767]
1433. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 701,241]
1461. where_6_subquery_7(sK10) <- {14} [resolution 768,251]
1464. 10 | ~14 [avatar split clause 1461,767,448]
1485. eq(sK11,sK10) <- {4} [resolution 363,252]
1498. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {8} [resolution 446,1402]
1501. select_4(sK10,sK11) <- {0, 8} [subsumption resolution 1498,410]
1552. main_query_2(sK10,sK11) <- {0, 8} [resolution 1501,259]
1554. $false <- {0, 3, 8} [subsumption resolution 1552,322]
1555. ~0 | 3 | ~8 [avatar contradiction clause 1554]
1571. select_4(sK10,sK11) <- {2} [resolution 319,258]
1585. lte(sK11,sK10) <- {4} [resolution 1485,265]
1586. lte(sK10,sK11) <- {4} [resolution 1485,266]
1627. where_3(sK10,sK11) | ~lte(sK10,sK11) <- {4} [resolution 1585,238]
1632. where_3(sK10,sK11) <- {4} [subsumption resolution 1627,1586]
2015. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 1433,1571]
2019. select_1(sK10,sK11) <- {2, 4} [subsumption resolution 2015,1632]
2029. main_query_1(sK10,sK11) <- {2, 4} [resolution 2019,243]
2031. $false <- {1, 2, 4} [subsumption resolution 2029,316]
2032. 1 | ~2 | ~4 [avatar contradiction clause 2031]
2033. $false [avatar sat refutation 323,324,370,409,453,1431,1464,1555,2032]
% SZS output end Proof for query12_query14_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5628
% Time elapsed: 0.042 s
% ------------------------------
% ------------------------------
