% Refutation found. Thanks to Tanya!
% SZS status Theorem for query06_query07_tptp4X
% SZS output start Proof for query06_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : ($true & from_5(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (subquery_from_3(X0) <=> select_4(X0)) [input]
4. ! [X0] : (from_2(X0) <=> subquery_from_3(X0)) [input]
5. ! [X0] : (select_1(X0) <=> ($true & from_2(X0))) [input]
6. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
7. ! [X5,X6,X7,X8,X9] : (from_7(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
8. ! [X5] : (where_8(X5) <=> (~gt(X5,X5) & ~lt(X5,X5))) [input]
9. ! [X5] : (select_6(X5) <=> ? [X6,X7,X8,X9] : (where_8(X5) & from_7(X5,X6,X7,X8,X9))) [input]
10. ! [X5] : (main_query_2(X5) <=> select_6(X5)) [input]
11. ! [X10] : eq(X10,X10) [input]
17. ! [X10,X11] : (lt(X10,X11) <=> (~eq(X10,X11) & lte(X10,X11))) [input]
18. ! [X10,X11] : (gt(X10,X11) <=> lt(X11,X10)) [input]
44. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
45. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 44]
46. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4)) [true and false elimination 2]
47. ! [X0] : (select_1(X0) <=> from_2(X0)) [true and false elimination 5]
48. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 7]
49. ! [X0] : (where_8(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [rectify 8]
50. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4))) [rectify 9]
51. ! [X0] : (main_query_2(X0) <=> select_6(X0)) [rectify 10]
52. ! [X0] : eq(X0,X0) [rectify 11]
58. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 17]
59. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 18]
136. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 45]
137. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 1]
138. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4) | ~select_4(X0))) [nnf transformation 46]
139. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) | ~select_4(X0))) [rectify 138]
140. ! [X0] : (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) => from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
141. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0))) [skolemisation 139,140]
142. ! [X0] : ((subquery_from_3(X0) | ~select_4(X0)) & (select_4(X0) | ~subquery_from_3(X0))) [nnf transformation 3]
143. ! [X0] : ((from_2(X0) | ~subquery_from_3(X0)) & (subquery_from_3(X0) | ~from_2(X0))) [nnf transformation 4]
144. ! [X0] : ((select_1(X0) | ~from_2(X0)) & (from_2(X0) | ~select_1(X0))) [nnf transformation 47]
145. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 6]
146. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 48]
147. ! [X0] : ((where_8(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_8(X0))) [nnf transformation 49]
148. ! [X0] : ((where_8(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_8(X0))) [flattening 147]
149. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4)) | ~select_6(X0))) [nnf transformation 50]
150. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) | ~select_6(X0))) [rectify 149]
151. ! [X0] : (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) => (where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
152. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & ((where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_6(X0))) [skolemisation 150,151]
153. ! [X0] : ((main_query_2(X0) | ~select_6(X0)) & (select_6(X0) | ~main_query_2(X0))) [nnf transformation 51]
156. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 58]
157. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 156]
158. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 59]
159. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 136]
160. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
161. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 159,160]
162. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 137]
163. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 137]
164. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [cnf transformation 141]
165. ~from_5(X0,X1,X2,X3,X4) | select_4(X0) [cnf transformation 141]
166. ~subquery_from_3(X0) | select_4(X0) [cnf transformation 142]
167. ~select_4(X0) | subquery_from_3(X0) [cnf transformation 142]
168. ~from_2(X0) | subquery_from_3(X0) [cnf transformation 143]
169. ~subquery_from_3(X0) | from_2(X0) [cnf transformation 143]
170. ~select_1(X0) | from_2(X0) [cnf transformation 144]
171. ~from_2(X0) | select_1(X0) [cnf transformation 144]
172. ~main_query_1(X0) | select_1(X0) [cnf transformation 145]
173. ~select_1(X0) | main_query_1(X0) [cnf transformation 145]
174. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 146]
175. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 146]
178. gt(X0,X0) | where_8(X0) | lt(X0,X0) [cnf transformation 148]
179. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 152]
181. ~from_7(X0,X1,X2,X3,X4) | ~where_8(X0) | select_6(X0) [cnf transformation 152]
182. ~main_query_2(X0) | select_6(X0) [cnf transformation 153]
183. ~select_6(X0) | main_query_2(X0) [cnf transformation 153]
184. eq(X0,X0) [cnf transformation 52]
193. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 157]
195. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 158]
220. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 161]
221. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 161]
246. 0 <=> main_query_1(sK8) [avatar definition]
247. main_query_1(sK8) <- {0} [avatar component clause 246]
249. 1 <=> ~main_query_1(sK8) [avatar definition]
250. ~main_query_1(sK8) <- {1} [avatar component clause 249]
252. 2 <=> main_query_2(sK8) [avatar definition]
253. main_query_2(sK8) <- {2} [avatar component clause 252]
255. 3 <=> ~main_query_2(sK8) [avatar definition]
256. ~main_query_2(sK8) <- {3} [avatar component clause 255]
257. ~1 | ~3 [avatar split clause 221,255,249]
258. 0 | 2 [avatar split clause 220,252,246]
259. select_6(sK8) <- {2} [resolution 182,253]
265. ~lt(X0,X0) [resolution 193,184]
267. where_8(X1) | lt(X1,X1) | lt(X1,X1) [resolution 178,195]
268. where_8(X1) | lt(X1,X1) [duplicate literal removal 267]
269. where_8(X1) [subsumption resolution 268,265]
274. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [resolution 164,162]
276. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 179,174]
278. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [resolution 274,175]
281. from_5(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 276,163]
283. ~select_4(X1) | ~where_8(X1) | select_6(X1) [resolution 278,181]
284. ~select_4(X1) | select_6(X1) [subsumption resolution 283,269]
286. ~select_6(X1) | select_4(X1) [resolution 281,165]
287. select_4(sK8) <- {2} [resolution 286,259]
289. subquery_from_3(sK8) <- {2} [resolution 287,167]
290. from_2(sK8) <- {2} [resolution 289,169]
292. select_1(sK8) <- {2} [resolution 290,171]
294. main_query_1(sK8) <- {2} [resolution 292,173]
296. $false <- {1, 2} [subsumption resolution 294,250]
297. 1 | ~2 [avatar contradiction clause 296]
298. select_1(sK8) <- {0} [resolution 247,172]
300. from_2(sK8) <- {0} [resolution 298,170]
302. subquery_from_3(sK8) <- {0} [resolution 300,168]
304. select_4(sK8) <- {0} [resolution 302,166]
305. select_6(sK8) <- {0} [resolution 304,284]
308. main_query_2(sK8) <- {0} [resolution 305,183]
310. $false <- {0, 3} [subsumption resolution 308,256]
311. ~0 | 3 [avatar contradiction clause 310]
312. $false [avatar sat refutation 257,258,297,311]
% SZS output end Proof for query06_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.012 s
% ------------------------------
% ------------------------------
