% Refutation found. Thanks to Tanya!
% SZS status Theorem for query10_query12_tptp4X
% SZS output start Proof for query10_query12_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1] : (where_6(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [input]
7. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
13. ! [X7,X8] : ((lte(X8,X7) & lte(X7,X8)) <=> eq(X7,X8)) [input]
54. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
55. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 54]
56. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
57. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
62. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
103. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 55]
179. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 103]
180. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
181. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
182. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
183. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 182]
184. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
185. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 183,184]
186. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 56]
187. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
188. ! [X0,X1] : ((where_6(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 6]
189. ! [X0,X1] : ((where_6(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [flattening 188]
190. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 7]
191. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 190]
192. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
193. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 191,192]
194. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 57]
195. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 62]
196. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 195]
197. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 179]
198. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
199. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 197,198]
200. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 180]
201. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 180]
202. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 181]
203. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 181]
204. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 185]
205. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 185]
206. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 185]
207. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 186]
208. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 186]
209. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 187]
210. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 187]
211. ~where_6(X0,X1) | lte(X0,X1) [cnf transformation 189]
212. ~where_6(X0,X1) | lte(X1,X0) [cnf transformation 189]
213. ~lte(X1,X0) | where_6(X0,X1) | ~lte(X0,X1) [cnf transformation 189]
214. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 193]
215. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 193]
216. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 193]
217. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 194]
218. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 194]
223. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 196]
224. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 196]
225. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 196]
262. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 199]
263. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 199]
300. 0 <=> main_query_1(sK6,sK7) [avatar definition]
301. main_query_1(sK6,sK7) <- {0} [avatar component clause 300]
303. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
304. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 303]
306. 2 <=> main_query_2(sK6,sK7) [avatar definition]
307. main_query_2(sK6,sK7) <- {2} [avatar component clause 306]
309. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
310. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 309]
311. ~1 | ~3 [avatar split clause 263,309,303]
312. 0 | 2 [avatar split clause 262,306,300]
315. select_4(sK6,sK7) <- {2} [resolution 217,307]
316. where_6(sK6,sK7) <- {2} [resolution 315,215]
318. lte(sK7,sK6) <- {2} [resolution 316,212]
319. lte(sK6,sK7) <- {2} [resolution 316,211]
332. eq(sK6,sK7) | ~lte(sK6,sK7) <- {2} [resolution 223,318]
334. eq(sK6,sK7) <- {2} [subsumption resolution 332,319]
345. where_3(sK6,sK7) <- {2} [resolution 334,203]
360. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 204,200]
364. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 214,209]
365. from_5(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 360,210]
368. from_2(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 364,201]
369. ~where_6(X0,X1) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 365,216]
371. ~select_4(X0,X1) | ~where_3(X0,X1) | select_1(X0,X1) [resolution 368,206]
389. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 371,315]
390. select_1(sK6,sK7) <- {2} [subsumption resolution 389,345]
391. main_query_1(sK6,sK7) <- {2} [resolution 390,208]
393. $false <- {1, 2} [subsumption resolution 391,304]
394. 1 | ~2 [avatar contradiction clause 393]
395. select_1(sK6,sK7) <- {0} [resolution 301,207]
397. where_3(sK6,sK7) <- {0} [resolution 395,205]
398. eq(sK6,sK7) <- {0} [resolution 397,202]
400. lte(sK7,sK6) <- {0} [resolution 398,225]
401. lte(sK6,sK7) <- {0} [resolution 398,224]
406. where_6(sK6,sK7) | ~lte(sK6,sK7) <- {0} [resolution 400,213]
407. where_6(sK6,sK7) <- {0} [subsumption resolution 406,401]
419. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 407,369]
422. select_4(sK6,sK7) <- {0} [subsumption resolution 419,395]
430. main_query_2(sK6,sK7) <- {0} [resolution 422,218]
432. $false <- {0, 3} [subsumption resolution 430,310]
433. ~0 | 3 [avatar contradiction clause 432]
434. $false [avatar sat refutation 311,312,394,433]
% SZS output end Proof for query10_query12_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.018 s
% ------------------------------
% ------------------------------
