% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_13_query_15_tptp4X
% SZS output start Proof for query_13_query_15_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (where_6_subquery_7(X5) <=> select_8(X5)) [input]
9. ! [X0,X1] : (where_6(X0,X1) <=> (where_6_subquery_7(X0) & eq(X0,X1))) [input]
10. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
13. ! [X7,X8] : (eq(X7,X8) => eq(X8,X7)) [input]
16. ! [X7,X8] : ((lte(X8,X7) & lte(X7,X8)) <=> eq(X7,X8)) [input]
40. ! [X7,X5,X6] : ((main_query_2(X5,X6) & eq(X5,X7)) => main_query_2(X7,X6)) [input]
60. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
61. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 60]
62. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
63. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 7]
64. ! [X0] : (where_6_subquery_7(X0) <=> select_8(X0)) [rectify 8]
65. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
67. ! [X0,X1] : (eq(X0,X1) => eq(X1,X0)) [rectify 13]
70. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 16]
94. ! [X0,X1,X2] : ((main_query_2(X1,X2) & eq(X1,X0)) => main_query_2(X0,X2)) [rectify 40]
114. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 61]
115. ! [X0,X1] : (eq(X1,X0) | ~eq(X0,X1)) [ennf transformation 67]
154. ! [X0,X1,X2] : (main_query_2(X0,X2) | (~main_query_2(X1,X2) | ~eq(X1,X0))) [ennf transformation 94]
155. ! [X0,X1,X2] : (main_query_2(X0,X2) | ~main_query_2(X1,X2) | ~eq(X1,X0)) [flattening 154]
194. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 114]
195. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
196. ! [X0,X1] : ((where_3(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [nnf transformation 2]
197. ! [X0,X1] : ((where_3(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [flattening 196]
198. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
199. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 198]
200. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
201. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 199,200]
202. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 62]
203. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
204. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
205. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 63]
206. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 205]
207. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
208. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 206,207]
209. ! [X0] : ((where_6_subquery_7(X0) | ~select_8(X0)) & (select_8(X0) | ~where_6_subquery_7(X0))) [nnf transformation 64]
210. ! [X0,X1] : ((where_6(X0,X1) | (~where_6_subquery_7(X0) | ~eq(X0,X1))) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 9]
211. ! [X0,X1] : ((where_6(X0,X1) | ~where_6_subquery_7(X0) | ~eq(X0,X1)) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 210]
212. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 10]
213. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 212]
214. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
215. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 213,214]
216. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 65]
217. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 70]
218. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 217]
219. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 194]
220. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
221. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 219,220]
222. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 195]
223. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 195]
224. ~where_3(X0,X1) | lte(X0,X1) [cnf transformation 197]
225. ~where_3(X0,X1) | lte(X1,X0) [cnf transformation 197]
226. ~lte(X1,X0) | where_3(X0,X1) | ~lte(X0,X1) [cnf transformation 197]
227. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 201]
228. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 201]
229. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 201]
230. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 202]
231. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 202]
232. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 203]
233. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 203]
235. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 204]
237. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 208]
239. ~select_8(X0) | where_6_subquery_7(X0) [cnf transformation 209]
240. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 211]
241. ~where_6(X0,X1) | where_6_subquery_7(X0) [cnf transformation 211]
242. ~eq(X0,X1) | ~where_6_subquery_7(X0) | where_6(X0,X1) [cnf transformation 211]
243. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 215]
244. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 215]
245. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 215]
246. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 216]
247. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 216]
249. ~eq(X0,X1) | eq(X1,X0) [cnf transformation 115]
252. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 218]
253. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 218]
254. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 218]
274. ~main_query_2(X1,X2) | main_query_2(X0,X2) | ~eq(X1,X0) [cnf transformation 155]
294. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 221]
295. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 221]
297. 0 <=> main_query_1(sK10,sK11) [avatar definition]
298. main_query_1(sK10,sK11) <- {0} [avatar component clause 297]
300. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
301. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 300]
303. 2 <=> main_query_2(sK10,sK11) [avatar definition]
304. main_query_2(sK10,sK11) <- {2} [avatar component clause 303]
306. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
307. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 306]
308. ~1 | ~3 [avatar split clause 295,306,300]
309. 0 | 2 [avatar split clause 294,303,297]
310. select_4(sK10,sK11) <- {2} [resolution 246,304]
312. where_6(sK10,sK11) <- {2} [resolution 310,244]
314. eq(sK10,sK11) <- {2} [resolution 312,240]
323. eq(sK11,sK10) <- {2} [resolution 314,249]
339. ~where_6_subquery_7(sK11) | where_6(sK11,sK10) <- {2} [resolution 323,242]
347. 4 <=> where_6(sK11,sK10) [avatar definition]
348. where_6(sK11,sK10) <- {4} [avatar component clause 347]
353. 7 <=> ~where_6_subquery_7(sK11) [avatar definition]
354. ~where_6_subquery_7(sK11) <- {7} [avatar component clause 353]
355. 4 | ~7 | ~2 [avatar split clause 339,303,353,347]
374. ~eq(sK10,X0) | main_query_2(X0,sK11) <- {2} [resolution 274,304]
375. main_query_2(sK11,sK11) <- {2} [resolution 374,314]
380. select_4(sK11,sK11) <- {2} [resolution 375,246]
387. where_6(sK11,sK11) <- {2} [resolution 380,244]
392. where_6_subquery_7(sK11) <- {2} [resolution 387,241]
393. $false <- {2, 7} [subsumption resolution 392,354]
394. ~2 | 7 [avatar contradiction clause 393]
395. select_1(sK10,sK11) <- {0} [resolution 298,230]
398. where_3(sK10,sK11) <- {0} [resolution 395,228]
400. lte(sK11,sK10) <- {0} [resolution 398,225]
401. lte(sK10,sK11) <- {0} [resolution 398,224]
404. eq(sK10,sK11) | ~lte(sK10,sK11) <- {0} [resolution 400,252]
406. eq(sK10,sK11) <- {0} [subsumption resolution 404,401]
414. ~where_6_subquery_7(sK10) | where_6(sK10,sK11) <- {0} [resolution 406,242]
430. 8 <=> where_6(sK10,sK11) [avatar definition]
431. where_6(sK10,sK11) <- {8} [avatar component clause 430]
433. 10 <=> where_6_subquery_7(sK10) [avatar definition]
436. 11 <=> ~where_6_subquery_7(sK10) [avatar definition]
438. 8 | ~11 | ~0 [avatar split clause 414,297,436,430]
474. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 227,222]
476. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 243,232]
634. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 474,233]
635. from_9(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 474,235]
654. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 476,223]
690. 14 <=> select_8(sK10) [avatar definition]
691. select_8(sK10) <- {14} [avatar component clause 690]
1301. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 634,245]
1309. ~select_1(X2,X3) | select_8(X2) [resolution 635,237]
1315. select_8(sK10) <- {0} [resolution 1309,395]
1330. 14 | ~0 [avatar split clause 1315,297,690]
1332. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 654,229]
1360. where_6_subquery_7(sK10) <- {14} [resolution 691,239]
1363. 10 | ~14 [avatar split clause 1360,690,433]
1384. eq(sK11,sK10) <- {4} [resolution 348,240]
1397. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {8} [resolution 431,1301]
1400. select_4(sK10,sK11) <- {0, 8} [subsumption resolution 1397,395]
1451. main_query_2(sK10,sK11) <- {0, 8} [resolution 1400,247]
1453. $false <- {0, 3, 8} [subsumption resolution 1451,307]
1454. ~0 | 3 | ~8 [avatar contradiction clause 1453]
1470. select_4(sK10,sK11) <- {2} [resolution 304,246]
1484. lte(sK11,sK10) <- {4} [resolution 1384,253]
1485. lte(sK10,sK11) <- {4} [resolution 1384,254]
1526. where_3(sK10,sK11) | ~lte(sK10,sK11) <- {4} [resolution 1484,226]
1531. where_3(sK10,sK11) <- {4} [subsumption resolution 1526,1485]
1914. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 1332,1470]
1918. select_1(sK10,sK11) <- {2, 4} [subsumption resolution 1914,1531]
1928. main_query_1(sK10,sK11) <- {2, 4} [resolution 1918,231]
1930. $false <- {1, 2, 4} [subsumption resolution 1928,301]
1931. 1 | ~2 | ~4 [avatar contradiction clause 1930]
1932. $false [avatar sat refutation 308,309,355,394,438,1330,1363,1454,1931]
% SZS output end Proof for query_13_query_15_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5500
% Time elapsed: 0.030 s
% ------------------------------
% ------------------------------
