% Refutation found. Thanks to Tanya!
% SZS status Theorem for query07_query09_tptp4X
% SZS output start Proof for query07_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X0,X5,X6,X7,X8] : (from_8(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
6. ! [X0] : (select_7(X0) <=> ? [X5,X6,X7,X8] : ($true & from_8(X0,X5,X6,X7,X8))) [input]
7. ! [X0] : (subquery_from_6(X0) <=> select_7(X0)) [input]
8. ! [X0] : (from_5(X0) <=> subquery_from_6(X0)) [input]
9. ! [X0] : (select_4(X0) <=> ($true & from_5(X0))) [input]
10. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [input]
11. ! [X9] : eq(X9,X9) [input]
17. ! [X9,X10] : (lt(X9,X10) <=> (~eq(X9,X10) & lte(X9,X10))) [input]
18. ! [X9,X10] : (gt(X9,X10) <=> lt(X10,X9)) [input]
44. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
45. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 44]
46. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
47. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
48. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 47]
49. ! [X0] : (select_4(X0) <=> from_5(X0)) [true and false elimination 9]
50. ! [X0] : eq(X0,X0) [rectify 11]
56. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 17]
57. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 18]
134. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 45]
135. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
136. ! [X0] : ((where_3(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
137. ! [X0] : ((where_3(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [flattening 136]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
139. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 138]
140. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
141. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 139,140]
142. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
143. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 46]
144. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 48]
145. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 144]
146. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
147. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0))) [skolemisation 145,146]
148. ! [X0] : ((subquery_from_6(X0) | ~select_7(X0)) & (select_7(X0) | ~subquery_from_6(X0))) [nnf transformation 7]
149. ! [X0] : ((from_5(X0) | ~subquery_from_6(X0)) & (subquery_from_6(X0) | ~from_5(X0))) [nnf transformation 8]
150. ! [X0] : ((select_4(X0) | ~from_5(X0)) & (from_5(X0) | ~select_4(X0))) [nnf transformation 49]
151. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 10]
154. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 56]
155. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 154]
156. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 57]
157. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 134]
158. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
159. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 157,158]
160. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 135]
161. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 135]
164. gt(X0,X0) | where_3(X0) | lt(X0,X0) [cnf transformation 137]
165. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 141]
167. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 141]
168. ~main_query_1(X0) | select_1(X0) [cnf transformation 142]
169. ~select_1(X0) | main_query_1(X0) [cnf transformation 142]
170. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 143]
171. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 143]
172. from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [cnf transformation 147]
173. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 147]
174. ~subquery_from_6(X0) | select_7(X0) [cnf transformation 148]
175. ~select_7(X0) | subquery_from_6(X0) [cnf transformation 148]
176. ~from_5(X0) | subquery_from_6(X0) [cnf transformation 149]
177. ~subquery_from_6(X0) | from_5(X0) [cnf transformation 149]
178. ~select_4(X0) | from_5(X0) [cnf transformation 150]
179. ~from_5(X0) | select_4(X0) [cnf transformation 150]
180. ~main_query_2(X0) | select_4(X0) [cnf transformation 151]
181. ~select_4(X0) | main_query_2(X0) [cnf transformation 151]
182. eq(X0,X0) [cnf transformation 50]
191. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 155]
193. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 156]
218. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 159]
219. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 159]
244. 0 <=> main_query_1(sK8) [avatar definition]
245. main_query_1(sK8) <- {0} [avatar component clause 244]
247. 1 <=> ~main_query_1(sK8) [avatar definition]
248. ~main_query_1(sK8) <- {1} [avatar component clause 247]
250. 2 <=> main_query_2(sK8) [avatar definition]
251. main_query_2(sK8) <- {2} [avatar component clause 250]
253. 3 <=> ~main_query_2(sK8) [avatar definition]
254. ~main_query_2(sK8) <- {3} [avatar component clause 253]
255. ~1 | ~3 [avatar split clause 219,253,247]
256. 0 | 2 [avatar split clause 218,250,244]
257. select_4(sK8) <- {2} [resolution 180,251]
258. from_5(sK8) <- {2} [resolution 257,178]
261. subquery_from_6(sK8) <- {2} [resolution 258,176]
263. select_7(sK8) <- {2} [resolution 261,174]
268. ~lt(X0,X0) [resolution 191,182]
270. where_3(X1) | lt(X1,X1) | lt(X1,X1) [resolution 164,193]
271. where_3(X1) | lt(X1,X1) [duplicate literal removal 270]
272. where_3(X1) [subsumption resolution 271,268]
277. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 165,160]
279. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [resolution 172,170]
281. from_8(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 277,171]
284. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_7(X1) [resolution 279,161]
286. ~select_1(X1) | select_7(X1) [resolution 281,173]
288. ~select_7(X1) | ~where_3(X1) | select_1(X1) [resolution 284,167]
289. ~select_7(X1) | select_1(X1) [subsumption resolution 288,272]
290. select_1(sK8) <- {2} [resolution 289,263]
292. main_query_1(sK8) <- {2} [resolution 290,169]
294. $false <- {1, 2} [subsumption resolution 292,248]
295. 1 | ~2 [avatar contradiction clause 294]
296. select_1(sK8) <- {0} [resolution 245,168]
297. select_7(sK8) <- {0} [resolution 296,286]
301. subquery_from_6(sK8) <- {0} [resolution 297,175]
302. from_5(sK8) <- {0} [resolution 301,177]
304. select_4(sK8) <- {0} [resolution 302,179]
306. main_query_2(sK8) <- {0} [resolution 304,181]
308. $false <- {0, 3} [subsumption resolution 306,254]
309. ~0 | 3 [avatar contradiction clause 308]
310. $false [avatar sat refutation 255,256,295,309]
% SZS output end Proof for query07_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
