% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_03_query_07_tptp4X
% SZS output start Proof for query_03_query_07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
4. ! [X7,X8,X9,X10,X11] : (from_4(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
5. ! [X7] : (where_5(X7) <=> (~gt(X7,X7) & ~lt(X7,X7))) [input]
6. ! [X7,X8] : (select_3(X7,X8) <=> ? [X9,X10,X11] : (where_5(X7) & from_4(X7,X8,X9,X10,X11))) [input]
7. ! [X5,X6] : (main_query_2(X5,X6) <=> select_3(X5,X6)) [input]
8. ! [X12] : eq(X12,X12) [input]
14. ! [X12,X13] : (lt(X12,X13) <=> (~eq(X12,X13) & lte(X12,X13))) [input]
15. ! [X12,X13] : (gt(X12,X13) <=> lt(X13,X12)) [input]
46. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
47. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 46]
48. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 3]
50. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
51. ! [X0] : (where_5(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [rectify 5]
52. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
53. ! [X0,X1] : (main_query_2(X0,X1) <=> select_3(X0,X1)) [rectify 7]
54. ! [X0] : eq(X0,X0) [rectify 8]
60. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 14]
61. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 15]
92. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 47]
154. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 92]
155. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
156. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0,X1))) [nnf transformation 48]
157. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) | ~select_1(X0,X1))) [rectify 156]
158. ! [X1,X0] : (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) => from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) [choice axiom]
159. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1))) [skolemisation 157,158]
160. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 49]
161. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 50]
162. ! [X0] : ((where_5(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [nnf transformation 51]
163. ! [X0] : ((where_5(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [flattening 162]
164. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0,X1))) [nnf transformation 52]
165. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) | ~select_3(X0,X1))) [rectify 164]
166. ! [X1,X0] : (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) => (where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
167. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_3(X0,X1))) [skolemisation 165,166]
168. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 53]
171. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 60]
172. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 171]
173. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 61]
174. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 154]
175. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
176. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 174,175]
177. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 155]
178. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 155]
179. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 159]
180. ~from_2(X0,X1,X2,X3,X4) | select_1(X0,X1) [cnf transformation 159]
181. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 160]
182. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 160]
183. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 161]
184. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 161]
187. gt(X0,X0) | where_5(X0) | lt(X0,X0) [cnf transformation 163]
188. from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [cnf transformation 167]
190. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0,X1) [cnf transformation 167]
191. ~main_query_2(X0,X1) | select_3(X0,X1) [cnf transformation 168]
192. ~select_3(X0,X1) | main_query_2(X0,X1) [cnf transformation 168]
193. eq(X0,X0) [cnf transformation 54]
202. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 172]
204. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 173]
234. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 176]
235. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 176]
237. 0 <=> main_query_1(sK6,sK7) [avatar definition]
238. main_query_1(sK6,sK7) <- {0} [avatar component clause 237]
240. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
241. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 240]
243. 2 <=> main_query_2(sK6,sK7) [avatar definition]
244. main_query_2(sK6,sK7) <- {2} [avatar component clause 243]
246. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
247. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 246]
248. ~1 | ~3 [avatar split clause 235,246,240]
249. 0 | 2 [avatar split clause 234,243,237]
250. select_3(sK6,sK7) <- {2} [resolution 191,244]
256. ~lt(X0,X0) [resolution 202,193]
258. where_5(X1) | lt(X1,X1) | lt(X1,X1) [resolution 187,204]
259. where_5(X1) | lt(X1,X1) [duplicate literal removal 258]
260. where_5(X1) [subsumption resolution 259,256]
273. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 179,177]
275. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 188,183]
288. from_4(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 273,184]
294. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 275,178]
309. ~select_1(X2,X3) | ~where_5(X2) | select_3(X2,X3) [resolution 288,190]
315. ~select_1(X2,X3) | select_3(X2,X3) [subsumption resolution 309,260]
331. ~select_3(X2,X3) | select_1(X2,X3) [resolution 294,180]
344. select_1(sK6,sK7) <- {2} [resolution 331,250]
353. main_query_1(sK6,sK7) <- {2} [resolution 344,182]
354. $false <- {1, 2} [subsumption resolution 353,241]
355. 1 | ~2 [avatar contradiction clause 354]
356. select_1(sK6,sK7) <- {0} [resolution 238,181]
357. select_3(sK6,sK7) <- {0} [resolution 356,315]
360. main_query_2(sK6,sK7) <- {0} [resolution 357,192]
362. $false <- {0, 3} [subsumption resolution 360,247]
363. ~0 | 3 [avatar contradiction clause 362]
364. $false [avatar sat refutation 248,249,355,363]
% SZS output end Proof for query_03_query_07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
