% Refutation found. Thanks to Tanya!
% SZS status Theorem for query04_query07_tptp4X
% SZS output start Proof for query04_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
7. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
9. ! [X6] : eq(X6,X6) [input]
15. ! [X6,X7] : (lt(X6,X7) <=> (~eq(X6,X7) & lte(X6,X7))) [input]
16. ! [X6,X7] : (gt(X6,X7) <=> lt(X7,X6)) [input]
48. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
49. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 48]
50. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
51. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
52. ! [X0] : eq(X0,X0) [rectify 9]
58. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 15]
59. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 16]
91. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 49]
155. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 91]
156. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
157. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
158. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
159. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 158]
160. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
161. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 159,160]
162. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 50]
163. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
164. ! [X0] : ((where_6(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [nnf transformation 6]
165. ! [X0] : ((where_6(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [flattening 164]
166. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 7]
167. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 166]
168. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
169. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 167,168]
170. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 51]
173. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 58]
174. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 173]
175. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 59]
176. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 155]
177. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
178. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 176,177]
179. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 156]
180. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 156]
182. ~eq(X0,X0) | where_3(X0) [cnf transformation 157]
183. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 161]
185. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 161]
186. ~main_query_1(X0) | select_1(X0) [cnf transformation 162]
187. ~select_1(X0) | main_query_1(X0) [cnf transformation 162]
188. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 163]
189. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 163]
192. gt(X0,X0) | where_6(X0) | lt(X0,X0) [cnf transformation 165]
193. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 169]
195. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 169]
196. ~main_query_2(X0) | select_4(X0) [cnf transformation 170]
197. ~select_4(X0) | main_query_2(X0) [cnf transformation 170]
198. eq(X0,X0) [cnf transformation 52]
207. ~lt(X0,X1) | ~eq(X0,X1) [cnf transformation 174]
209. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 175]
240. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 178]
241. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 178]
272. 0 <=> main_query_1(sK8) [avatar definition]
273. main_query_1(sK8) <- {0} [avatar component clause 272]
275. 1 <=> ~main_query_1(sK8) [avatar definition]
276. ~main_query_1(sK8) <- {1} [avatar component clause 275]
278. 2 <=> main_query_2(sK8) [avatar definition]
279. main_query_2(sK8) <- {2} [avatar component clause 278]
281. 3 <=> ~main_query_2(sK8) [avatar definition]
282. ~main_query_2(sK8) <- {3} [avatar component clause 281]
283. ~1 | ~3 [avatar split clause 241,281,275]
284. 0 | 2 [avatar split clause 240,278,272]
285. select_4(sK8) <- {2} [resolution 196,279]
289. where_3(X1) [resolution 182,198]
297. where_6(X1) | lt(X1,X1) | lt(X1,X1) [resolution 192,209]
298. lt(X1,X1) | where_6(X1) [duplicate literal removal 297]
303. where_6(X2) | ~eq(X2,X2) [resolution 298,207]
305. where_6(X2) [subsumption resolution 303,198]
311. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 183,179]
313. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 193,188]
315. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 311,189]
318. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 313,180]
320. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 315,195]
321. ~select_1(X1) | select_4(X1) [subsumption resolution 320,305]
323. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 318,185]
324. ~select_4(X1) | select_1(X1) [subsumption resolution 323,289]
325. select_1(sK8) <- {2} [resolution 324,285]
327. main_query_1(sK8) <- {2} [resolution 325,187]
329. $false <- {1, 2} [subsumption resolution 327,276]
330. 1 | ~2 [avatar contradiction clause 329]
331. select_1(sK8) <- {0} [resolution 273,186]
332. select_4(sK8) <- {0} [resolution 331,321]
336. main_query_2(sK8) <- {0} [resolution 332,197]
338. $false <- {0, 3} [subsumption resolution 336,282]
339. ~0 | 3 [avatar contradiction clause 338]
340. $false [avatar sat refutation 283,284,330,339]
% SZS output end Proof for query04_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.015 s
% ------------------------------
% ------------------------------
