% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_01_query_05_tptp4X
% SZS output start Proof for query_01_query_05_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [input]
6. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X5,X6] : (main_query_2(X5,X6) <=> select_3(X5,X6)) [input]
8. ! [X7] : eq(X7,X7) [input]
46. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
47. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 46]
48. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 3]
50. ! [X0,X1] : (main_query_2(X0,X1) <=> select_3(X0,X1)) [rectify 7]
51. ! [X0] : eq(X0,X0) [rectify 8]
89. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 47]
151. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 89]
152. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
153. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0,X1))) [nnf transformation 48]
154. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) | ~select_1(X0,X1))) [rectify 153]
155. ! [X1,X0] : (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) => from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) [choice axiom]
156. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1))) [skolemisation 154,155]
157. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 49]
158. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
159. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 5]
160. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0,X1))) [nnf transformation 6]
161. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) | ~select_3(X0,X1))) [rectify 160]
162. ! [X1,X0] : (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) => (where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
163. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_3(X0,X1))) [skolemisation 161,162]
164. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 50]
167. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 151]
168. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
169. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 167,168]
170. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 152]
171. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 152]
172. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 156]
173. ~from_2(X0,X1,X2,X3,X4) | select_1(X0,X1) [cnf transformation 156]
174. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 157]
175. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 157]
176. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 158]
177. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 158]
179. ~eq(X0,X0) | where_5(X0) [cnf transformation 159]
180. from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [cnf transformation 163]
182. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0,X1) [cnf transformation 163]
183. ~main_query_2(X0,X1) | select_3(X0,X1) [cnf transformation 164]
184. ~select_3(X0,X1) | main_query_2(X0,X1) [cnf transformation 164]
185. eq(X0,X0) [cnf transformation 51]
221. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 169]
222. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 169]
224. 0 <=> main_query_1(sK6,sK7) [avatar definition]
225. main_query_1(sK6,sK7) <- {0} [avatar component clause 224]
227. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
228. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 227]
230. 2 <=> main_query_2(sK6,sK7) [avatar definition]
231. main_query_2(sK6,sK7) <- {2} [avatar component clause 230]
233. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
234. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 233]
235. ~1 | ~3 [avatar split clause 222,233,227]
236. 0 | 2 [avatar split clause 221,230,224]
238. where_5(X1) [resolution 179,185]
239. select_3(sK6,sK7) <- {2} [resolution 183,231]
272. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 172,170]
274. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 180,176]
287. from_4(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 272,177]
293. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 274,171]
308. ~select_1(X2,X3) | ~where_5(X2) | select_3(X2,X3) [resolution 287,182]
314. ~select_1(X2,X3) | select_3(X2,X3) [subsumption resolution 308,238]
330. ~select_3(X2,X3) | select_1(X2,X3) [resolution 293,173]
343. select_1(sK6,sK7) <- {2} [resolution 330,239]
352. main_query_1(sK6,sK7) <- {2} [resolution 343,175]
353. $false <- {1, 2} [subsumption resolution 352,228]
354. 1 | ~2 [avatar contradiction clause 353]
355. select_1(sK6,sK7) <- {0} [resolution 225,174]
356. select_3(sK6,sK7) <- {0} [resolution 355,314]
361. main_query_2(sK6,sK7) <- {0} [resolution 356,184]
363. $false <- {0, 3} [subsumption resolution 361,234]
364. ~0 | 3 [avatar contradiction clause 363]
365. $false [avatar sat refutation 235,236,354,364]
% SZS output end Proof for query_01_query_05_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
