% Refutation found. Thanks to Tanya!
% SZS status Theorem for query01_query03_tptp4X
% SZS output start Proof for query01_query03_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X5,X6,X7,X8] : (from_4(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
5. ! [X0] : (select_3(X0) <=> ? [X5,X6,X7,X8] : ($true & from_4(X0,X5,X6,X7,X8))) [input]
6. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
38. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
39. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 38]
40. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
41. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
42. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [rectify 5]
43. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 42]
120. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 39]
121. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
122. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 40]
123. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 122]
124. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
125. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 123,124]
126. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
127. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 41]
128. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 43]
129. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 128]
130. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
131. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0))) [skolemisation 129,130]
132. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 6]
135. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 120]
136. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
137. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 135,136]
138. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 121]
139. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 121]
140. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 125]
141. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 125]
142. ~main_query_1(X0) | select_1(X0) [cnf transformation 126]
143. ~select_1(X0) | main_query_1(X0) [cnf transformation 126]
144. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 127]
145. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 127]
146. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 131]
147. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 131]
148. ~main_query_2(X0) | select_3(X0) [cnf transformation 132]
149. ~select_3(X0) | main_query_2(X0) [cnf transformation 132]
179. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 137]
180. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 137]
182. 0 <=> main_query_1(sK8) [avatar definition]
183. main_query_1(sK8) <- {0} [avatar component clause 182]
185. 1 <=> ~main_query_1(sK8) [avatar definition]
186. ~main_query_1(sK8) <- {1} [avatar component clause 185]
188. 2 <=> main_query_2(sK8) [avatar definition]
189. main_query_2(sK8) <- {2} [avatar component clause 188]
191. 3 <=> ~main_query_2(sK8) [avatar definition]
192. ~main_query_2(sK8) <- {3} [avatar component clause 191]
193. ~1 | ~3 [avatar split clause 180,191,185]
194. 0 | 2 [avatar split clause 179,188,182]
195. select_3(sK8) <- {2} [resolution 148,189]
226. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 140,138]
228. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 146,144]
230. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 226,145]
234. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 228,139]
237. ~select_1(X1) | select_3(X1) [resolution 230,147]
245. ~select_3(X9) | select_1(X9) [resolution 234,141]
248. select_1(sK8) <- {2} [resolution 245,195]
252. main_query_1(sK8) <- {2} [resolution 248,143]
253. $false <- {1, 2} [subsumption resolution 252,186]
254. 1 | ~2 [avatar contradiction clause 253]
255. select_1(sK8) <- {0} [resolution 183,142]
258. select_3(sK8) <- {0} [resolution 255,237]
263. main_query_2(sK8) <- {0} [resolution 258,149]
264. $false <- {0, 3} [subsumption resolution 263,192]
265. ~0 | 3 [avatar contradiction clause 264]
266. $false [avatar sat refutation 193,194,254,265]
% SZS output end Proof for query01_query03_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
