% Refutation found. Thanks to Tanya!
% SZS status Theorem for query11_query14_tptp4X
% SZS output start Proof for query11_query14_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8,X9,X10,X11] : (from_9(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
7. ! [X7] : (select_8(X7) <=> ? [X8,X9,X10,X11] : ($true & from_9(X7,X8,X9,X10,X11))) [input]
8. ! [X5] : (where_6_subquery_7(X5) <=> select_8(X5)) [input]
9. ! [X7,X8] : (where_6(X7,X8) <=> (where_6_subquery_7(X7) & eq(X7,X8))) [input]
10. ! [X7,X8] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X7,X8) & from_5(X7,X8,X9,X10,X11))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
60. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
61. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 60]
62. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
63. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
64. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 6]
65. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [rectify 7]
66. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 65]
67. ! [X0] : (where_6_subquery_7(X0) <=> select_8(X0)) [rectify 8]
68. ! [X0,X1] : (where_6(X0,X1) <=> (where_6_subquery_7(X0) & eq(X0,X1))) [rectify 9]
69. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 10]
70. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
119. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 61]
199. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 119]
200. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
201. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
202. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
203. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 202]
204. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
205. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 203,204]
206. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 62]
207. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 63]
208. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 64]
209. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 66]
210. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 209]
211. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
212. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 210,211]
213. ! [X0] : ((where_6_subquery_7(X0) | ~select_8(X0)) & (select_8(X0) | ~where_6_subquery_7(X0))) [nnf transformation 67]
214. ! [X0,X1] : ((where_6(X0,X1) | (~where_6_subquery_7(X0) | ~eq(X0,X1))) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 68]
215. ! [X0,X1] : ((where_6(X0,X1) | ~where_6_subquery_7(X0) | ~eq(X0,X1)) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 214]
216. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 69]
217. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 216]
218. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
219. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 217,218]
220. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 70]
223. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 199]
224. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
225. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 223,224]
226. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 200]
227. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 200]
228. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 201]
229. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 201]
230. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 205]
231. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 205]
232. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 205]
233. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 206]
234. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 206]
235. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 207]
236. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 207]
238. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 208]
240. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 212]
242. ~select_8(X0) | where_6_subquery_7(X0) [cnf transformation 213]
243. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 215]
245. ~eq(X0,X1) | ~where_6_subquery_7(X0) | where_6(X0,X1) [cnf transformation 215]
246. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 219]
247. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 219]
248. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 219]
249. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 220]
250. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 220]
297. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 225]
298. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 225]
300. 0 <=> main_query_1(sK10,sK11) [avatar definition]
301. main_query_1(sK10,sK11) <- {0} [avatar component clause 300]
303. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
304. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 303]
306. 2 <=> main_query_2(sK10,sK11) [avatar definition]
307. main_query_2(sK10,sK11) <- {2} [avatar component clause 306]
309. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
310. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 309]
311. ~1 | ~3 [avatar split clause 298,309,303]
312. 0 | 2 [avatar split clause 297,306,300]
392. select_1(sK10,sK11) <- {0} [resolution 301,233]
395. where_3(sK10,sK11) <- {0} [resolution 392,231]
397. eq(sK10,sK11) <- {0} [resolution 395,228]
400. ~where_6_subquery_7(sK10) | where_6(sK10,sK11) <- {0} [resolution 397,245]
411. 8 <=> where_6(sK10,sK11) [avatar definition]
412. where_6(sK10,sK11) <- {8} [avatar component clause 411]
414. 10 <=> where_6_subquery_7(sK10) [avatar definition]
417. 11 <=> ~where_6_subquery_7(sK10) [avatar definition]
419. 8 | ~11 | ~0 [avatar split clause 400,300,417,411]
458. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 230,226]
462. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 246,235]
602. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 458,236]
603. from_9(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 458,238]
620. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 462,227]
655. 14 <=> select_8(sK10) [avatar definition]
656. select_8(sK10) <- {14} [avatar component clause 655]
1203. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 602,248]
1211. ~select_1(X2,X3) | select_8(X2) [resolution 603,240]
1217. select_8(sK10) <- {0} [resolution 1211,392]
1232. 14 | ~0 [avatar split clause 1217,300,655]
1234. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 620,232]
1252. where_6_subquery_7(sK10) <- {14} [resolution 656,242]
1255. 10 | ~14 [avatar split clause 1252,655,414]
1282. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {8} [resolution 412,1203]
1287. select_4(sK10,sK11) <- {0, 8} [subsumption resolution 1282,392]
1312. main_query_2(sK10,sK11) <- {0, 8} [resolution 1287,250]
1314. $false <- {0, 3, 8} [subsumption resolution 1312,310]
1315. ~0 | 3 | ~8 [avatar contradiction clause 1314]
1331. select_4(sK10,sK11) <- {2} [resolution 307,249]
1335. where_6(sK10,sK11) <- {2} [resolution 1331,247]
1336. 8 | ~2 [avatar split clause 1335,306,411]
1340. eq(sK10,sK11) <- {8} [resolution 412,243]
1343. where_3(sK10,sK11) <- {8} [resolution 1340,229]
1761. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 1234,1331]
1765. select_1(sK10,sK11) <- {2, 8} [subsumption resolution 1761,1343]
1774. main_query_1(sK10,sK11) <- {2, 8} [resolution 1765,234]
1780. $false <- {1, 2, 8} [subsumption resolution 1774,304]
1781. 1 | ~2 | ~8 [avatar contradiction clause 1780]
1782. $false [avatar sat refutation 311,312,419,1232,1255,1315,1336,1781]
% SZS output end Proof for query11_query14_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5500
% Time elapsed: 0.074 s
% ------------------------------
% ------------------------------
