% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query05_tptp4X
% SZS output start Proof for query00_query05_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [input]
6. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
8. ! [X6] : eq(X6,X6) [input]
39. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
40. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 39]
41. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
42. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
43. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
44. ! [X0] : eq(X0,X0) [rectify 8]
75. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 40]
121. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 75]
122. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
123. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 41]
124. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 123]
125. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
126. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 124,125]
127. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 42]
128. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
129. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 5]
130. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 6]
131. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 130]
132. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
133. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 131,132]
134. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 43]
137. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 121]
138. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
139. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 137,138]
140. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 122]
141. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 122]
142. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 126]
143. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 126]
144. ~main_query_1(X0) | select_1(X0) [cnf transformation 127]
145. ~select_1(X0) | main_query_1(X0) [cnf transformation 127]
146. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 128]
147. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 128]
149. ~eq(X0,X0) | where_5(X0) [cnf transformation 129]
150. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 133]
152. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 133]
153. ~main_query_2(X0) | select_3(X0) [cnf transformation 134]
154. ~select_3(X0) | main_query_2(X0) [cnf transformation 134]
155. eq(X0,X0) [cnf transformation 44]
184. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 139]
185. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 139]
187. 0 <=> main_query_1(sK8) [avatar definition]
188. main_query_1(sK8) <- {0} [avatar component clause 187]
190. 1 <=> ~main_query_1(sK8) [avatar definition]
191. ~main_query_1(sK8) <- {1} [avatar component clause 190]
193. 2 <=> main_query_2(sK8) [avatar definition]
194. main_query_2(sK8) <- {2} [avatar component clause 193]
196. 3 <=> ~main_query_2(sK8) [avatar definition]
197. ~main_query_2(sK8) <- {3} [avatar component clause 196]
198. ~1 | ~3 [avatar split clause 185,196,190]
199. 0 | 2 [avatar split clause 184,193,187]
200. select_3(sK8) <- {2} [resolution 153,194]
204. where_5(X1) [resolution 149,155]
238. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 142,140]
240. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 150,146]
243. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 238,147]
247. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 240,141]
253. ~select_1(X7) | ~where_5(X7) | select_3(X7) [resolution 243,152]
254. ~select_1(X7) | select_3(X7) [subsumption resolution 253,204]
260. ~select_3(X1) | select_1(X1) [resolution 247,143]
263. select_1(sK8) <- {2} [resolution 260,200]
267. main_query_1(sK8) <- {2} [resolution 263,145]
268. $false <- {1, 2} [subsumption resolution 267,191]
269. 1 | ~2 [avatar contradiction clause 268]
270. select_1(sK8) <- {0} [resolution 188,144]
273. select_3(sK8) <- {0} [resolution 270,254]
278. main_query_2(sK8) <- {0} [resolution 273,154]
280. $false <- {0, 3} [subsumption resolution 278,197]
281. ~0 | 3 [avatar contradiction clause 280]
282. $false [avatar sat refutation 198,199,269,281]
% SZS output end Proof for query00_query05_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
