% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_02_query_04_tptp4X
% SZS output start Proof for query_02_query_04_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0,X5] : (main_query_1(X0,X5) <=> select_1(X0,X5)) [input]
4. ! [X0,X6,X7,X8,X9] : (from_4(X0,X6,X7,X8,X9) <=> tab1(X0,X6,X7,X8,X9)) [input]
5. ! [X0,X6] : (select_3(X0,X6) <=> ? [X7,X8,X9] : ($true & from_4(X0,X6,X7,X8,X9))) [input]
6. ! [X0,X6] : (main_query_2(X0,X6) <=> select_3(X0,X6)) [input]
45. ! [X0,X5] : (main_query_1(X0,X5) <=> main_query_2(X0,X5)) [input]
46. ~! [X0,X5] : (main_query_1(X0,X5) <=> main_query_2(X0,X5)) [negated conjecture 45]
47. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
48. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 3]
49. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
50. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [rectify 5]
51. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 50]
52. ! [X0,X1] : (main_query_2(X0,X1) <=> select_3(X0,X1)) [rectify 6]
91. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 46]
153. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 91]
154. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
155. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0,X1))) [nnf transformation 47]
156. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) | ~select_1(X0,X1))) [rectify 155]
157. ! [X1,X0] : (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) => from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) [choice axiom]
158. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1))) [skolemisation 156,157]
159. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 48]
160. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 49]
161. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0,X1))) [nnf transformation 51]
162. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_4(X0,X1,X5,X6,X7) | ~select_3(X0,X1))) [rectify 161]
163. ! [X1,X0] : (? [X5,X6,X7] : from_4(X0,X1,X5,X6,X7) => from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) [choice axiom]
164. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1))) [skolemisation 162,163]
165. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 52]
168. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 153]
169. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
170. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 168,169]
171. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 154]
172. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 154]
173. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 158]
174. ~from_2(X0,X1,X2,X3,X4) | select_1(X0,X1) [cnf transformation 158]
175. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 159]
176. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 159]
177. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 160]
178. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 160]
179. from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [cnf transformation 164]
180. ~from_4(X0,X1,X2,X3,X4) | select_3(X0,X1) [cnf transformation 164]
181. ~main_query_2(X0,X1) | select_3(X0,X1) [cnf transformation 165]
182. ~select_3(X0,X1) | main_query_2(X0,X1) [cnf transformation 165]
219. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 170]
220. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 170]
222. 0 <=> main_query_1(sK6,sK7) [avatar definition]
223. main_query_1(sK6,sK7) <- {0} [avatar component clause 222]
225. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
226. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 225]
228. 2 <=> main_query_2(sK6,sK7) [avatar definition]
229. main_query_2(sK6,sK7) <- {2} [avatar component clause 228]
231. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
232. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 231]
233. ~1 | ~3 [avatar split clause 220,231,225]
234. 0 | 2 [avatar split clause 219,228,222]
235. select_3(sK6,sK7) <- {2} [resolution 181,229]
251. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 173,171]
253. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 179,177]
266. from_4(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 251,178]
272. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 253,172]
287. ~select_1(X2,X3) | select_3(X2,X3) [resolution 266,180]
308. ~select_3(X2,X3) | select_1(X2,X3) [resolution 272,174]
321. select_1(sK6,sK7) <- {2} [resolution 308,235]
330. main_query_1(sK6,sK7) <- {2} [resolution 321,176]
331. $false <- {1, 2} [subsumption resolution 330,226]
332. 1 | ~2 [avatar contradiction clause 331]
333. select_1(sK6,sK7) <- {0} [resolution 223,175]
334. select_3(sK6,sK7) <- {0} [resolution 333,287]
337. main_query_2(sK6,sK7) <- {0} [resolution 334,182]
338. $false <- {0, 3} [subsumption resolution 337,232]
339. ~0 | 3 [avatar contradiction clause 338]
340. $false [avatar sat refutation 233,234,332,339]
% SZS output end Proof for query_02_query_04_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
