% Refutation found. Thanks to Tanya!
% SZS status Theorem for query03_query04_tptp4X
% SZS output start Proof for query03_query04_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_4(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (where_5(X6) <=> eq(X6,X6)) [input]
6. ! [X6] : (select_3(X6) <=> ? [X7,X8,X9,X10] : (where_5(X6) & from_4(X6,X7,X8,X9,X10))) [input]
7. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
8. ! [X11] : eq(X11,X11) [input]
41. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
42. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
45. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
46. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [rectify 5]
47. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
48. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
49. ! [X0] : eq(X0,X0) [rectify 8]
82. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 42]
134. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 82]
135. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 136]
138. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
139. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 137,138]
140. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 44]
141. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 45]
142. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 46]
143. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 47]
144. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 143]
145. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
146. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 144,145]
147. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 48]
150. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 134]
151. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
152. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 150,151]
153. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 135]
154. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 135]
155. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 139]
156. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 139]
157. ~main_query_1(X0) | select_1(X0) [cnf transformation 140]
158. ~select_1(X0) | main_query_1(X0) [cnf transformation 140]
159. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 141]
160. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 141]
162. ~eq(X0,X0) | where_5(X0) [cnf transformation 142]
163. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 146]
165. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 146]
166. ~main_query_2(X0) | select_3(X0) [cnf transformation 147]
167. ~select_3(X0) | main_query_2(X0) [cnf transformation 147]
168. eq(X0,X0) [cnf transformation 49]
199. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 152]
200. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 152]
225. 0 <=> main_query_1(sK8) [avatar definition]
226. main_query_1(sK8) <- {0} [avatar component clause 225]
228. 1 <=> ~main_query_1(sK8) [avatar definition]
229. ~main_query_1(sK8) <- {1} [avatar component clause 228]
231. 2 <=> main_query_2(sK8) [avatar definition]
232. main_query_2(sK8) <- {2} [avatar component clause 231]
234. 3 <=> ~main_query_2(sK8) [avatar definition]
235. ~main_query_2(sK8) <- {3} [avatar component clause 234]
236. ~1 | ~3 [avatar split clause 200,234,228]
237. 0 | 2 [avatar split clause 199,231,225]
238. select_3(sK8) <- {2} [resolution 166,232]
242. where_5(X1) [resolution 162,168]
253. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 155,153]
255. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 163,159]
257. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 253,160]
260. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 255,154]
262. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 257,165]
263. ~select_1(X1) | select_3(X1) [subsumption resolution 262,242]
265. ~select_3(X1) | select_1(X1) [resolution 260,156]
266. select_1(sK8) <- {2} [resolution 265,238]
268. main_query_1(sK8) <- {2} [resolution 266,158]
269. $false <- {1, 2} [subsumption resolution 268,229]
270. 1 | ~2 [avatar contradiction clause 269]
271. select_1(sK8) <- {0} [resolution 226,157]
272. select_3(sK8) <- {0} [resolution 271,263]
275. main_query_2(sK8) <- {0} [resolution 272,167]
277. $false <- {0, 3} [subsumption resolution 275,235]
278. ~0 | 3 [avatar contradiction clause 277]
279. $false [avatar sat refutation 236,237,270,278]
% SZS output end Proof for query03_query04_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.012 s
% ------------------------------
% ------------------------------
