% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query01_tptp4X
% SZS output start Proof for query00_query01_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [input]
6. ! [X6] : (main_query_2(X6) <=> select_3(X6)) [input]
44. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
45. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 44]
46. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
47. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
48. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 5]
49. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 6]
87. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 45]
147. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 87]
148. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
149. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 46]
150. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 149]
151. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
152. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 150,151]
153. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 47]
154. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
155. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 48]
156. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 155]
157. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
158. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0))) [skolemisation 156,157]
159. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 49]
162. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 147]
163. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
164. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 162,163]
165. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 148]
166. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 148]
167. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 152]
168. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 152]
169. ~main_query_1(X0) | select_1(X0) [cnf transformation 153]
170. ~select_1(X0) | main_query_1(X0) [cnf transformation 153]
171. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 154]
172. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 154]
173. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 158]
174. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 158]
175. ~main_query_2(X0) | select_3(X0) [cnf transformation 159]
176. ~select_3(X0) | main_query_2(X0) [cnf transformation 159]
212. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 164]
213. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 164]
242. 0 <=> main_query_1(sK8) [avatar definition]
243. main_query_1(sK8) <- {0} [avatar component clause 242]
245. 1 <=> ~main_query_1(sK8) [avatar definition]
246. ~main_query_1(sK8) <- {1} [avatar component clause 245]
248. 2 <=> main_query_2(sK8) [avatar definition]
249. main_query_2(sK8) <- {2} [avatar component clause 248]
251. 3 <=> ~main_query_2(sK8) [avatar definition]
252. ~main_query_2(sK8) <- {3} [avatar component clause 251]
253. ~1 | ~3 [avatar split clause 213,251,245]
254. 0 | 2 [avatar split clause 212,248,242]
255. select_3(sK8) <- {2} [resolution 175,249]
263. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 167,165]
265. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 173,171]
267. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 263,172]
270. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 265,166]
272. ~select_1(X1) | select_3(X1) [resolution 267,174]
274. ~select_3(X1) | select_1(X1) [resolution 270,168]
275. select_1(sK8) <- {2} [resolution 274,255]
277. main_query_1(sK8) <- {2} [resolution 275,170]
278. $false <- {1, 2} [subsumption resolution 277,246]
279. 1 | ~2 [avatar contradiction clause 278]
280. select_1(sK8) <- {0} [resolution 243,169]
281. select_3(sK8) <- {0} [resolution 280,272]
284. main_query_2(sK8) <- {0} [resolution 281,176]
285. $false <- {0, 3} [subsumption resolution 284,252]
286. ~0 | 3 [avatar contradiction clause 285]
287. $false [avatar sat refutation 253,254,279,286]
% SZS output end Proof for query00_query01_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.011 s
% ------------------------------
% ------------------------------
