% Refutation found. Thanks to Tanya!
% SZS status Theorem for query06_query09_tptp4X
% SZS output start Proof for query06_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : ($true & from_5(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (subquery_from_3(X0) <=> select_4(X0)) [input]
4. ! [X0] : (from_2(X0) <=> subquery_from_3(X0)) [input]
5. ! [X0] : (select_1(X0) <=> ($true & from_2(X0))) [input]
6. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
7. ! [X5,X1,X2,X3,X4] : (from_10(X5,X1,X2,X3,X4) <=> tab1(X5,X1,X2,X3,X4)) [input]
8. ! [X5] : (select_9(X5) <=> ? [X1,X2,X3,X4] : ($true & from_10(X5,X1,X2,X3,X4))) [input]
9. ! [X5] : (subquery_from_8(X5) <=> select_9(X5)) [input]
10. ! [X5] : (from_7(X5) <=> subquery_from_8(X5)) [input]
11. ! [X5] : (select_6(X5) <=> ($true & from_7(X5))) [input]
12. ! [X5] : (main_query_2(X5) <=> select_6(X5)) [input]
46. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
47. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 46]
48. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0] : (select_1(X0) <=> from_2(X0)) [true and false elimination 5]
50. ! [X0,X1,X2,X3,X4] : (from_10(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 7]
51. ! [X0] : (select_9(X0) <=> ? [X1,X2,X3,X4] : ($true & from_10(X0,X1,X2,X3,X4))) [rectify 8]
52. ! [X0] : (select_9(X0) <=> ? [X1,X2,X3,X4] : from_10(X0,X1,X2,X3,X4)) [true and false elimination 51]
53. ! [X0] : (subquery_from_8(X0) <=> select_9(X0)) [rectify 9]
54. ! [X0] : (from_7(X0) <=> subquery_from_8(X0)) [rectify 10]
55. ! [X0] : (select_6(X0) <=> ($true & from_7(X0))) [rectify 11]
56. ! [X0] : (select_6(X0) <=> from_7(X0)) [true and false elimination 55]
57. ! [X0] : (main_query_2(X0) <=> select_6(X0)) [rectify 12]
142. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 47]
143. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 1]
144. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4) | ~select_4(X0))) [nnf transformation 48]
145. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) | ~select_4(X0))) [rectify 144]
146. ! [X0] : (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) => from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
147. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0))) [skolemisation 145,146]
148. ! [X0] : ((subquery_from_3(X0) | ~select_4(X0)) & (select_4(X0) | ~subquery_from_3(X0))) [nnf transformation 3]
149. ! [X0] : ((from_2(X0) | ~subquery_from_3(X0)) & (subquery_from_3(X0) | ~from_2(X0))) [nnf transformation 4]
150. ! [X0] : ((select_1(X0) | ~from_2(X0)) & (from_2(X0) | ~select_1(X0))) [nnf transformation 49]
151. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 6]
152. ! [X0,X1,X2,X3,X4] : ((from_10(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_10(X0,X1,X2,X3,X4))) [nnf transformation 50]
153. ! [X0] : ((select_9(X0) | ! [X1,X2,X3,X4] : ~from_10(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_10(X0,X1,X2,X3,X4) | ~select_9(X0))) [nnf transformation 52]
154. ! [X0] : ((select_9(X0) | ! [X1,X2,X3,X4] : ~from_10(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_10(X0,X5,X6,X7,X8) | ~select_9(X0))) [rectify 153]
155. ! [X0] : (? [X5,X6,X7,X8] : from_10(X0,X5,X6,X7,X8) => from_10(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
156. ! [X0] : ((select_9(X0) | ! [X1,X2,X3,X4] : ~from_10(X0,X1,X2,X3,X4)) & (from_10(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_9(X0))) [skolemisation 154,155]
157. ! [X0] : ((subquery_from_8(X0) | ~select_9(X0)) & (select_9(X0) | ~subquery_from_8(X0))) [nnf transformation 53]
158. ! [X0] : ((from_7(X0) | ~subquery_from_8(X0)) & (subquery_from_8(X0) | ~from_7(X0))) [nnf transformation 54]
159. ! [X0] : ((select_6(X0) | ~from_7(X0)) & (from_7(X0) | ~select_6(X0))) [nnf transformation 56]
160. ! [X0] : ((main_query_2(X0) | ~select_6(X0)) & (select_6(X0) | ~main_query_2(X0))) [nnf transformation 57]
163. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 142]
164. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
165. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 163,164]
166. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 143]
167. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 143]
168. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [cnf transformation 147]
169. ~from_5(X0,X1,X2,X3,X4) | select_4(X0) [cnf transformation 147]
170. ~subquery_from_3(X0) | select_4(X0) [cnf transformation 148]
171. ~select_4(X0) | subquery_from_3(X0) [cnf transformation 148]
172. ~from_2(X0) | subquery_from_3(X0) [cnf transformation 149]
173. ~subquery_from_3(X0) | from_2(X0) [cnf transformation 149]
174. ~select_1(X0) | from_2(X0) [cnf transformation 150]
175. ~from_2(X0) | select_1(X0) [cnf transformation 150]
176. ~main_query_1(X0) | select_1(X0) [cnf transformation 151]
177. ~select_1(X0) | main_query_1(X0) [cnf transformation 151]
178. ~from_10(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 152]
179. ~tab1(X0,X1,X2,X3,X4) | from_10(X0,X1,X2,X3,X4) [cnf transformation 152]
180. from_10(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_9(X0) [cnf transformation 156]
181. ~from_10(X0,X1,X2,X3,X4) | select_9(X0) [cnf transformation 156]
182. ~subquery_from_8(X0) | select_9(X0) [cnf transformation 157]
183. ~select_9(X0) | subquery_from_8(X0) [cnf transformation 157]
184. ~from_7(X0) | subquery_from_8(X0) [cnf transformation 158]
185. ~subquery_from_8(X0) | from_7(X0) [cnf transformation 158]
186. ~select_6(X0) | from_7(X0) [cnf transformation 159]
187. ~from_7(X0) | select_6(X0) [cnf transformation 159]
188. ~main_query_2(X0) | select_6(X0) [cnf transformation 160]
189. ~select_6(X0) | main_query_2(X0) [cnf transformation 160]
221. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 165]
222. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 165]
247. 0 <=> main_query_1(sK8) [avatar definition]
248. main_query_1(sK8) <- {0} [avatar component clause 247]
250. 1 <=> ~main_query_1(sK8) [avatar definition]
251. ~main_query_1(sK8) <- {1} [avatar component clause 250]
253. 2 <=> main_query_2(sK8) [avatar definition]
254. main_query_2(sK8) <- {2} [avatar component clause 253]
256. 3 <=> ~main_query_2(sK8) [avatar definition]
257. ~main_query_2(sK8) <- {3} [avatar component clause 256]
258. ~1 | ~3 [avatar split clause 222,256,250]
259. 0 | 2 [avatar split clause 221,253,247]
260. select_6(sK8) <- {2} [resolution 188,254]
262. from_7(sK8) <- {2} [resolution 260,186]
265. subquery_from_8(sK8) <- {2} [resolution 262,184]
268. select_9(sK8) <- {2} [resolution 265,182]
274. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [resolution 168,166]
276. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_9(X0) [resolution 180,178]
278. from_10(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [resolution 274,179]
281. from_5(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_9(X1) [resolution 276,167]
283. ~select_4(X1) | select_9(X1) [resolution 278,181]
285. ~select_9(X1) | select_4(X1) [resolution 281,169]
286. select_4(sK8) <- {2} [resolution 285,268]
288. subquery_from_3(sK8) <- {2} [resolution 286,171]
289. from_2(sK8) <- {2} [resolution 288,173]
291. select_1(sK8) <- {2} [resolution 289,175]
293. main_query_1(sK8) <- {2} [resolution 291,177]
295. $false <- {1, 2} [subsumption resolution 293,251]
296. 1 | ~2 [avatar contradiction clause 295]
297. select_1(sK8) <- {0} [resolution 248,176]
299. from_2(sK8) <- {0} [resolution 297,174]
301. subquery_from_3(sK8) <- {0} [resolution 299,172]
303. select_4(sK8) <- {0} [resolution 301,170]
304. select_9(sK8) <- {0} [resolution 303,283]
307. subquery_from_8(sK8) <- {0} [resolution 304,183]
308. from_7(sK8) <- {0} [resolution 307,185]
310. select_6(sK8) <- {0} [resolution 308,187]
312. main_query_2(sK8) <- {0} [resolution 310,189]
314. $false <- {0, 3} [subsumption resolution 312,257]
315. ~0 | 3 [avatar contradiction clause 314]
316. $false [avatar sat refutation 258,259,296,315]
% SZS output end Proof for query06_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.012 s
% ------------------------------
% ------------------------------
