% Refutation found. Thanks to Tanya!
% SZS status Theorem for query20_query80_tptp4X
% SZS output start Proof for query20_query80_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
3. ! [X5,X6] : (where_9(X5,X6) <=> eq(X5,X6)) [input]
4. ! [X5,X6] : (select_7(X5,X6) <=> ? [X0,X1,X2,X3,X4] : (where_9(X5,X6) & from_8(X0,X1,X2,X3,X4))) [input]
5. ! [X5,X6] : (query_6(X5,X6) <=> select_7(X5,X6)) [input]
6. ! [X7,X8] : (from_10(X7,X8) <=> query_6(X7,X8)) [input]
7. ! [X7] : (select_5(X7) <=> ? [X8] : ($true & from_10(X7,X8))) [input]
8. ! [X5] : (query_4(X5) <=> select_5(X5)) [input]
9. ! [X5] : (where_3(X5) <=> query_4(X5)) [input]
10. ! [X5] : (select_1(X5) <=> ? [X0,X1,X2,X3,X4] : (where_3(X5) & from_2(X0,X1,X2,X3,X4))) [input]
11. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
12. ! [X9,X10,X11,X12,X13] : (from_12(X9,X10,X11,X12,X13) <=> tab1(X9,X10,X11,X12,X13)) [input]
13. ! [X0] : (select_11(X0) <=> ? [X9,X10,X11,X12,X13] : ($true & from_12(X9,X10,X11,X12,X13))) [input]
14. ! [X14] : (main_query_2(X14) <=> select_11(X14)) [input]
15. ! [X15] : eq(X15,X15) [input]
61. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
62. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 61]
63. ! [X0,X1] : (where_9(X0,X1) <=> eq(X0,X1)) [rectify 3]
64. ! [X0,X1] : (select_7(X0,X1) <=> ? [X2,X3,X4,X5,X6] : (where_9(X0,X1) & from_8(X2,X3,X4,X5,X6))) [rectify 4]
65. ! [X0,X1] : (query_6(X0,X1) <=> select_7(X0,X1)) [rectify 5]
66. ! [X0,X1] : (from_10(X0,X1) <=> query_6(X0,X1)) [rectify 6]
67. ! [X0] : (select_5(X0) <=> ? [X1] : ($true & from_10(X0,X1))) [rectify 7]
68. ! [X0] : (select_5(X0) <=> ? [X1] : from_10(X0,X1)) [true and false elimination 67]
69. ! [X0] : (query_4(X0) <=> select_5(X0)) [rectify 8]
70. ! [X0] : (where_3(X0) <=> query_4(X0)) [rectify 9]
71. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : (where_3(X0) & from_2(X1,X2,X3,X4,X5))) [rectify 10]
72. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 11]
73. ! [X0,X1,X2,X3,X4] : (from_12(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 12]
74. ! [X0] : (select_11(X0) <=> ? [X1,X2,X3,X4,X5] : ($true & from_12(X1,X2,X3,X4,X5))) [rectify 13]
75. ! [X0] : (select_11(X0) <=> ? [X1,X2,X3,X4,X5] : from_12(X1,X2,X3,X4,X5)) [true and false elimination 74]
76. ! [X0] : (main_query_2(X0) <=> select_11(X0)) [rectify 14]
77. ! [X0] : eq(X0,X0) [rectify 15]
123. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 62]
201. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 123]
202. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
203. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 2]
204. ! [X0,X1] : ((where_9(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_9(X0,X1))) [nnf transformation 63]
205. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_9(X0,X1) | ~from_8(X2,X3,X4,X5,X6))) & (? [X2,X3,X4,X5,X6] : (where_9(X0,X1) & from_8(X2,X3,X4,X5,X6)) | ~select_7(X0,X1))) [nnf transformation 64]
206. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_9(X0,X1) | ~from_8(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_9(X0,X1) & from_8(X7,X8,X9,X10,X11)) | ~select_7(X0,X1))) [rectify 205]
207. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_9(X0,X1) & from_8(X7,X8,X9,X10,X11)) => (where_9(X0,X1) & from_8(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)))) [choice axiom]
208. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_9(X0,X1) | ~from_8(X2,X3,X4,X5,X6))) & ((where_9(X0,X1) & from_8(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1))) | ~select_7(X0,X1))) [skolemisation 206,207]
209. ! [X0,X1] : ((query_6(X0,X1) | ~select_7(X0,X1)) & (select_7(X0,X1) | ~query_6(X0,X1))) [nnf transformation 65]
210. ! [X0,X1] : ((from_10(X0,X1) | ~query_6(X0,X1)) & (query_6(X0,X1) | ~from_10(X0,X1))) [nnf transformation 66]
211. ! [X0] : ((select_5(X0) | ! [X1] : ~from_10(X0,X1)) & (? [X1] : from_10(X0,X1) | ~select_5(X0))) [nnf transformation 68]
212. ! [X0] : ((select_5(X0) | ! [X1] : ~from_10(X0,X1)) & (? [X2] : from_10(X0,X2) | ~select_5(X0))) [rectify 211]
213. ! [X0] : (? [X2] : from_10(X0,X2) => from_10(X0,sK5(X0))) [choice axiom]
214. ! [X0] : ((select_5(X0) | ! [X1] : ~from_10(X0,X1)) & (from_10(X0,sK5(X0)) | ~select_5(X0))) [skolemisation 212,213]
215. ! [X0] : ((query_4(X0) | ~select_5(X0)) & (select_5(X0) | ~query_4(X0))) [nnf transformation 69]
216. ! [X0] : ((where_3(X0) | ~query_4(X0)) & (query_4(X0) | ~where_3(X0))) [nnf transformation 70]
217. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : (~where_3(X0) | ~from_2(X1,X2,X3,X4,X5))) & (? [X1,X2,X3,X4,X5] : (where_3(X0) & from_2(X1,X2,X3,X4,X5)) | ~select_1(X0))) [nnf transformation 71]
218. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : (~where_3(X0) | ~from_2(X1,X2,X3,X4,X5))) & (? [X6,X7,X8,X9,X10] : (where_3(X0) & from_2(X6,X7,X8,X9,X10)) | ~select_1(X0))) [rectify 217]
219. ! [X0] : (? [X6,X7,X8,X9,X10] : (where_3(X0) & from_2(X6,X7,X8,X9,X10)) => (where_3(X0) & from_2(sK6(X0),sK7(X0),sK8(X0),sK9(X0),sK10(X0)))) [choice axiom]
220. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : (~where_3(X0) | ~from_2(X1,X2,X3,X4,X5))) & ((where_3(X0) & from_2(sK6(X0),sK7(X0),sK8(X0),sK9(X0),sK10(X0))) | ~select_1(X0))) [skolemisation 218,219]
221. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 72]
222. ! [X0,X1,X2,X3,X4] : ((from_12(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_12(X0,X1,X2,X3,X4))) [nnf transformation 73]
223. ! [X0] : ((select_11(X0) | ! [X1,X2,X3,X4,X5] : ~from_12(X1,X2,X3,X4,X5)) & (? [X1,X2,X3,X4,X5] : from_12(X1,X2,X3,X4,X5) | ~select_11(X0))) [nnf transformation 75]
224. ! [X0] : ((select_11(X0) | ! [X1,X2,X3,X4,X5] : ~from_12(X1,X2,X3,X4,X5)) & (? [X6,X7,X8,X9,X10] : from_12(X6,X7,X8,X9,X10) | ~select_11(X0))) [rectify 223]
225. ? [X6,X7,X8,X9,X10] : from_12(X6,X7,X8,X9,X10) => from_12(sK11,sK12,sK13,sK14,sK15) [choice axiom]
226. ! [X0] : ((select_11(X0) | ! [X1,X2,X3,X4,X5] : ~from_12(X1,X2,X3,X4,X5)) & (from_12(sK11,sK12,sK13,sK14,sK15) | ~select_11(X0))) [skolemisation 224,225]
227. ! [X0] : ((main_query_2(X0) | ~select_11(X0)) & (select_11(X0) | ~main_query_2(X0))) [nnf transformation 76]
230. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 201]
231. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK16) | ~main_query_1(sK16)) & (main_query_2(sK16) | main_query_1(sK16))) [choice axiom]
232. (~main_query_2(sK16) | ~main_query_1(sK16)) & (main_query_2(sK16) | main_query_1(sK16)) [skolemisation 230,231]
233. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 202]
234. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 202]
236. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 203]
238. ~eq(X0,X1) | where_9(X0,X1) [cnf transformation 204]
241. select_7(X0,X1) | ~where_9(X0,X1) | ~from_8(X2,X3,X4,X5,X6) [cnf transformation 208]
243. ~select_7(X0,X1) | query_6(X0,X1) [cnf transformation 209]
245. ~query_6(X0,X1) | from_10(X0,X1) [cnf transformation 210]
247. ~from_10(X0,X1) | select_5(X0) [cnf transformation 214]
249. ~select_5(X0) | query_4(X0) [cnf transformation 215]
251. ~query_4(X0) | where_3(X0) [cnf transformation 216]
252. from_2(sK6(X0),sK7(X0),sK8(X0),sK9(X0),sK10(X0)) | ~select_1(X0) [cnf transformation 220]
254. select_1(X0) | ~where_3(X0) | ~from_2(X1,X2,X3,X4,X5) [cnf transformation 220]
255. ~main_query_1(X0) | select_1(X0) [cnf transformation 221]
256. ~select_1(X0) | main_query_1(X0) [cnf transformation 221]
257. ~from_12(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 222]
258. ~tab1(X0,X1,X2,X3,X4) | from_12(X0,X1,X2,X3,X4) [cnf transformation 222]
259. from_12(sK11,sK12,sK13,sK14,sK15) | ~select_11(X0) [cnf transformation 226]
260. select_11(X0) | ~from_12(X1,X2,X3,X4,X5) [cnf transformation 226]
261. ~main_query_2(X0) | select_11(X0) [cnf transformation 227]
262. ~select_11(X0) | main_query_2(X0) [cnf transformation 227]
263. eq(X0,X0) [cnf transformation 77]
307. main_query_2(sK16) | main_query_1(sK16) [cnf transformation 232]
308. ~main_query_2(sK16) | ~main_query_1(sK16) [cnf transformation 232]
346. 0 <=> ! [X3,X5,X2,X4,X6] : ~from_8(X2,X3,X4,X5,X6) [avatar definition]
347. ~from_8(X2,X3,X4,X5,X6) <- {0} [avatar component clause 346]
349. 2 <=> ! [X1,X0] : (select_7(X0,X1) | ~where_9(X0,X1)) [avatar definition]
350. ~where_9(X0,X1) | select_7(X0,X1) <- {2} [avatar component clause 349]
351. 0 | 2 [avatar split clause 241,349,346]
353. 4 <=> ! [X1,X3,X5,X2,X4] : ~from_2(X1,X2,X3,X4,X5) [avatar definition]
354. ~from_2(X1,X2,X3,X4,X5) <- {4} [avatar component clause 353]
356. 6 <=> ! [X0] : (select_1(X0) | ~where_3(X0)) [avatar definition]
357. ~where_3(X0) | select_1(X0) <- {6} [avatar component clause 356]
358. 4 | 6 [avatar split clause 254,356,353]
360. 8 <=> ! [X1,X3,X5,X2,X4] : ~from_12(X1,X2,X3,X4,X5) [avatar definition]
361. ~from_12(X1,X2,X3,X4,X5) <- {8} [avatar component clause 360]
363. 10 <=> ! [X0] : select_11(X0) [avatar definition]
364. select_11(X0) <- {10} [avatar component clause 363]
365. 8 | 10 [avatar split clause 260,363,360]
367. 12 <=> ! [X0] : ~select_11(X0) [avatar definition]
368. ~select_11(X0) <- {12} [avatar component clause 367]
373. 14 <=> from_12(sK11,sK12,sK13,sK14,sK15) [avatar definition]
374. from_12(sK11,sK12,sK13,sK14,sK15) <- {14} [avatar component clause 373]
375. 12 | 14 [avatar split clause 259,373,367]
377. 16 <=> main_query_1(sK16) [avatar definition]
378. main_query_1(sK16) <- {16} [avatar component clause 377]
380. 17 <=> ~main_query_1(sK16) [avatar definition]
381. ~main_query_1(sK16) <- {17} [avatar component clause 380]
383. 18 <=> main_query_2(sK16) [avatar definition]
384. main_query_2(sK16) <- {18} [avatar component clause 383]
386. 19 <=> ~main_query_2(sK16) [avatar definition]
387. ~main_query_2(sK16) <- {19} [avatar component clause 386]
388. ~17 | ~19 [avatar split clause 308,386,380]
389. 16 | 18 [avatar split clause 307,383,377]
397. where_9(X0,X0) [resolution 238,263]
434. tab1(sK6(X0),sK7(X0),sK8(X0),sK9(X0),sK10(X0)) | ~select_1(X0) [resolution 252,233]
438. select_11(sK16) <- {18} [resolution 384,261]
439. $false <- {12, 18} [subsumption resolution 438,368]
440. ~12 | ~18 [avatar contradiction clause 439]
444. select_1(sK16) <- {16} [resolution 378,255]
449. from_12(sK6(X0),sK7(X0),sK8(X0),sK9(X0),sK10(X0)) | ~select_1(X0) [resolution 434,258]
456. ~select_1(X0) <- {8} [subsumption resolution 449,361]
457. $false <- {8, 16} [subsumption resolution 444,456]
458. ~8 | ~16 [avatar contradiction clause 457]
461. main_query_2(X0) <- {10} [resolution 364,262]
464. tab1(sK11,sK12,sK13,sK14,sK15) <- {14} [resolution 374,257]
466. from_8(sK11,sK12,sK13,sK14,sK15) <- {14} [resolution 464,236]
467. from_2(sK11,sK12,sK13,sK14,sK15) <- {14} [resolution 464,234]
468. $false <- {0, 14} [subsumption resolution 466,347]
469. ~0 | ~14 [avatar contradiction clause 468]
470. select_7(X0,X0) <- {2} [resolution 350,397]
471. query_6(X0,X0) <- {2} [resolution 470,243]
476. from_10(X0,X0) <- {2} [resolution 471,245]
481. select_5(X1) <- {2} [resolution 476,247]
483. query_4(X0) <- {2} [resolution 481,249]
485. where_3(X0) <- {2} [resolution 483,251]
488. select_1(X0) <- {2, 6} [resolution 485,357]
491. main_query_1(X0) <- {2, 6} [resolution 488,256]
495. $false <- {2, 6, 17} [resolution 491,381]
496. ~2 | ~6 | 17 [avatar contradiction clause 495]
498. $false <- {4, 14} [resolution 354,467]
499. ~4 | ~14 [avatar contradiction clause 498]
500. $false <- {10, 19} [subsumption resolution 387,461]
501. ~10 | 19 [avatar contradiction clause 500]
503. $false [avatar sat refutation 351,358,365,375,388,389,440,458,469,496,499,501]
% SZS output end Proof for query20_query80_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.014 s
% ------------------------------
% ------------------------------
