% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_13_query_16_tptp4X
% SZS output start Proof for query_13_query_16_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1] : (where_9(X0,X1) <=> eq(X0,X1)) [input]
7. ! [X0,X1] : (select_7(X0,X1) <=> ? [X2,X3,X4] : (where_9(X0,X1) & from_8(X0,X1,X2,X3,X4))) [input]
8. ! [X5,X6] : (subquery_from_6(X5,X6) <=> select_7(X5,X6)) [input]
9. ! [X7,X8] : (from_5(X7,X8) <=> subquery_from_6(X7,X8)) [input]
10. ! [X7,X8] : (select_4(X7,X8) <=> ($true & from_5(X7,X8))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
16. ! [X9,X10] : ((lte(X10,X9) & lte(X9,X10)) <=> eq(X9,X10)) [input]
60. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
61. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 60]
62. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
63. ! [X0,X1] : (subquery_from_6(X0,X1) <=> select_7(X0,X1)) [rectify 8]
64. ! [X0,X1] : (from_5(X0,X1) <=> subquery_from_6(X0,X1)) [rectify 9]
65. ! [X0,X1] : (select_4(X0,X1) <=> ($true & from_5(X0,X1))) [rectify 10]
66. ! [X0,X1] : (select_4(X0,X1) <=> from_5(X0,X1)) [true and false elimination 65]
67. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
72. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 16]
116. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 61]
198. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 116]
199. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
200. ! [X0,X1] : ((where_3(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [nnf transformation 2]
201. ! [X0,X1] : ((where_3(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [flattening 200]
202. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
203. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 202]
204. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
205. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 203,204]
206. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 62]
207. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 5]
208. ! [X0,X1] : ((where_9(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_9(X0,X1))) [nnf transformation 6]
209. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : (~where_9(X0,X1) | ~from_8(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_9(X0,X1) & from_8(X0,X1,X2,X3,X4)) | ~select_7(X0,X1))) [nnf transformation 7]
210. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : (~where_9(X0,X1) | ~from_8(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_9(X0,X1) & from_8(X0,X1,X5,X6,X7)) | ~select_7(X0,X1))) [rectify 209]
211. ! [X1,X0] : (? [X5,X6,X7] : (where_9(X0,X1) & from_8(X0,X1,X5,X6,X7)) => (where_9(X0,X1) & from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
212. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : (~where_9(X0,X1) | ~from_8(X0,X1,X2,X3,X4))) & ((where_9(X0,X1) & from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_7(X0,X1))) [skolemisation 210,211]
213. ! [X0,X1] : ((subquery_from_6(X0,X1) | ~select_7(X0,X1)) & (select_7(X0,X1) | ~subquery_from_6(X0,X1))) [nnf transformation 63]
214. ! [X0,X1] : ((from_5(X0,X1) | ~subquery_from_6(X0,X1)) & (subquery_from_6(X0,X1) | ~from_5(X0,X1))) [nnf transformation 64]
215. ! [X0,X1] : ((select_4(X0,X1) | ~from_5(X0,X1)) & (from_5(X0,X1) | ~select_4(X0,X1))) [nnf transformation 66]
216. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 67]
217. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 72]
218. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 217]
219. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 198]
220. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
221. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 219,220]
222. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 199]
223. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 199]
224. ~where_3(X0,X1) | lte(X0,X1) [cnf transformation 201]
225. ~where_3(X0,X1) | lte(X1,X0) [cnf transformation 201]
226. ~lte(X1,X0) | where_3(X0,X1) | ~lte(X0,X1) [cnf transformation 201]
227. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 205]
228. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 205]
229. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 205]
230. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 206]
231. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 206]
232. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 207]
233. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 207]
234. ~where_9(X0,X1) | eq(X0,X1) [cnf transformation 208]
235. ~eq(X0,X1) | where_9(X0,X1) [cnf transformation 208]
236. from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1) [cnf transformation 212]
237. ~select_7(X0,X1) | where_9(X0,X1) [cnf transformation 212]
238. ~from_8(X0,X1,X2,X3,X4) | ~where_9(X0,X1) | select_7(X0,X1) [cnf transformation 212]
239. ~subquery_from_6(X0,X1) | select_7(X0,X1) [cnf transformation 213]
240. ~select_7(X0,X1) | subquery_from_6(X0,X1) [cnf transformation 213]
241. ~from_5(X0,X1) | subquery_from_6(X0,X1) [cnf transformation 214]
242. ~subquery_from_6(X0,X1) | from_5(X0,X1) [cnf transformation 214]
243. ~select_4(X0,X1) | from_5(X0,X1) [cnf transformation 215]
244. ~from_5(X0,X1) | select_4(X0,X1) [cnf transformation 215]
245. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 216]
246. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 216]
251. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 218]
252. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 218]
253. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 218]
293. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 221]
294. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 221]
296. 0 <=> main_query_1(sK6,sK7) [avatar definition]
297. main_query_1(sK6,sK7) <- {0} [avatar component clause 296]
299. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
300. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 299]
302. 2 <=> main_query_2(sK6,sK7) [avatar definition]
303. main_query_2(sK6,sK7) <- {2} [avatar component clause 302]
305. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
306. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 305]
307. ~1 | ~3 [avatar split clause 294,305,299]
308. 0 | 2 [avatar split clause 293,302,296]
311. select_4(sK6,sK7) <- {2} [resolution 245,303]
312. from_5(sK6,sK7) <- {2} [resolution 311,243]
315. subquery_from_6(sK6,sK7) <- {2} [resolution 312,241]
318. select_7(sK6,sK7) <- {2} [resolution 315,239]
321. where_9(sK6,sK7) <- {2} [resolution 318,237]
323. eq(sK6,sK7) <- {2} [resolution 321,234]
330. lte(sK7,sK6) <- {2} [resolution 323,253]
331. lte(sK6,sK7) <- {2} [resolution 323,252]
336. where_3(sK6,sK7) | ~lte(sK6,sK7) <- {2} [resolution 330,226]
337. where_3(sK6,sK7) <- {2} [subsumption resolution 336,331]
447. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 227,222]
451. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_7(X2,X3) [resolution 236,232]
499. from_8(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 447,233]
509. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1) [resolution 451,223]
606. ~where_9(X2,X3) | ~select_1(X2,X3) | select_7(X2,X3) [resolution 499,238]
615. ~select_7(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 509,229]
806. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 615,318]
810. select_1(sK6,sK7) <- {2} [subsumption resolution 806,337]
817. main_query_1(sK6,sK7) <- {2} [resolution 810,231]
819. $false <- {1, 2} [subsumption resolution 817,300]
820. 1 | ~2 [avatar contradiction clause 819]
821. select_1(sK6,sK7) <- {0} [resolution 297,230]
823. where_3(sK6,sK7) <- {0} [resolution 821,228]
824. lte(sK7,sK6) <- {0} [resolution 823,225]
825. lte(sK6,sK7) <- {0} [resolution 823,224]
834. eq(sK6,sK7) | ~lte(sK6,sK7) <- {0} [resolution 824,251]
836. eq(sK6,sK7) <- {0} [subsumption resolution 834,825]
844. where_9(sK6,sK7) <- {0} [resolution 836,235]
895. ~select_1(sK6,sK7) | select_7(sK6,sK7) <- {0} [resolution 844,606]
897. select_7(sK6,sK7) <- {0} [subsumption resolution 895,821]
924. subquery_from_6(sK6,sK7) <- {0} [resolution 897,240]
946. from_5(sK6,sK7) <- {0} [resolution 924,242]
957. select_4(sK6,sK7) <- {0} [resolution 946,244]
962. main_query_2(sK6,sK7) <- {0} [resolution 957,246]
964. $false <- {0, 3} [subsumption resolution 962,306]
965. ~0 | 3 [avatar contradiction clause 964]
966. $false [avatar sat refutation 307,308,820,965]
% SZS output end Proof for query_13_query_16_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.030 s
% ------------------------------
% ------------------------------
