% Refutation found. Thanks to Tanya!
% SZS status Theorem for query04_query06_tptp4X
% SZS output start Proof for query04_query06_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X5,X6,X7,X8,X9] : (from_8(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
6. ! [X5] : (select_7(X5) <=> ? [X6,X7,X8,X9] : ($true & from_8(X5,X6,X7,X8,X9))) [input]
7. ! [X10] : (subquery_from_6(X10) <=> select_7(X10)) [input]
8. ! [X11] : (from_5(X11) <=> subquery_from_6(X11)) [input]
9. ! [X11] : (select_4(X11) <=> ($true & from_5(X11))) [input]
10. ! [X10] : (main_query_2(X10) <=> select_4(X10)) [input]
11. ! [X12] : eq(X12,X12) [input]
45. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
46. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 45]
47. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
48. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
49. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 48]
50. ! [X0] : (subquery_from_6(X0) <=> select_7(X0)) [rectify 7]
51. ! [X0] : (from_5(X0) <=> subquery_from_6(X0)) [rectify 8]
52. ! [X0] : (select_4(X0) <=> ($true & from_5(X0))) [rectify 9]
53. ! [X0] : (select_4(X0) <=> from_5(X0)) [true and false elimination 52]
54. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 10]
55. ! [X0] : eq(X0,X0) [rectify 11]
140. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 46]
141. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
142. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
143. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
144. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 143]
145. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
146. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 144,145]
147. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
148. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 47]
149. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 49]
150. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 149]
151. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
152. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0))) [skolemisation 150,151]
153. ! [X0] : ((subquery_from_6(X0) | ~select_7(X0)) & (select_7(X0) | ~subquery_from_6(X0))) [nnf transformation 50]
154. ! [X0] : ((from_5(X0) | ~subquery_from_6(X0)) & (subquery_from_6(X0) | ~from_5(X0))) [nnf transformation 51]
155. ! [X0] : ((select_4(X0) | ~from_5(X0)) & (from_5(X0) | ~select_4(X0))) [nnf transformation 53]
156. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 54]
159. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 140]
160. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
161. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 159,160]
162. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 141]
163. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 141]
165. ~eq(X0,X0) | where_3(X0) [cnf transformation 142]
166. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 146]
168. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 146]
169. ~main_query_1(X0) | select_1(X0) [cnf transformation 147]
170. ~select_1(X0) | main_query_1(X0) [cnf transformation 147]
171. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 148]
172. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 148]
173. from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [cnf transformation 152]
174. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 152]
175. ~subquery_from_6(X0) | select_7(X0) [cnf transformation 153]
176. ~select_7(X0) | subquery_from_6(X0) [cnf transformation 153]
177. ~from_5(X0) | subquery_from_6(X0) [cnf transformation 154]
178. ~subquery_from_6(X0) | from_5(X0) [cnf transformation 154]
179. ~select_4(X0) | from_5(X0) [cnf transformation 155]
180. ~from_5(X0) | select_4(X0) [cnf transformation 155]
181. ~main_query_2(X0) | select_4(X0) [cnf transformation 156]
182. ~select_4(X0) | main_query_2(X0) [cnf transformation 156]
183. eq(X0,X0) [cnf transformation 55]
215. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 161]
216. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 161]
218. 0 <=> main_query_1(sK8) [avatar definition]
219. main_query_1(sK8) <- {0} [avatar component clause 218]
221. 1 <=> ~main_query_1(sK8) [avatar definition]
222. ~main_query_1(sK8) <- {1} [avatar component clause 221]
224. 2 <=> main_query_2(sK8) [avatar definition]
225. main_query_2(sK8) <- {2} [avatar component clause 224]
227. 3 <=> ~main_query_2(sK8) [avatar definition]
228. ~main_query_2(sK8) <- {3} [avatar component clause 227]
229. ~1 | ~3 [avatar split clause 216,227,221]
230. 0 | 2 [avatar split clause 215,224,218]
231. select_4(sK8) <- {2} [resolution 181,225]
232. from_5(sK8) <- {2} [resolution 231,179]
235. subquery_from_6(sK8) <- {2} [resolution 232,177]
237. select_7(sK8) <- {2} [resolution 235,175]
239. where_3(X1) [resolution 165,183]
274. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 166,162]
276. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [resolution 173,171]
278. from_8(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 274,172]
283. from_2(X3,sK4(X3),sK5(X3),sK6(X3),sK7(X3)) | ~select_7(X3) [resolution 276,163]
287. ~select_1(X1) | select_7(X1) [resolution 278,174]
293. ~select_7(X1) | ~where_3(X1) | select_1(X1) [resolution 283,168]
294. ~select_7(X1) | select_1(X1) [subsumption resolution 293,239]
297. select_1(sK8) <- {2} [resolution 294,237]
301. main_query_1(sK8) <- {2} [resolution 297,170]
303. $false <- {1, 2} [subsumption resolution 301,222]
304. 1 | ~2 [avatar contradiction clause 303]
305. select_1(sK8) <- {0} [resolution 219,169]
308. select_7(sK8) <- {0} [resolution 305,287]
314. subquery_from_6(sK8) <- {0} [resolution 308,176]
317. from_5(sK8) <- {0} [resolution 314,178]
321. select_4(sK8) <- {0} [resolution 317,180]
323. main_query_2(sK8) <- {0} [resolution 321,182]
325. $false <- {0, 3} [subsumption resolution 323,228]
326. ~0 | 3 [avatar contradiction clause 325]
327. $false [avatar sat refutation 229,230,304,326]
% SZS output end Proof for query04_query06_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.010 s
% ------------------------------
% ------------------------------
