% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query05_tptp4X
% SZS output start Proof for query00_query05_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [input]
6. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
8. ! [X6] : eq(X6,X6) [input]
46. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
47. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 46]
48. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
50. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
51. ! [X0] : eq(X0,X0) [rectify 8]
89. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 47]
151. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 89]
152. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
153. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 48]
154. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 153]
155. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
156. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 154,155]
157. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 49]
158. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
159. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 5]
160. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 6]
161. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 160]
162. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
163. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 161,162]
164. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 50]
167. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 151]
168. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
169. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 167,168]
170. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 152]
171. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 152]
172. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 156]
173. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 156]
174. ~main_query_1(X0) | select_1(X0) [cnf transformation 157]
175. ~select_1(X0) | main_query_1(X0) [cnf transformation 157]
176. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 158]
177. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 158]
179. ~eq(X0,X0) | where_5(X0) [cnf transformation 159]
180. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 163]
182. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 163]
183. ~main_query_2(X0) | select_3(X0) [cnf transformation 164]
184. ~select_3(X0) | main_query_2(X0) [cnf transformation 164]
185. eq(X0,X0) [cnf transformation 51]
221. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 169]
222. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 169]
252. 0 <=> main_query_1(sK8) [avatar definition]
253. main_query_1(sK8) <- {0} [avatar component clause 252]
255. 1 <=> ~main_query_1(sK8) [avatar definition]
256. ~main_query_1(sK8) <- {1} [avatar component clause 255]
258. 2 <=> main_query_2(sK8) [avatar definition]
259. main_query_2(sK8) <- {2} [avatar component clause 258]
261. 3 <=> ~main_query_2(sK8) [avatar definition]
262. ~main_query_2(sK8) <- {3} [avatar component clause 261]
263. ~1 | ~3 [avatar split clause 222,261,255]
264. 0 | 2 [avatar split clause 221,258,252]
265. select_3(sK8) <- {2} [resolution 183,259]
269. where_5(X1) [resolution 179,185]
280. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 172,170]
282. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 180,176]
284. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 280,177]
287. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 282,171]
289. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 284,182]
290. ~select_1(X1) | select_3(X1) [subsumption resolution 289,269]
292. ~select_3(X1) | select_1(X1) [resolution 287,173]
293. select_1(sK8) <- {2} [resolution 292,265]
295. main_query_1(sK8) <- {2} [resolution 293,175]
296. $false <- {1, 2} [subsumption resolution 295,256]
297. 1 | ~2 [avatar contradiction clause 296]
298. select_1(sK8) <- {0} [resolution 253,174]
299. select_3(sK8) <- {0} [resolution 298,290]
302. main_query_2(sK8) <- {0} [resolution 299,184]
304. $false <- {0, 3} [subsumption resolution 302,262]
305. ~0 | 3 [avatar contradiction clause 304]
306. $false [avatar sat refutation 263,264,297,305]
% SZS output end Proof for query00_query05_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
