% Refutation found. Thanks to Tanya!
% SZS status Theorem for query05_query09_tptp4X
% SZS output start Proof for query05_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X6,X7,X8,X9,X10] : (from_8(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
6. ! [X6] : (select_7(X6) <=> ? [X7,X8,X9,X10] : ($true & from_8(X6,X7,X8,X9,X10))) [input]
7. ! [X6] : (subquery_from_6(X6) <=> select_7(X6)) [input]
8. ! [X11] : (from_5(X11) <=> subquery_from_6(X11)) [input]
9. ! [X11] : (select_4(X11) <=> ($true & from_5(X11))) [input]
10. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
11. ! [X12] : eq(X12,X12) [input]
42. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
43. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 42]
44. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
45. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
46. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
47. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 46]
48. ! [X0] : (subquery_from_6(X0) <=> select_7(X0)) [rectify 7]
49. ! [X0] : (from_5(X0) <=> subquery_from_6(X0)) [rectify 8]
50. ! [X0] : (select_4(X0) <=> ($true & from_5(X0))) [rectify 9]
51. ! [X0] : (select_4(X0) <=> from_5(X0)) [true and false elimination 50]
52. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 10]
53. ! [X0] : eq(X0,X0) [rectify 11]
84. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 43]
130. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 84]
131. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
132. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
133. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 133]
135. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 134,135]
137. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 44]
138. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 45]
139. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 47]
140. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 139]
141. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
142. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0))) [skolemisation 140,141]
143. ! [X0] : ((subquery_from_6(X0) | ~select_7(X0)) & (select_7(X0) | ~subquery_from_6(X0))) [nnf transformation 48]
144. ! [X0] : ((from_5(X0) | ~subquery_from_6(X0)) & (subquery_from_6(X0) | ~from_5(X0))) [nnf transformation 49]
145. ! [X0] : ((select_4(X0) | ~from_5(X0)) & (from_5(X0) | ~select_4(X0))) [nnf transformation 51]
146. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 52]
149. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 130]
150. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
151. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 149,150]
152. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 131]
153. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 131]
155. ~eq(X0,X0) | where_3(X0) [cnf transformation 132]
156. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 136]
158. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 136]
159. ~main_query_1(X0) | select_1(X0) [cnf transformation 137]
160. ~select_1(X0) | main_query_1(X0) [cnf transformation 137]
161. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 138]
162. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 138]
163. from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [cnf transformation 142]
164. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 142]
165. ~subquery_from_6(X0) | select_7(X0) [cnf transformation 143]
166. ~select_7(X0) | subquery_from_6(X0) [cnf transformation 143]
167. ~from_5(X0) | subquery_from_6(X0) [cnf transformation 144]
168. ~subquery_from_6(X0) | from_5(X0) [cnf transformation 144]
169. ~select_4(X0) | from_5(X0) [cnf transformation 145]
170. ~from_5(X0) | select_4(X0) [cnf transformation 145]
171. ~main_query_2(X0) | select_4(X0) [cnf transformation 146]
172. ~select_4(X0) | main_query_2(X0) [cnf transformation 146]
173. eq(X0,X0) [cnf transformation 53]
202. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 151]
203. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 151]
205. 0 <=> main_query_1(sK8) [avatar definition]
206. main_query_1(sK8) <- {0} [avatar component clause 205]
208. 1 <=> ~main_query_1(sK8) [avatar definition]
209. ~main_query_1(sK8) <- {1} [avatar component clause 208]
211. 2 <=> main_query_2(sK8) [avatar definition]
212. main_query_2(sK8) <- {2} [avatar component clause 211]
214. 3 <=> ~main_query_2(sK8) [avatar definition]
215. ~main_query_2(sK8) <- {3} [avatar component clause 214]
216. ~1 | ~3 [avatar split clause 203,214,208]
217. 0 | 2 [avatar split clause 202,211,205]
218. select_4(sK8) <- {2} [resolution 171,212]
219. from_5(sK8) <- {2} [resolution 218,169]
222. subquery_from_6(sK8) <- {2} [resolution 219,167]
224. select_7(sK8) <- {2} [resolution 222,165]
226. where_3(X1) [resolution 155,173]
261. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 156,152]
263. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [resolution 163,161]
266. from_8(X2,sK0(X2),sK1(X2),sK2(X2),sK3(X2)) | ~select_1(X2) [resolution 261,162]
272. from_2(X5,sK4(X5),sK5(X5),sK6(X5),sK7(X5)) | ~select_7(X5) [resolution 263,153]
276. ~select_1(X1) | select_7(X1) [resolution 266,164]
282. ~select_7(X1) | ~where_3(X1) | select_1(X1) [resolution 272,158]
283. ~select_7(X1) | select_1(X1) [subsumption resolution 282,226]
286. select_1(sK8) <- {2} [resolution 283,224]
290. main_query_1(sK8) <- {2} [resolution 286,160]
292. $false <- {1, 2} [subsumption resolution 290,209]
293. 1 | ~2 [avatar contradiction clause 292]
294. select_1(sK8) <- {0} [resolution 206,159]
297. select_7(sK8) <- {0} [resolution 294,276]
303. subquery_from_6(sK8) <- {0} [resolution 297,166]
306. from_5(sK8) <- {0} [resolution 303,168]
308. select_4(sK8) <- {0} [resolution 306,170]
310. main_query_2(sK8) <- {0} [resolution 308,172]
312. $false <- {0, 3} [subsumption resolution 310,215]
313. ~0 | 3 [avatar contradiction clause 312]
314. $false [avatar sat refutation 216,217,293,313]
% SZS output end Proof for query05_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.023 s
% ------------------------------
% ------------------------------
