% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_01_query_08_tptp4X
% SZS output start Proof for query_01_query_08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
4. ! [X7,X8,X9,X10,X11] : (from_7(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
5. ! [X7,X8] : (select_6(X7,X8) <=> ? [X9,X10,X11] : ($true & from_7(X7,X8,X9,X10,X11))) [input]
6. ! [X5,X6] : (subquery_from_5(X5,X6) <=> select_6(X5,X6)) [input]
7. ! [X12,X13] : (from_4(X12,X13) <=> subquery_from_5(X12,X13)) [input]
8. ! [X12,X13] : (select_3(X12,X13) <=> ($true & from_4(X12,X13))) [input]
9. ! [X5,X6] : (main_query_2(X5,X6) <=> select_3(X5,X6)) [input]
54. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
55. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 54]
56. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 3]
58. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
59. ! [X0,X1] : (select_6(X0,X1) <=> ? [X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [rectify 5]
60. ! [X0,X1] : (select_6(X0,X1) <=> ? [X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 59]
61. ! [X0,X1] : (subquery_from_5(X0,X1) <=> select_6(X0,X1)) [rectify 6]
62. ! [X0,X1] : (from_4(X0,X1) <=> subquery_from_5(X0,X1)) [rectify 7]
63. ! [X0,X1] : (select_3(X0,X1) <=> ($true & from_4(X0,X1))) [rectify 8]
64. ! [X0,X1] : (select_3(X0,X1) <=> from_4(X0,X1)) [true and false elimination 63]
65. ! [X0,X1] : (main_query_2(X0,X1) <=> select_3(X0,X1)) [rectify 9]
110. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 55]
184. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 110]
185. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
186. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0,X1))) [nnf transformation 56]
187. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) | ~select_1(X0,X1))) [rectify 186]
188. ! [X1,X0] : (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) => from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) [choice axiom]
189. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1))) [skolemisation 187,188]
190. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
191. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 58]
192. ! [X0,X1] : ((select_6(X0,X1) | ! [X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0,X1))) [nnf transformation 60]
193. ! [X0,X1] : ((select_6(X0,X1) | ! [X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_7(X0,X1,X5,X6,X7) | ~select_6(X0,X1))) [rectify 192]
194. ! [X1,X0] : (? [X5,X6,X7] : from_7(X0,X1,X5,X6,X7) => from_7(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) [choice axiom]
195. ! [X0,X1] : ((select_6(X0,X1) | ! [X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_6(X0,X1))) [skolemisation 193,194]
196. ! [X0,X1] : ((subquery_from_5(X0,X1) | ~select_6(X0,X1)) & (select_6(X0,X1) | ~subquery_from_5(X0,X1))) [nnf transformation 61]
197. ! [X0,X1] : ((from_4(X0,X1) | ~subquery_from_5(X0,X1)) & (subquery_from_5(X0,X1) | ~from_4(X0,X1))) [nnf transformation 62]
198. ! [X0,X1] : ((select_3(X0,X1) | ~from_4(X0,X1)) & (from_4(X0,X1) | ~select_3(X0,X1))) [nnf transformation 64]
199. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 65]
202. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 184]
203. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
204. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 202,203]
205. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 185]
206. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 185]
207. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 189]
208. ~from_2(X0,X1,X2,X3,X4) | select_1(X0,X1) [cnf transformation 189]
209. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 190]
210. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 190]
211. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 191]
212. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 191]
213. from_7(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_6(X0,X1) [cnf transformation 195]
214. ~from_7(X0,X1,X2,X3,X4) | select_6(X0,X1) [cnf transformation 195]
215. ~subquery_from_5(X0,X1) | select_6(X0,X1) [cnf transformation 196]
216. ~select_6(X0,X1) | subquery_from_5(X0,X1) [cnf transformation 196]
217. ~from_4(X0,X1) | subquery_from_5(X0,X1) [cnf transformation 197]
218. ~subquery_from_5(X0,X1) | from_4(X0,X1) [cnf transformation 197]
219. ~select_3(X0,X1) | from_4(X0,X1) [cnf transformation 198]
220. ~from_4(X0,X1) | select_3(X0,X1) [cnf transformation 198]
221. ~main_query_2(X0,X1) | select_3(X0,X1) [cnf transformation 199]
222. ~select_3(X0,X1) | main_query_2(X0,X1) [cnf transformation 199]
265. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 204]
266. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 204]
268. 0 <=> main_query_1(sK6,sK7) [avatar definition]
269. main_query_1(sK6,sK7) <- {0} [avatar component clause 268]
271. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
272. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 271]
274. 2 <=> main_query_2(sK6,sK7) [avatar definition]
275. main_query_2(sK6,sK7) <- {2} [avatar component clause 274]
277. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
278. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 277]
279. ~1 | ~3 [avatar split clause 266,277,271]
280. 0 | 2 [avatar split clause 265,274,268]
281. select_3(sK6,sK7) <- {2} [resolution 221,275]
283. from_4(sK6,sK7) <- {2} [resolution 281,219]
286. subquery_from_5(sK6,sK7) <- {2} [resolution 283,217]
289. select_6(sK6,sK7) <- {2} [resolution 286,215]
309. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 207,205]
311. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_6(X0,X1) [resolution 213,211]
324. from_7(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 309,212]
330. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_6(X0,X1) [resolution 311,206]
338. ~select_1(X2,X3) | select_6(X2,X3) [resolution 324,214]
359. ~select_6(X2,X3) | select_1(X2,X3) [resolution 330,208]
372. select_1(sK6,sK7) <- {2} [resolution 359,289]
381. main_query_1(sK6,sK7) <- {2} [resolution 372,210]
382. $false <- {1, 2} [subsumption resolution 381,272]
383. 1 | ~2 [avatar contradiction clause 382]
384. select_1(sK6,sK7) <- {0} [resolution 269,209]
385. select_6(sK6,sK7) <- {0} [resolution 384,338]
388. subquery_from_5(sK6,sK7) <- {0} [resolution 385,216]
396. from_4(sK6,sK7) <- {0} [resolution 388,218]
405. select_3(sK6,sK7) <- {0} [resolution 396,220]
414. main_query_2(sK6,sK7) <- {0} [resolution 405,222]
416. $false <- {0, 3} [subsumption resolution 414,278]
417. ~0 | 3 [avatar contradiction clause 416]
418. $false [avatar sat refutation 279,280,383,417]
% SZS output end Proof for query_01_query_08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.010 s
% ------------------------------
% ------------------------------
