% Refutation found. Thanks to Tanya!
% SZS status Theorem for query20_query21_tptp4X
% SZS output start Proof for query20_query21_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
3. ! [X0,X1] : (where_10(X0,X1) <=> eq(X0,X1)) [input]
4. ! [X0,X1] : (select_8(X0,X1) <=> ? [X2,X3,X4] : (where_10(X0,X1) & from_9(X0,X1,X2,X3,X4))) [input]
5. ! [X5,X6] : (subquery_from_7(X5,X6) <=> select_8(X5,X6)) [input]
6. ! [X7,X8] : (from_6(X7,X8) <=> subquery_from_7(X7,X8)) [input]
7. ! [X7] : (select_5(X7) <=> ? [X8] : ($true & from_6(X7,X8))) [input]
8. ! [X5] : (where_3_subquery_4(X5) <=> select_5(X5)) [input]
9. ! [X0] : (where_3(X0) <=> where_3_subquery_4(X0)) [input]
10. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
11. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
12. ! [X0,X1,X2,X3,X4] : (from_15(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
13. ! [X0,X1] : (where_16(X0,X1) <=> eq(X0,X1)) [input]
14. ! [X0,X1] : (select_14(X0,X1) <=> ? [X2,X3,X4] : (where_16(X0,X1) & from_15(X0,X1,X2,X3,X4))) [input]
15. ! [X5,X6] : (subquery_from_13(X5,X6) <=> select_14(X5,X6)) [input]
16. ! [X7,X8] : (from_12(X7,X8) <=> subquery_from_13(X7,X8)) [input]
17. ! [X7] : (select_11(X7) <=> ? [X8] : ($true & from_12(X7,X8))) [input]
18. ! [X5] : (main_query_2(X5) <=> select_11(X5)) [input]
19. ! [X9] : eq(X9,X9) [input]
20. ! [X9,X10] : (eq(X9,X10) => eq(X10,X9)) [input]
38. ! [X9,X7,X8] : ((from_6(X7,X8) & eq(X8,X9)) => from_6(X7,X9)) [input]
50. ! [X9,X7,X8] : ((from_12(X7,X8) & eq(X8,X9)) => from_12(X7,X9)) [input]
71. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
72. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 71]
73. ! [X0,X1] : (subquery_from_7(X0,X1) <=> select_8(X0,X1)) [rectify 5]
74. ! [X0,X1] : (from_6(X0,X1) <=> subquery_from_7(X0,X1)) [rectify 6]
75. ! [X0] : (select_5(X0) <=> ? [X1] : ($true & from_6(X0,X1))) [rectify 7]
76. ! [X0] : (select_5(X0) <=> ? [X1] : from_6(X0,X1)) [true and false elimination 75]
77. ! [X0] : (where_3_subquery_4(X0) <=> select_5(X0)) [rectify 8]
78. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 11]
79. ! [X0,X1] : (subquery_from_13(X0,X1) <=> select_14(X0,X1)) [rectify 15]
80. ! [X0,X1] : (from_12(X0,X1) <=> subquery_from_13(X0,X1)) [rectify 16]
81. ! [X0] : (select_11(X0) <=> ? [X1] : ($true & from_12(X0,X1))) [rectify 17]
82. ! [X0] : (select_11(X0) <=> ? [X1] : from_12(X0,X1)) [true and false elimination 81]
83. ! [X0] : (main_query_2(X0) <=> select_11(X0)) [rectify 18]
84. ! [X0] : eq(X0,X0) [rectify 19]
85. ! [X0,X1] : (eq(X0,X1) => eq(X1,X0)) [rectify 20]
103. ! [X0,X1,X2] : ((from_6(X1,X2) & eq(X2,X0)) => from_6(X1,X0)) [rectify 38]
115. ! [X0,X1,X2] : ((from_12(X1,X2) & eq(X2,X0)) => from_12(X1,X0)) [rectify 50]
136. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 72]
137. ! [X0,X1] : (eq(X1,X0) | ~eq(X0,X1)) [ennf transformation 85]
158. ! [X0,X1,X2] : (from_6(X1,X0) | (~from_6(X1,X2) | ~eq(X2,X0))) [ennf transformation 103]
159. ! [X0,X1,X2] : (from_6(X1,X0) | ~from_6(X1,X2) | ~eq(X2,X0)) [flattening 158]
182. ! [X0,X1,X2] : (from_12(X1,X0) | (~from_12(X1,X2) | ~eq(X2,X0))) [ennf transformation 115]
183. ! [X0,X1,X2] : (from_12(X1,X0) | ~from_12(X1,X2) | ~eq(X2,X0)) [flattening 182]
224. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 136]
225. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
226. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 2]
227. ! [X0,X1] : ((where_10(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_10(X0,X1))) [nnf transformation 3]
228. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4] : (~where_10(X0,X1) | ~from_9(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_10(X0,X1) & from_9(X0,X1,X2,X3,X4)) | ~select_8(X0,X1))) [nnf transformation 4]
229. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4] : (~where_10(X0,X1) | ~from_9(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_10(X0,X1) & from_9(X0,X1,X5,X6,X7)) | ~select_8(X0,X1))) [rectify 228]
230. ! [X1,X0] : (? [X5,X6,X7] : (where_10(X0,X1) & from_9(X0,X1,X5,X6,X7)) => (where_10(X0,X1) & from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
231. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4] : (~where_10(X0,X1) | ~from_9(X0,X1,X2,X3,X4))) & ((where_10(X0,X1) & from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_8(X0,X1))) [skolemisation 229,230]
232. ! [X0,X1] : ((subquery_from_7(X0,X1) | ~select_8(X0,X1)) & (select_8(X0,X1) | ~subquery_from_7(X0,X1))) [nnf transformation 73]
233. ! [X0,X1] : ((from_6(X0,X1) | ~subquery_from_7(X0,X1)) & (subquery_from_7(X0,X1) | ~from_6(X0,X1))) [nnf transformation 74]
234. ! [X0] : ((select_5(X0) | ! [X1] : ~from_6(X0,X1)) & (? [X1] : from_6(X0,X1) | ~select_5(X0))) [nnf transformation 76]
235. ! [X0] : ((select_5(X0) | ! [X1] : ~from_6(X0,X1)) & (? [X2] : from_6(X0,X2) | ~select_5(X0))) [rectify 234]
236. ! [X0] : (? [X2] : from_6(X0,X2) => from_6(X0,sK3(X0))) [choice axiom]
237. ! [X0] : ((select_5(X0) | ! [X1] : ~from_6(X0,X1)) & (from_6(X0,sK3(X0)) | ~select_5(X0))) [skolemisation 235,236]
238. ! [X0] : ((where_3_subquery_4(X0) | ~select_5(X0)) & (select_5(X0) | ~where_3_subquery_4(X0))) [nnf transformation 77]
239. ! [X0] : ((where_3(X0) | ~where_3_subquery_4(X0)) & (where_3_subquery_4(X0) | ~where_3(X0))) [nnf transformation 9]
240. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 10]
241. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 240]
242. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
243. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_1(X0))) [skolemisation 241,242]
244. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 78]
245. ! [X0,X1,X2,X3,X4] : ((from_15(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_15(X0,X1,X2,X3,X4))) [nnf transformation 12]
246. ! [X0,X1] : ((where_16(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_16(X0,X1))) [nnf transformation 13]
247. ! [X0,X1] : ((select_14(X0,X1) | ! [X2,X3,X4] : (~where_16(X0,X1) | ~from_15(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_16(X0,X1) & from_15(X0,X1,X2,X3,X4)) | ~select_14(X0,X1))) [nnf transformation 14]
248. ! [X0,X1] : ((select_14(X0,X1) | ! [X2,X3,X4] : (~where_16(X0,X1) | ~from_15(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_16(X0,X1) & from_15(X0,X1,X5,X6,X7)) | ~select_14(X0,X1))) [rectify 247]
249. ! [X1,X0] : (? [X5,X6,X7] : (where_16(X0,X1) & from_15(X0,X1,X5,X6,X7)) => (where_16(X0,X1) & from_15(X0,X1,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)))) [choice axiom]
250. ! [X0,X1] : ((select_14(X0,X1) | ! [X2,X3,X4] : (~where_16(X0,X1) | ~from_15(X0,X1,X2,X3,X4))) & ((where_16(X0,X1) & from_15(X0,X1,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1))) | ~select_14(X0,X1))) [skolemisation 248,249]
251. ! [X0,X1] : ((subquery_from_13(X0,X1) | ~select_14(X0,X1)) & (select_14(X0,X1) | ~subquery_from_13(X0,X1))) [nnf transformation 79]
252. ! [X0,X1] : ((from_12(X0,X1) | ~subquery_from_13(X0,X1)) & (subquery_from_13(X0,X1) | ~from_12(X0,X1))) [nnf transformation 80]
253. ! [X0] : ((select_11(X0) | ! [X1] : ~from_12(X0,X1)) & (? [X1] : from_12(X0,X1) | ~select_11(X0))) [nnf transformation 82]
254. ! [X0] : ((select_11(X0) | ! [X1] : ~from_12(X0,X1)) & (? [X2] : from_12(X0,X2) | ~select_11(X0))) [rectify 253]
255. ! [X0] : (? [X2] : from_12(X0,X2) => from_12(X0,sK11(X0))) [choice axiom]
256. ! [X0] : ((select_11(X0) | ! [X1] : ~from_12(X0,X1)) & (from_12(X0,sK11(X0)) | ~select_11(X0))) [skolemisation 254,255]
257. ! [X0] : ((main_query_2(X0) | ~select_11(X0)) & (select_11(X0) | ~main_query_2(X0))) [nnf transformation 83]
260. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 224]
261. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK12) | ~main_query_1(sK12)) & (main_query_2(sK12) | main_query_1(sK12))) [choice axiom]
262. (~main_query_2(sK12) | ~main_query_1(sK12)) & (main_query_2(sK12) | main_query_1(sK12)) [skolemisation 260,261]
264. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 225]
265. ~from_9(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 226]
266. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 226]
267. ~where_10(X0,X1) | eq(X0,X1) [cnf transformation 227]
268. ~eq(X0,X1) | where_10(X0,X1) [cnf transformation 227]
269. from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_8(X0,X1) [cnf transformation 231]
270. ~select_8(X0,X1) | where_10(X0,X1) [cnf transformation 231]
271. ~from_9(X0,X1,X2,X3,X4) | ~where_10(X0,X1) | select_8(X0,X1) [cnf transformation 231]
272. ~subquery_from_7(X0,X1) | select_8(X0,X1) [cnf transformation 232]
273. ~select_8(X0,X1) | subquery_from_7(X0,X1) [cnf transformation 232]
274. ~from_6(X0,X1) | subquery_from_7(X0,X1) [cnf transformation 233]
275. ~subquery_from_7(X0,X1) | from_6(X0,X1) [cnf transformation 233]
276. from_6(X0,sK3(X0)) | ~select_5(X0) [cnf transformation 237]
277. ~from_6(X0,X1) | select_5(X0) [cnf transformation 237]
278. ~where_3_subquery_4(X0) | select_5(X0) [cnf transformation 238]
279. ~select_5(X0) | where_3_subquery_4(X0) [cnf transformation 238]
280. ~where_3(X0) | where_3_subquery_4(X0) [cnf transformation 239]
281. ~where_3_subquery_4(X0) | where_3(X0) [cnf transformation 239]
283. ~select_1(X0) | where_3(X0) [cnf transformation 243]
284. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 243]
285. ~main_query_1(X0) | select_1(X0) [cnf transformation 244]
286. ~select_1(X0) | main_query_1(X0) [cnf transformation 244]
287. ~from_15(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 245]
288. ~tab1(X0,X1,X2,X3,X4) | from_15(X0,X1,X2,X3,X4) [cnf transformation 245]
289. ~where_16(X0,X1) | eq(X0,X1) [cnf transformation 246]
290. ~eq(X0,X1) | where_16(X0,X1) [cnf transformation 246]
291. from_15(X0,X1,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) | ~select_14(X0,X1) [cnf transformation 250]
292. ~select_14(X0,X1) | where_16(X0,X1) [cnf transformation 250]
293. ~from_15(X0,X1,X2,X3,X4) | ~where_16(X0,X1) | select_14(X0,X1) [cnf transformation 250]
294. ~subquery_from_13(X0,X1) | select_14(X0,X1) [cnf transformation 251]
295. ~select_14(X0,X1) | subquery_from_13(X0,X1) [cnf transformation 251]
296. ~from_12(X0,X1) | subquery_from_13(X0,X1) [cnf transformation 252]
297. ~subquery_from_13(X0,X1) | from_12(X0,X1) [cnf transformation 252]
298. from_12(X0,sK11(X0)) | ~select_11(X0) [cnf transformation 256]
299. ~from_12(X0,X1) | select_11(X0) [cnf transformation 256]
300. ~main_query_2(X0) | select_11(X0) [cnf transformation 257]
301. ~select_11(X0) | main_query_2(X0) [cnf transformation 257]
302. eq(X0,X0) [cnf transformation 84]
303. ~eq(X0,X1) | eq(X1,X0) [cnf transformation 137]
319. ~from_6(X1,X2) | from_6(X1,X0) | ~eq(X2,X0) [cnf transformation 159]
331. ~from_12(X1,X2) | from_12(X1,X0) | ~eq(X2,X0) [cnf transformation 183]
352. main_query_2(sK12) | main_query_1(sK12) [cnf transformation 262]
353. ~main_query_2(sK12) | ~main_query_1(sK12) [cnf transformation 262]
355. 0 <=> main_query_1(sK12) [avatar definition]
356. main_query_1(sK12) <- {0} [avatar component clause 355]
358. 1 <=> ~main_query_1(sK12) [avatar definition]
361. 2 <=> main_query_2(sK12) [avatar definition]
362. main_query_2(sK12) <- {2} [avatar component clause 361]
364. 3 <=> ~main_query_2(sK12) [avatar definition]
365. ~main_query_2(sK12) <- {3} [avatar component clause 364]
366. ~1 | ~3 [avatar split clause 353,364,358]
367. 0 | 2 [avatar split clause 352,361,355]
368. select_11(sK12) <- {2} [resolution 300,362]
370. where_10(X0,X0) [resolution 268,302]
372. subquery_from_7(X0,sK3(X0)) | ~select_5(X0) [resolution 276,274]
374. where_16(X0,X0) [resolution 290,302]
376. subquery_from_13(X0,sK11(X0)) | ~select_11(X0) [resolution 298,296]
383. select_8(X1,sK3(X1)) | ~select_5(X1) [resolution 372,272]
386. select_14(X1,sK11(X1)) | ~select_11(X1) [resolution 376,294]
391. where_10(X1,sK3(X1)) | ~select_5(X1) [resolution 383,270]
396. where_16(X1,sK11(X1)) | ~select_11(X1) [resolution 386,292]
398. eq(X0,sK3(X0)) | ~select_5(X0) [resolution 391,267]
400. eq(X0,sK11(X0)) | ~select_11(X0) [resolution 396,289]
405. eq(sK3(X4),X4) | ~select_5(X4) [resolution 398,303]
408. ~eq(sK3(X0),X1) | from_6(X0,X1) | ~select_5(X0) [resolution 319,276]
412. eq(sK11(X4),X4) | ~select_11(X4) [resolution 400,303]
418. ~eq(sK11(X0),X1) | from_12(X0,X1) | ~select_11(X0) [resolution 331,298]
503. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_8(X2,X3) [resolution 269,265]
506. tab1(X2,X3,sK8(X2,X3),sK9(X2,X3),sK10(X2,X3)) | ~select_14(X2,X3) [resolution 291,287]
507. from_6(X0,X0) | ~select_5(X0) | ~select_5(X0) [resolution 408,405]
511. from_6(X0,X0) | ~select_5(X0) [duplicate literal removal 507]
516. subquery_from_7(X4,X4) | ~select_5(X4) [resolution 511,274]
522. select_8(X5,X5) | ~select_5(X5) [resolution 516,272]
650. from_12(X0,X0) | ~select_11(X0) | ~select_11(X0) [resolution 418,412]
655. from_12(X0,X0) | ~select_11(X0) [duplicate literal removal 650]
668. subquery_from_13(X4,X4) | ~select_11(X4) [resolution 655,296]
672. from_15(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_8(X4,X5) [resolution 503,288]
679. from_2(X0,X1,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) | ~select_14(X0,X1) [resolution 506,264]
680. from_9(X2,X3,sK8(X2,X3),sK9(X2,X3),sK10(X2,X3)) | ~select_14(X2,X3) [resolution 506,266]
690. select_14(X5,X5) | ~select_11(X5) [resolution 668,294]
1114. ~where_16(X2,X3) | ~select_8(X2,X3) | select_14(X2,X3) [resolution 672,293]
1123. ~select_14(X2,X3) | ~where_3(X2) | select_1(X2) [resolution 679,284]
1131. ~select_14(X2,X3) | ~where_10(X2,X3) | select_8(X2,X3) [resolution 680,271]
1168. ~select_11(X0) | select_1(X0) | ~where_3(X0) [resolution 1123,386]
1181. select_1(sK12) | ~where_3(sK12) <- {2} [resolution 1168,368]
1186. 4 <=> where_3(sK12) [avatar definition]
1187. where_3(sK12) <- {4} [avatar component clause 1186]
1189. 5 <=> ~where_3(sK12) [avatar definition]
1190. ~where_3(sK12) <- {5} [avatar component clause 1189]
1195. 6 <=> select_1(sK12) [avatar definition]
1196. select_1(sK12) <- {6} [avatar component clause 1195]
1197. ~5 | 6 | ~2 [avatar split clause 1181,361,1195,1189]
1921. ~select_8(X0,X0) | select_14(X0,X0) [resolution 1114,374]
1940. select_14(X1,X1) | ~select_5(X1) [resolution 1921,522]
1946. ~where_10(X5,X5) | select_8(X5,X5) | ~select_11(X5) [resolution 1131,690]
1954. select_8(X5,X5) | ~select_11(X5) [subsumption resolution 1946,370]
1958. subquery_from_13(X1,X1) | ~select_5(X1) [resolution 1940,295]
1966. subquery_from_7(X1,X1) | ~select_11(X1) [resolution 1954,273]
1975. from_12(X4,X4) | ~select_5(X4) [resolution 1958,297]
1991. from_6(X4,X4) | ~select_11(X4) [resolution 1966,275]
2011. ~select_5(X5) | select_11(X5) [resolution 1975,299]
2043. ~select_11(X5) | select_5(X5) [resolution 1991,277]
2064. select_5(sK12) <- {2} [resolution 2043,368]
2091. where_3_subquery_4(sK12) <- {2} [resolution 2064,279]
2099. where_3(sK12) <- {2} [resolution 2091,281]
2101. $false <- {2, 5} [subsumption resolution 2099,1190]
2102. ~2 | 5 [avatar contradiction clause 2101]
2103. select_1(sK12) <- {0} [resolution 356,285]
2104. 6 | ~0 [avatar split clause 2103,355,1195]
2117. where_3_subquery_4(sK12) <- {4} [resolution 1187,280]
2125. main_query_1(sK12) <- {6} [resolution 1196,286]
2126. where_3(sK12) <- {6} [resolution 1196,283]
2129. 0 | ~6 [avatar split clause 2125,1195,355]
2139. select_5(sK12) <- {4} [resolution 2117,278]
2147. select_11(sK12) <- {4} [resolution 2139,2011]
2158. main_query_2(sK12) <- {4} [resolution 2147,301]
2159. $false <- {3, 4} [subsumption resolution 2158,365]
2160. 3 | ~4 [avatar contradiction clause 2159]
2161. 4 | ~6 [avatar split clause 2126,1195,1186]
2162. $false [avatar sat refutation 366,367,1197,2102,2104,2129,2160,2161]
% SZS output end Proof for query20_query21_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5756
% Time elapsed: 0.035 s
% ------------------------------
% ------------------------------
