% Refutation found. Thanks to Tanya!
% SZS status Theorem for query12_query14_tptp4X
% SZS output start Proof for query12_query14_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (query_7(X5) <=> select_8(X5)) [input]
9. ! [X0,X1] : (where_6(X0,X1) <=> (query_7(X0) & eq(X0,X1))) [input]
10. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
16. ! [X7,X8] : ((lte(X8,X7) & lte(X7,X8)) <=> eq(X7,X8)) [input]
62. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
63. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 62]
64. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
65. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 7]
66. ! [X0] : (query_7(X0) <=> select_8(X0)) [rectify 8]
67. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
72. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 16]
118. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 63]
204. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 118]
205. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
206. ! [X0,X1] : ((where_3(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [nnf transformation 2]
207. ! [X0,X1] : ((where_3(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_3(X0,X1))) [flattening 206]
208. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
209. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 208]
210. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
211. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 209,210]
212. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 64]
213. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
214. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
215. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 65]
216. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 215]
217. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
218. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 216,217]
219. ! [X0] : ((query_7(X0) | ~select_8(X0)) & (select_8(X0) | ~query_7(X0))) [nnf transformation 66]
220. ! [X0,X1] : ((where_6(X0,X1) | (~query_7(X0) | ~eq(X0,X1))) & ((query_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 9]
221. ! [X0,X1] : ((where_6(X0,X1) | ~query_7(X0) | ~eq(X0,X1)) & ((query_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 220]
222. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 10]
223. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 222]
224. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
225. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 223,224]
226. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 67]
227. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 72]
228. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 227]
229. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 204]
230. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
231. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 229,230]
232. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 205]
233. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 205]
234. ~where_3(X0,X1) | lte(X0,X1) [cnf transformation 207]
235. ~where_3(X0,X1) | lte(X1,X0) [cnf transformation 207]
236. ~lte(X1,X0) | where_3(X0,X1) | ~lte(X0,X1) [cnf transformation 207]
237. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 211]
238. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 211]
239. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 211]
240. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 212]
241. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 212]
242. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 213]
243. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 213]
245. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 214]
247. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 218]
249. ~select_8(X0) | query_7(X0) [cnf transformation 219]
250. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 221]
252. ~eq(X0,X1) | ~query_7(X0) | where_6(X0,X1) [cnf transformation 221]
253. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 225]
254. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 225]
255. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 225]
256. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 226]
257. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 226]
262. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 228]
263. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 228]
264. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 228]
306. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 231]
307. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 231]
349. 0 <=> main_query_1(sK10,sK11) [avatar definition]
350. main_query_1(sK10,sK11) <- {0} [avatar component clause 349]
352. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
353. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 352]
355. 2 <=> main_query_2(sK10,sK11) [avatar definition]
356. main_query_2(sK10,sK11) <- {2} [avatar component clause 355]
358. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
359. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 358]
360. ~1 | ~3 [avatar split clause 307,358,352]
361. 0 | 2 [avatar split clause 306,355,349]
362. select_4(sK10,sK11) <- {2} [resolution 256,356]
364. where_6(sK10,sK11) <- {2} [resolution 362,254]
366. eq(sK10,sK11) <- {2} [resolution 364,250]
372. lte(sK11,sK10) <- {2} [resolution 366,264]
373. lte(sK10,sK11) <- {2} [resolution 366,263]
378. where_3(sK10,sK11) | ~lte(sK10,sK11) <- {2} [resolution 236,372]
380. where_3(sK10,sK11) <- {2} [subsumption resolution 378,373]
428. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 237,232]
433. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 253,242]
436. from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 428,245]
437. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 428,243]
441. from_2(X4,X5,sK7(X4,X5),sK8(X4,X5),sK9(X4,X5)) | ~select_4(X4,X5) [resolution 433,233]
445. ~select_1(X2,X3) | select_8(X2) [resolution 436,247]
446. ~where_6(X0,X1) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 437,255]
453. ~select_4(X0,X1) | ~where_3(X0,X1) | select_1(X0,X1) [resolution 441,239]
455. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 453,362]
456. select_1(sK10,sK11) <- {2} [subsumption resolution 455,380]
458. main_query_1(sK10,sK11) <- {2} [resolution 456,241]
460. $false <- {1, 2} [subsumption resolution 458,353]
461. 1 | ~2 [avatar contradiction clause 460]
462. select_1(sK10,sK11) <- {0} [resolution 350,240]
463. select_8(sK10) <- {0} [resolution 462,445]
465. where_3(sK10,sK11) <- {0} [resolution 462,238]
466. query_7(sK10) <- {0} [resolution 463,249]
467. lte(sK11,sK10) <- {0} [resolution 465,235]
468. lte(sK10,sK11) <- {0} [resolution 465,234]
471. eq(sK10,sK11) | ~lte(sK10,sK11) <- {0} [resolution 467,262]
473. eq(sK10,sK11) <- {0} [subsumption resolution 471,468]
480. ~query_7(sK10) | where_6(sK10,sK11) <- {0} [resolution 473,252]
484. where_6(sK10,sK11) <- {0} [subsumption resolution 480,466]
494. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {0} [resolution 484,446]
497. select_4(sK10,sK11) <- {0} [subsumption resolution 494,462]
502. main_query_2(sK10,sK11) <- {0} [resolution 497,257]
504. $false <- {0, 3} [subsumption resolution 502,359]
505. ~0 | 3 [avatar contradiction clause 504]
506. $false [avatar sat refutation 360,361,461,505]
% SZS output end Proof for query12_query14_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.016 s
% ------------------------------
% ------------------------------
