% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_11_query_12_tptp4X
% SZS output start Proof for query_11_query_12_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8] : (where_6(X7,X8) <=> eq(X7,X8)) [input]
7. ! [X7,X8] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X7,X8) & from_5(X7,X8,X9,X10,X11))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
52. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
53. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 52]
54. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
55. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
56. ! [X0,X1] : (where_6(X0,X1) <=> eq(X0,X1)) [rectify 6]
57. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
58. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
102. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 53]
172. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 102]
173. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
174. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
175. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
176. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 175]
177. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
178. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 176,177]
179. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 54]
180. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 55]
181. ! [X0,X1] : ((where_6(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_6(X0,X1))) [nnf transformation 56]
182. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 57]
183. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 182]
184. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
185. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 183,184]
186. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 58]
189. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 172]
190. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
191. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 189,190]
192. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 173]
193. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 173]
194. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 174]
195. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 174]
196. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 178]
197. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 178]
198. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 178]
199. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 179]
200. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 179]
201. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 180]
202. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 180]
203. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 181]
204. ~eq(X0,X1) | where_6(X0,X1) [cnf transformation 181]
205. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 185]
206. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 185]
207. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 185]
208. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 186]
209. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 186]
251. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 191]
252. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 191]
254. 0 <=> main_query_1(sK6,sK7) [avatar definition]
255. main_query_1(sK6,sK7) <- {0} [avatar component clause 254]
257. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
258. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 257]
260. 2 <=> main_query_2(sK6,sK7) [avatar definition]
261. main_query_2(sK6,sK7) <- {2} [avatar component clause 260]
263. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
264. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 263]
265. ~1 | ~3 [avatar split clause 252,263,257]
266. 0 | 2 [avatar split clause 251,260,254]
271. select_4(sK6,sK7) <- {2} [resolution 208,261]
273. where_6(sK6,sK7) <- {2} [resolution 271,206]
275. eq(sK6,sK7) <- {2} [resolution 273,203]
280. where_3(sK6,sK7) <- {2} [resolution 275,195]
367. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 196,192]
371. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 205,201]
466. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 367,202]
485. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 371,193]
860. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 466,207]
870. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 485,198]
883. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 870,271]
887. select_1(sK6,sK7) <- {2} [subsumption resolution 883,280]
900. main_query_1(sK6,sK7) <- {2} [resolution 887,200]
902. $false <- {1, 2} [subsumption resolution 900,258]
903. 1 | ~2 [avatar contradiction clause 902]
906. select_1(sK6,sK7) <- {0} [resolution 255,199]
910. where_3(sK6,sK7) <- {0} [resolution 906,197]
911. eq(sK6,sK7) <- {0} [resolution 910,194]
914. where_6(sK6,sK7) <- {0} [resolution 911,204]
924. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 914,860]
928. select_4(sK6,sK7) <- {0} [subsumption resolution 924,906]
1014. main_query_2(sK6,sK7) <- {0} [resolution 928,209]
1016. $false <- {0, 3} [subsumption resolution 1014,264]
1017. ~0 | 3 [avatar contradiction clause 1016]
1018. $false [avatar sat refutation 265,266,903,1017]
% SZS output end Proof for query_11_query_12_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.019 s
% ------------------------------
% ------------------------------
