% Refutation found. Thanks to Tanya!
% SZS status Theorem for query01_query08_tptp4X
% SZS output start Proof for query01_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
6. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X6] : (main_query_2(X6) <=> select_3(X6)) [input]
11. ! [X7] : lte(X7,X7) [input]
17. ! [X7,X8] : (gte(X7,X8) <=> lte(X8,X7)) [input]
41. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
42. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
45. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
49. ! [X0] : lte(X0,X0) [rectify 11]
55. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 17]
79. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 42]
131. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 79]
132. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
133. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 133]
135. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 134,135]
137. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 44]
138. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
139. ! [X0] : ((where_5(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [nnf transformation 5]
140. ! [X0] : ((where_5(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [flattening 139]
141. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 6]
142. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 141]
143. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
144. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 142,143]
145. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 45]
148. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 55]
149. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 131]
150. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
151. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 149,150]
152. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 132]
153. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 132]
154. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 136]
155. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 136]
156. ~main_query_1(X0) | select_1(X0) [cnf transformation 137]
157. ~select_1(X0) | main_query_1(X0) [cnf transformation 137]
158. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 138]
159. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 138]
162. ~gte(X0,X0) | where_5(X0) | ~lte(X0,X0) [cnf transformation 140]
163. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 144]
165. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 144]
166. ~main_query_2(X0) | select_3(X0) [cnf transformation 145]
167. ~select_3(X0) | main_query_2(X0) [cnf transformation 145]
171. lte(X0,X0) [cnf transformation 49]
177. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 148]
201. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 151]
202. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 151]
227. 0 <=> main_query_1(sK8) [avatar definition]
228. main_query_1(sK8) <- {0} [avatar component clause 227]
230. 1 <=> ~main_query_1(sK8) [avatar definition]
231. ~main_query_1(sK8) <- {1} [avatar component clause 230]
233. 2 <=> main_query_2(sK8) [avatar definition]
234. main_query_2(sK8) <- {2} [avatar component clause 233]
236. 3 <=> ~main_query_2(sK8) [avatar definition]
237. ~main_query_2(sK8) <- {3} [avatar component clause 236]
238. ~1 | ~3 [avatar split clause 202,236,230]
239. 0 | 2 [avatar split clause 201,233,227]
240. select_3(sK8) <- {2} [resolution 166,234]
248. gte(X1,X1) [resolution 177,171]
250. where_5(X0) | ~lte(X0,X0) [resolution 162,248]
252. where_5(X0) [subsumption resolution 250,171]
258. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 154,152]
260. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 163,158]
262. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 258,159]
265. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 260,153]
267. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 262,165]
268. ~select_1(X1) | select_3(X1) [subsumption resolution 267,252]
270. ~select_3(X1) | select_1(X1) [resolution 265,155]
271. select_1(sK8) <- {2} [resolution 270,240]
273. main_query_1(sK8) <- {2} [resolution 271,157]
274. $false <- {1, 2} [subsumption resolution 273,231]
275. 1 | ~2 [avatar contradiction clause 274]
276. select_1(sK8) <- {0} [resolution 228,156]
277. select_3(sK8) <- {0} [resolution 276,268]
280. main_query_2(sK8) <- {0} [resolution 277,167]
282. $false <- {0, 3} [subsumption resolution 280,237]
283. ~0 | 3 [avatar contradiction clause 282]
284. $false [avatar sat refutation 238,239,275,283]
% SZS output end Proof for query01_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
