% Refutation found. Thanks to Tanya!
% SZS status Theorem for query04_query05_tptp4X
% SZS output start Proof for query04_query05_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X5,X1,X2,X3,X4] : (from_5(X5,X1,X2,X3,X4) <=> tab1(X5,X1,X2,X3,X4)) [input]
6. ! [X5] : (where_6(X5) <=> eq(X5,X5)) [input]
7. ! [X5] : (select_4(X5) <=> ? [X1,X2,X3,X4] : (where_6(X5) & from_5(X5,X1,X2,X3,X4))) [input]
8. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
9. ! [X6] : eq(X6,X6) [input]
42. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
43. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 42]
44. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
45. ! [X0] : (where_6(X0) <=> eq(X0,X0)) [rectify 6]
46. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
47. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
48. ! [X0] : eq(X0,X0) [rectify 9]
132. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 43]
133. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
134. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 135]
137. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 136,137]
139. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
140. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 44]
141. ! [X0] : ((where_6(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_6(X0))) [nnf transformation 45]
142. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 46]
143. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 142]
144. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
145. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 143,144]
146. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 47]
149. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 132]
150. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
151. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 149,150]
152. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 133]
153. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 133]
154. eq(X0,X0) | ~where_3(X0) [cnf transformation 134]
155. ~eq(X0,X0) | where_3(X0) [cnf transformation 134]
156. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 138]
158. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 138]
159. ~main_query_1(X0) | select_1(X0) [cnf transformation 139]
160. ~select_1(X0) | main_query_1(X0) [cnf transformation 139]
161. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 140]
162. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 140]
164. ~eq(X0,X0) | where_6(X0) [cnf transformation 141]
165. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 145]
167. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 145]
168. ~main_query_2(X0) | select_4(X0) [cnf transformation 146]
169. ~select_4(X0) | main_query_2(X0) [cnf transformation 146]
170. eq(X0,X0) [cnf transformation 48]
201. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 151]
202. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 151]
227. 0 <=> main_query_1(sK8) [avatar definition]
228. main_query_1(sK8) <- {0} [avatar component clause 227]
230. 1 <=> ~main_query_1(sK8) [avatar definition]
231. ~main_query_1(sK8) <- {1} [avatar component clause 230]
233. 2 <=> main_query_2(sK8) [avatar definition]
234. main_query_2(sK8) <- {2} [avatar component clause 233]
236. 3 <=> ~main_query_2(sK8) [avatar definition]
237. ~main_query_2(sK8) <- {3} [avatar component clause 236]
238. ~1 | ~3 [avatar split clause 202,236,230]
239. 0 | 2 [avatar split clause 201,233,227]
240. select_4(sK8) <- {2} [resolution 168,234]
244. where_3(X1) [resolution 155,170]
247. where_6(X1) | ~where_3(X1) [resolution 164,154]
249. where_6(X1) [subsumption resolution 247,244]
264. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 156,152]
266. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 165,161]
268. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 264,162]
271. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 266,153]
273. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 268,167]
274. ~select_1(X1) | select_4(X1) [subsumption resolution 273,249]
276. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 271,158]
277. ~select_4(X1) | select_1(X1) [subsumption resolution 276,244]
278. select_1(sK8) <- {2} [resolution 277,240]
280. main_query_1(sK8) <- {2} [resolution 278,160]
282. $false <- {1, 2} [subsumption resolution 280,231]
283. 1 | ~2 [avatar contradiction clause 282]
284. select_1(sK8) <- {0} [resolution 228,159]
285. select_4(sK8) <- {0} [resolution 284,274]
289. main_query_2(sK8) <- {0} [resolution 285,169]
291. $false <- {0, 3} [subsumption resolution 289,237]
292. ~0 | 3 [avatar contradiction clause 291]
293. $false [avatar sat refutation 238,239,283,292]
% SZS output end Proof for query04_query05_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
