% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query05_tptp4X
% SZS output start Proof for query00_query05_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [input]
6. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
8. ! [X6] : eq(X6,X6) [input]
42. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
43. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 42]
44. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
45. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
46. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
47. ! [X0] : eq(X0,X0) [rectify 8]
81. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 43]
133. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 81]
134. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 44]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 135]
137. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 136,137]
139. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 45]
140. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
141. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 5]
142. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 6]
143. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 142]
144. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
145. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 143,144]
146. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 46]
149. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 133]
150. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
151. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 149,150]
152. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 134]
153. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 134]
154. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 138]
155. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 138]
156. ~main_query_1(X0) | select_1(X0) [cnf transformation 139]
157. ~select_1(X0) | main_query_1(X0) [cnf transformation 139]
158. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 140]
159. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 140]
161. ~eq(X0,X0) | where_5(X0) [cnf transformation 141]
162. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 145]
164. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 145]
165. ~main_query_2(X0) | select_3(X0) [cnf transformation 146]
166. ~select_3(X0) | main_query_2(X0) [cnf transformation 146]
167. eq(X0,X0) [cnf transformation 47]
199. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 151]
200. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 151]
202. 0 <=> main_query_1(sK8) [avatar definition]
203. main_query_1(sK8) <- {0} [avatar component clause 202]
205. 1 <=> ~main_query_1(sK8) [avatar definition]
206. ~main_query_1(sK8) <- {1} [avatar component clause 205]
208. 2 <=> main_query_2(sK8) [avatar definition]
209. main_query_2(sK8) <- {2} [avatar component clause 208]
211. 3 <=> ~main_query_2(sK8) [avatar definition]
212. ~main_query_2(sK8) <- {3} [avatar component clause 211]
213. ~1 | ~3 [avatar split clause 200,211,205]
214. 0 | 2 [avatar split clause 199,208,202]
215. select_3(sK8) <- {2} [resolution 165,209]
219. where_5(X1) [resolution 161,167]
253. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 154,152]
255. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 162,158]
257. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 253,159]
261. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 255,153]
264. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 257,164]
265. ~select_1(X1) | select_3(X1) [subsumption resolution 264,219]
273. ~select_3(X9) | select_1(X9) [resolution 261,155]
276. select_1(sK8) <- {2} [resolution 273,215]
280. main_query_1(sK8) <- {2} [resolution 276,157]
281. $false <- {1, 2} [subsumption resolution 280,206]
282. 1 | ~2 [avatar contradiction clause 281]
283. select_1(sK8) <- {0} [resolution 203,156]
286. select_3(sK8) <- {0} [resolution 283,265]
291. main_query_2(sK8) <- {0} [resolution 286,166]
293. $false <- {0, 3} [subsumption resolution 291,212]
294. ~0 | 3 [avatar contradiction clause 293]
295. $false [avatar sat refutation 213,214,282,294]
% SZS output end Proof for query00_query05_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.017 s
% ------------------------------
% ------------------------------
