% Refutation found. Thanks to Tanya!
% SZS status Theorem for query10_query14_tptp4X
% SZS output start Proof for query10_query14_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X1,X0] : (where_3(X1,X0) <=> eq(X0,X1)) [input]
3. ! [X1,X0] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X1,X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (query_7(X5) <=> select_8(X5)) [input]
9. ! [X1,X0] : (where_6(X1,X0) <=> (query_7(X0) & eq(X0,X1))) [input]
10. ! [X1,X0] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X1,X0) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
62. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
63. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 62]
64. ! [X0,X1] : (where_3(X0,X1) <=> eq(X1,X0)) [rectify 2]
65. ! [X0,X1] : (select_1(X1,X0) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X1,X0,X2,X3,X4))) [rectify 3]
66. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
67. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 7]
68. ! [X0] : (query_7(X0) <=> select_8(X0)) [rectify 8]
69. ! [X0,X1] : (where_6(X0,X1) <=> (query_7(X1) & eq(X1,X0))) [rectify 9]
70. ! [X0,X1] : (select_4(X1,X0) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X1,X0,X2,X3,X4))) [rectify 10]
71. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
122. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 63]
208. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 122]
209. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
210. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X1,X0)) & (eq(X1,X0) | ~where_3(X0,X1))) [nnf transformation 64]
211. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X1,X0,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X1,X0,X2,X3,X4)) | ~select_1(X1,X0))) [nnf transformation 65]
212. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X1,X0,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X1,X0,X5,X6,X7)) | ~select_1(X1,X0))) [rectify 211]
213. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X1,X0,X5,X6,X7)) => (where_3(X0,X1) & from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
214. ! [X0,X1] : ((select_1(X1,X0) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X1,X0,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X1,X0))) [skolemisation 212,213]
215. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 66]
216. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
217. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
218. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 67]
219. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 218]
220. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
221. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 219,220]
222. ! [X0] : ((query_7(X0) | ~select_8(X0)) & (select_8(X0) | ~query_7(X0))) [nnf transformation 68]
223. ! [X0,X1] : ((where_6(X0,X1) | (~query_7(X1) | ~eq(X1,X0))) & ((query_7(X1) & eq(X1,X0)) | ~where_6(X0,X1))) [nnf transformation 69]
224. ! [X0,X1] : ((where_6(X0,X1) | ~query_7(X1) | ~eq(X1,X0)) & ((query_7(X1) & eq(X1,X0)) | ~where_6(X0,X1))) [flattening 223]
225. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X1,X0,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X1,X0,X2,X3,X4)) | ~select_4(X1,X0))) [nnf transformation 70]
226. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X1,X0,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X1,X0,X5,X6,X7)) | ~select_4(X1,X0))) [rectify 225]
227. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X1,X0,X5,X6,X7)) => (where_6(X0,X1) & from_5(X1,X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
228. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X1,X0,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X1,X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X1,X0))) [skolemisation 226,227]
229. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 71]
232. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 208]
233. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
234. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 232,233]
235. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 209]
236. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 209]
237. ~where_3(X0,X1) | eq(X1,X0) [cnf transformation 210]
238. ~eq(X1,X0) | where_3(X0,X1) [cnf transformation 210]
239. from_2(X1,X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X1,X0) [cnf transformation 214]
240. ~select_1(X1,X0) | where_3(X0,X1) [cnf transformation 214]
241. ~from_2(X1,X0,X2,X3,X4) | ~where_3(X0,X1) | select_1(X1,X0) [cnf transformation 214]
242. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 215]
243. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 215]
244. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 216]
245. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 216]
247. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 217]
249. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 221]
251. ~select_8(X0) | query_7(X0) [cnf transformation 222]
252. ~where_6(X0,X1) | eq(X1,X0) [cnf transformation 224]
254. ~eq(X1,X0) | ~query_7(X1) | where_6(X0,X1) [cnf transformation 224]
255. from_5(X1,X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X1,X0) [cnf transformation 228]
256. ~select_4(X1,X0) | where_6(X0,X1) [cnf transformation 228]
257. ~from_5(X1,X0,X2,X3,X4) | ~where_6(X0,X1) | select_4(X1,X0) [cnf transformation 228]
258. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 229]
259. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 229]
308. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 234]
309. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 234]
351. 0 <=> main_query_1(sK10,sK11) [avatar definition]
352. main_query_1(sK10,sK11) <- {0} [avatar component clause 351]
354. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
355. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 354]
357. 2 <=> main_query_2(sK10,sK11) [avatar definition]
358. main_query_2(sK10,sK11) <- {2} [avatar component clause 357]
360. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
361. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 360]
362. ~1 | ~3 [avatar split clause 309,360,354]
363. 0 | 2 [avatar split clause 308,357,351]
366. select_4(sK10,sK11) <- {2} [resolution 258,358]
367. where_6(sK11,sK10) <- {2} [resolution 366,256]
369. eq(sK10,sK11) <- {2} [resolution 367,252]
379. where_3(sK11,sK10) <- {2} [resolution 369,238]
427. tab1(X2,X3,sK0(X3,X2),sK1(X3,X2),sK2(X3,X2)) | ~select_1(X2,X3) [resolution 239,235]
431. tab1(X2,X3,sK7(X3,X2),sK8(X3,X2),sK9(X3,X2)) | ~select_4(X2,X3) [resolution 255,244]
432. from_9(X0,X1,sK0(X1,X0),sK1(X1,X0),sK2(X1,X0)) | ~select_1(X0,X1) [resolution 427,247]
433. from_5(X2,X3,sK0(X3,X2),sK1(X3,X2),sK2(X3,X2)) | ~select_1(X2,X3) [resolution 427,245]
439. from_2(X4,X5,sK7(X5,X4),sK8(X5,X4),sK9(X5,X4)) | ~select_4(X4,X5) [resolution 431,236]
441. ~select_1(X2,X3) | select_8(X2) [resolution 432,249]
442. ~where_6(X1,X0) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 433,257]
449. ~select_4(X0,X1) | ~where_3(X1,X0) | select_1(X0,X1) [resolution 439,241]
451. ~where_3(sK11,sK10) | select_1(sK10,sK11) <- {2} [resolution 449,366]
452. select_1(sK10,sK11) <- {2} [subsumption resolution 451,379]
454. main_query_1(sK10,sK11) <- {2} [resolution 452,243]
456. $false <- {1, 2} [subsumption resolution 454,355]
457. 1 | ~2 [avatar contradiction clause 456]
458. select_1(sK10,sK11) <- {0} [resolution 352,242]
459. select_8(sK10) <- {0} [resolution 458,441]
461. where_3(sK11,sK10) <- {0} [resolution 458,240]
462. query_7(sK10) <- {0} [resolution 459,251]
463. eq(sK10,sK11) <- {0} [resolution 461,237]
466. ~query_7(sK10) | where_6(sK11,sK10) <- {0} [resolution 463,254]
471. where_6(sK11,sK10) <- {0} [subsumption resolution 466,462]
482. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {0} [resolution 471,442]
485. select_4(sK10,sK11) <- {0} [subsumption resolution 482,458]
491. main_query_2(sK10,sK11) <- {0} [resolution 485,259]
493. $false <- {0, 3} [subsumption resolution 491,361]
494. ~0 | 3 [avatar contradiction clause 493]
495. $false [avatar sat refutation 362,363,457,494]
% SZS output end Proof for query10_query14_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
