% Refutation found. Thanks to Tanya!
% SZS status Theorem for query06_query09_tptp4X
% SZS output start Proof for query06_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : ($true & from_5(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (subquery_from_3(X5) <=> select_4(X5)) [input]
4. ! [X6] : (from_2(X6) <=> subquery_from_3(X6)) [input]
5. ! [X6] : (select_1(X6) <=> ($true & from_2(X6))) [input]
6. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
7. ! [X7,X1,X2,X3,X4] : (from_10(X7,X1,X2,X3,X4) <=> tab1(X7,X1,X2,X3,X4)) [input]
8. ! [X7] : (select_9(X7) <=> ? [X1,X2,X3,X4] : ($true & from_10(X7,X1,X2,X3,X4))) [input]
9. ! [X7] : (subquery_from_8(X7) <=> select_9(X7)) [input]
10. ! [X8] : (from_7(X8) <=> subquery_from_8(X8)) [input]
11. ! [X8] : (select_6(X8) <=> ($true & from_7(X8))) [input]
12. ! [X7] : (main_query_2(X7) <=> select_6(X7)) [input]
47. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
48. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 47]
49. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4)) [true and false elimination 2]
50. ! [X0] : (subquery_from_3(X0) <=> select_4(X0)) [rectify 3]
51. ! [X0] : (from_2(X0) <=> subquery_from_3(X0)) [rectify 4]
52. ! [X0] : (select_1(X0) <=> ($true & from_2(X0))) [rectify 5]
53. ! [X0] : (select_1(X0) <=> from_2(X0)) [true and false elimination 52]
54. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 6]
55. ! [X0,X1,X2,X3,X4] : (from_10(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 7]
56. ! [X0] : (select_9(X0) <=> ? [X1,X2,X3,X4] : ($true & from_10(X0,X1,X2,X3,X4))) [rectify 8]
57. ! [X0] : (select_9(X0) <=> ? [X1,X2,X3,X4] : from_10(X0,X1,X2,X3,X4)) [true and false elimination 56]
58. ! [X0] : (subquery_from_8(X0) <=> select_9(X0)) [rectify 9]
59. ! [X0] : (from_7(X0) <=> subquery_from_8(X0)) [rectify 10]
60. ! [X0] : (select_6(X0) <=> ($true & from_7(X0))) [rectify 11]
61. ! [X0] : (select_6(X0) <=> from_7(X0)) [true and false elimination 60]
62. ! [X0] : (main_query_2(X0) <=> select_6(X0)) [rectify 12]
97. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 48]
149. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 97]
150. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 1]
151. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4) | ~select_4(X0))) [nnf transformation 49]
152. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) | ~select_4(X0))) [rectify 151]
153. ! [X0] : (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) => from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
154. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0))) [skolemisation 152,153]
155. ! [X0] : ((subquery_from_3(X0) | ~select_4(X0)) & (select_4(X0) | ~subquery_from_3(X0))) [nnf transformation 50]
156. ! [X0] : ((from_2(X0) | ~subquery_from_3(X0)) & (subquery_from_3(X0) | ~from_2(X0))) [nnf transformation 51]
157. ! [X0] : ((select_1(X0) | ~from_2(X0)) & (from_2(X0) | ~select_1(X0))) [nnf transformation 53]
158. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 54]
159. ! [X0,X1,X2,X3,X4] : ((from_10(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_10(X0,X1,X2,X3,X4))) [nnf transformation 55]
160. ! [X0] : ((select_9(X0) | ! [X1,X2,X3,X4] : ~from_10(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_10(X0,X1,X2,X3,X4) | ~select_9(X0))) [nnf transformation 57]
161. ! [X0] : ((select_9(X0) | ! [X1,X2,X3,X4] : ~from_10(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_10(X0,X5,X6,X7,X8) | ~select_9(X0))) [rectify 160]
162. ! [X0] : (? [X5,X6,X7,X8] : from_10(X0,X5,X6,X7,X8) => from_10(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
163. ! [X0] : ((select_9(X0) | ! [X1,X2,X3,X4] : ~from_10(X0,X1,X2,X3,X4)) & (from_10(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_9(X0))) [skolemisation 161,162]
164. ! [X0] : ((subquery_from_8(X0) | ~select_9(X0)) & (select_9(X0) | ~subquery_from_8(X0))) [nnf transformation 58]
165. ! [X0] : ((from_7(X0) | ~subquery_from_8(X0)) & (subquery_from_8(X0) | ~from_7(X0))) [nnf transformation 59]
166. ! [X0] : ((select_6(X0) | ~from_7(X0)) & (from_7(X0) | ~select_6(X0))) [nnf transformation 61]
167. ! [X0] : ((main_query_2(X0) | ~select_6(X0)) & (select_6(X0) | ~main_query_2(X0))) [nnf transformation 62]
170. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 149]
171. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
172. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 170,171]
173. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 150]
174. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 150]
175. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [cnf transformation 154]
176. ~from_5(X0,X1,X2,X3,X4) | select_4(X0) [cnf transformation 154]
177. ~subquery_from_3(X0) | select_4(X0) [cnf transformation 155]
178. ~select_4(X0) | subquery_from_3(X0) [cnf transformation 155]
179. ~from_2(X0) | subquery_from_3(X0) [cnf transformation 156]
180. ~subquery_from_3(X0) | from_2(X0) [cnf transformation 156]
181. ~select_1(X0) | from_2(X0) [cnf transformation 157]
182. ~from_2(X0) | select_1(X0) [cnf transformation 157]
183. ~main_query_1(X0) | select_1(X0) [cnf transformation 158]
184. ~select_1(X0) | main_query_1(X0) [cnf transformation 158]
185. ~from_10(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 159]
186. ~tab1(X0,X1,X2,X3,X4) | from_10(X0,X1,X2,X3,X4) [cnf transformation 159]
187. from_10(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_9(X0) [cnf transformation 163]
188. ~from_10(X0,X1,X2,X3,X4) | select_9(X0) [cnf transformation 163]
189. ~subquery_from_8(X0) | select_9(X0) [cnf transformation 164]
190. ~select_9(X0) | subquery_from_8(X0) [cnf transformation 164]
191. ~from_7(X0) | subquery_from_8(X0) [cnf transformation 165]
192. ~subquery_from_8(X0) | from_7(X0) [cnf transformation 165]
193. ~select_6(X0) | from_7(X0) [cnf transformation 166]
194. ~from_7(X0) | select_6(X0) [cnf transformation 166]
195. ~main_query_2(X0) | select_6(X0) [cnf transformation 167]
196. ~select_6(X0) | main_query_2(X0) [cnf transformation 167]
229. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 172]
230. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 172]
232. 0 <=> main_query_1(sK8) [avatar definition]
233. main_query_1(sK8) <- {0} [avatar component clause 232]
235. 1 <=> ~main_query_1(sK8) [avatar definition]
236. ~main_query_1(sK8) <- {1} [avatar component clause 235]
238. 2 <=> main_query_2(sK8) [avatar definition]
239. main_query_2(sK8) <- {2} [avatar component clause 238]
241. 3 <=> ~main_query_2(sK8) [avatar definition]
242. ~main_query_2(sK8) <- {3} [avatar component clause 241]
243. ~1 | ~3 [avatar split clause 230,241,235]
244. 0 | 2 [avatar split clause 229,238,232]
245. select_6(sK8) <- {2} [resolution 195,239]
247. from_7(sK8) <- {2} [resolution 245,193]
250. subquery_from_8(sK8) <- {2} [resolution 247,191]
253. select_9(sK8) <- {2} [resolution 250,189]
282. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [resolution 175,173]
284. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_9(X0) [resolution 187,185]
286. from_10(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [resolution 282,186]
290. from_5(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_9(X1) [resolution 284,174]
293. ~select_4(X1) | select_9(X1) [resolution 286,188]
301. ~select_9(X9) | select_4(X9) [resolution 290,176]
304. select_4(sK8) <- {2} [resolution 301,253]
306. subquery_from_3(sK8) <- {2} [resolution 304,178]
307. from_2(sK8) <- {2} [resolution 306,180]
309. select_1(sK8) <- {2} [resolution 307,182]
311. main_query_1(sK8) <- {2} [resolution 309,184]
313. $false <- {1, 2} [subsumption resolution 311,236]
314. 1 | ~2 [avatar contradiction clause 313]
315. select_1(sK8) <- {0} [resolution 233,183]
317. from_2(sK8) <- {0} [resolution 315,181]
321. subquery_from_3(sK8) <- {0} [resolution 317,179]
325. select_4(sK8) <- {0} [resolution 321,177]
328. select_9(sK8) <- {0} [resolution 325,293]
333. subquery_from_8(sK8) <- {0} [resolution 328,190]
336. from_7(sK8) <- {0} [resolution 333,192]
340. select_6(sK8) <- {0} [resolution 336,194]
344. main_query_2(sK8) <- {0} [resolution 340,196]
346. $false <- {0, 3} [subsumption resolution 344,242]
347. ~0 | 3 [avatar contradiction clause 346]
348. $false [avatar sat refutation 243,244,314,347]
% SZS output end Proof for query06_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.011 s
% ------------------------------
% ------------------------------
