% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_05_query_06_tptp4X
% SZS output start Proof for query_05_query_06_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
7. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
12. ! [X7] : lte(X7,X7) [input]
18. ! [X7,X8] : (gte(X7,X8) <=> lte(X8,X7)) [input]
47. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
48. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 47]
49. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
50. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
51. ! [X0] : eq(X0,X0) [rectify 9]
54. ! [X0] : lte(X0,X0) [rectify 12]
60. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 18]
89. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 48]
151. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 89]
152. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
153. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
154. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
155. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 154]
156. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
157. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 155,156]
158. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 49]
159. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
160. ! [X0] : ((where_6(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [nnf transformation 6]
161. ! [X0] : ((where_6(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [flattening 160]
162. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 7]
163. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 162]
164. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
165. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 163,164]
166. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 50]
169. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 60]
170. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 151]
171. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
172. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 170,171]
173. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 152]
174. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 152]
176. ~eq(X0,X0) | where_3(X0) [cnf transformation 153]
177. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 157]
179. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0,X1) [cnf transformation 157]
180. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 158]
181. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 158]
182. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 159]
183. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 159]
186. ~gte(X0,X0) | where_6(X0) | ~lte(X0,X0) [cnf transformation 161]
187. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 165]
189. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0,X1) [cnf transformation 165]
190. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 166]
191. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 166]
192. eq(X0,X0) [cnf transformation 51]
195. lte(X0,X0) [cnf transformation 54]
201. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 169]
230. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 172]
231. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 172]
233. 0 <=> main_query_1(sK6,sK7) [avatar definition]
234. main_query_1(sK6,sK7) <- {0} [avatar component clause 233]
236. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
237. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 236]
239. 2 <=> main_query_2(sK6,sK7) [avatar definition]
240. main_query_2(sK6,sK7) <- {2} [avatar component clause 239]
242. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
243. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 242]
244. ~1 | ~3 [avatar split clause 231,242,236]
245. 0 | 2 [avatar split clause 230,239,233]
247. where_3(X1) [resolution 176,192]
248. select_4(sK6,sK7) <- {2} [resolution 190,240]
259. gte(X1,X1) [resolution 201,195]
261. where_6(X0) | ~lte(X0,X0) [resolution 186,259]
263. where_6(X0) [subsumption resolution 261,195]
282. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 177,173]
284. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 187,182]
297. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 282,183]
303. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 284,174]
311. ~select_1(X2,X3) | ~where_6(X2) | select_4(X2,X3) [resolution 297,189]
317. ~select_1(X2,X3) | select_4(X2,X3) [subsumption resolution 311,263]
333. ~select_4(X2,X3) | ~where_3(X2) | select_1(X2,X3) [resolution 303,179]
339. ~select_4(X2,X3) | select_1(X2,X3) [subsumption resolution 333,247]
347. select_1(sK6,sK7) <- {2} [resolution 339,248]
358. main_query_1(sK6,sK7) <- {2} [resolution 347,181]
360. $false <- {1, 2} [subsumption resolution 358,237]
361. 1 | ~2 [avatar contradiction clause 360]
364. select_1(sK6,sK7) <- {0} [resolution 234,180]
365. select_4(sK6,sK7) <- {0} [resolution 364,317]
373. main_query_2(sK6,sK7) <- {0} [resolution 365,191]
375. $false <- {0, 3} [subsumption resolution 373,243]
376. ~0 | 3 [avatar contradiction clause 375]
377. $false [avatar sat refutation 244,245,361,376]
% SZS output end Proof for query_05_query_06_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
