% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_02_query_06_tptp4X
% SZS output start Proof for query_02_query_06_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X5,X1,X2,X3,X4] : (from_4(X5,X1,X2,X3,X4) <=> tab1(X5,X1,X2,X3,X4)) [input]
5. ! [X5] : (where_5(X5) <=> eq(X5,X5)) [input]
6. ! [X5] : (select_3(X5) <=> ? [X1,X2,X3,X4] : (where_5(X5) & from_4(X5,X1,X2,X3,X4))) [input]
7. ! [X6] : (main_query_2(X6) <=> select_3(X6)) [input]
8. ! [X7] : eq(X7,X7) [input]
39. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
40. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 39]
41. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
42. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
43. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [rectify 5]
44. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
45. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
46. ! [X0] : eq(X0,X0) [rectify 8]
122. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 40]
123. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
124. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 41]
125. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 124]
126. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
127. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 125,126]
128. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
129. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 42]
130. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 43]
131. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 44]
132. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 131]
133. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
134. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 132,133]
135. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 45]
138. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 122]
139. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
140. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 138,139]
141. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 123]
142. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 123]
143. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 127]
144. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 127]
145. ~main_query_1(X0) | select_1(X0) [cnf transformation 128]
146. ~select_1(X0) | main_query_1(X0) [cnf transformation 128]
147. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 129]
148. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 129]
150. ~eq(X0,X0) | where_5(X0) [cnf transformation 130]
151. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 134]
153. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 134]
154. ~main_query_2(X0) | select_3(X0) [cnf transformation 135]
155. ~select_3(X0) | main_query_2(X0) [cnf transformation 135]
156. eq(X0,X0) [cnf transformation 46]
185. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 140]
186. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 140]
188. 0 <=> main_query_1(sK8) [avatar definition]
189. main_query_1(sK8) <- {0} [avatar component clause 188]
191. 1 <=> ~main_query_1(sK8) [avatar definition]
192. ~main_query_1(sK8) <- {1} [avatar component clause 191]
194. 2 <=> main_query_2(sK8) [avatar definition]
195. main_query_2(sK8) <- {2} [avatar component clause 194]
197. 3 <=> ~main_query_2(sK8) [avatar definition]
198. ~main_query_2(sK8) <- {3} [avatar component clause 197]
199. ~1 | ~3 [avatar split clause 186,197,191]
200. 0 | 2 [avatar split clause 185,194,188]
201. select_3(sK8) <- {2} [resolution 154,195]
205. where_5(X1) [resolution 150,156]
239. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 143,141]
241. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 151,147]
243. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 239,148]
246. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 241,142]
248. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 243,153]
249. ~select_1(X1) | select_3(X1) [subsumption resolution 248,205]
251. ~select_3(X1) | select_1(X1) [resolution 246,144]
254. select_1(sK8) <- {2} [resolution 251,201]
258. main_query_1(sK8) <- {2} [resolution 254,146]
259. $false <- {1, 2} [subsumption resolution 258,192]
260. 1 | ~2 [avatar contradiction clause 259]
261. select_1(sK8) <- {0} [resolution 189,145]
264. select_3(sK8) <- {0} [resolution 261,249]
269. main_query_2(sK8) <- {0} [resolution 264,155]
271. $false <- {0, 3} [subsumption resolution 269,198]
272. ~0 | 3 [avatar contradiction clause 271]
273. $false [avatar sat refutation 199,200,260,272]
% SZS output end Proof for query_02_query_06_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.020 s
% ------------------------------
% ------------------------------
