% Refutation found. Thanks to Tanya!
% SZS status Theorem for query07_query08_tptp4X
% SZS output start Proof for query07_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
7. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
12. ! [X7] : lte(X7,X7) [input]
15. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
16. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
18. ! [X7,X8] : (gte(X7,X8) <=> lte(X8,X7)) [input]
48. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
49. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 48]
50. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
51. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
52. ! [X0] : eq(X0,X0) [rectify 9]
55. ! [X0] : lte(X0,X0) [rectify 12]
58. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 15]
59. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 16]
61. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 18]
91. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 49]
155. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 91]
156. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
157. ! [X0] : ((where_3(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
158. ! [X0] : ((where_3(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [flattening 157]
159. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
160. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 159]
161. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
162. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 160,161]
163. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 50]
164. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
165. ! [X0] : ((where_6(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [nnf transformation 6]
166. ! [X0] : ((where_6(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [flattening 165]
167. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 7]
168. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 167]
169. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
170. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 168,169]
171. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 51]
174. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 58]
175. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 174]
176. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 59]
177. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 61]
178. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 155]
179. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
180. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 178,179]
181. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 156]
182. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 156]
185. gt(X0,X0) | where_3(X0) | lt(X0,X0) [cnf transformation 158]
186. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 162]
188. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 162]
189. ~main_query_1(X0) | select_1(X0) [cnf transformation 163]
190. ~select_1(X0) | main_query_1(X0) [cnf transformation 163]
191. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 164]
192. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 164]
195. ~gte(X0,X0) | where_6(X0) | ~lte(X0,X0) [cnf transformation 166]
196. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 170]
198. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 170]
199. ~main_query_2(X0) | select_4(X0) [cnf transformation 171]
200. ~select_4(X0) | main_query_2(X0) [cnf transformation 171]
201. eq(X0,X0) [cnf transformation 52]
204. lte(X0,X0) [cnf transformation 55]
210. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 175]
212. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 176]
215. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 177]
245. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 180]
246. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 180]
277. 0 <=> main_query_1(sK8) [avatar definition]
278. main_query_1(sK8) <- {0} [avatar component clause 277]
280. 1 <=> ~main_query_1(sK8) [avatar definition]
281. ~main_query_1(sK8) <- {1} [avatar component clause 280]
283. 2 <=> main_query_2(sK8) [avatar definition]
284. main_query_2(sK8) <- {2} [avatar component clause 283]
286. 3 <=> ~main_query_2(sK8) [avatar definition]
287. ~main_query_2(sK8) <- {3} [avatar component clause 286]
288. ~1 | ~3 [avatar split clause 246,286,280]
289. 0 | 2 [avatar split clause 245,283,277]
290. select_4(sK8) <- {2} [resolution 199,284]
296. ~lt(X0,X0) [resolution 210,201]
299. gte(X1,X1) [resolution 215,204]
302. where_3(X1) | lt(X1,X1) | lt(X1,X1) [resolution 185,212]
303. where_3(X1) | lt(X1,X1) [duplicate literal removal 302]
304. where_3(X1) [subsumption resolution 303,296]
305. where_6(X0) | ~lte(X0,X0) [resolution 195,299]
307. where_6(X0) [subsumption resolution 305,204]
315. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 186,181]
317. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 196,191]
319. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 315,192]
322. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 317,182]
324. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 319,198]
325. ~select_1(X1) | select_4(X1) [subsumption resolution 324,307]
327. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 322,188]
328. ~select_4(X1) | select_1(X1) [subsumption resolution 327,304]
329. select_1(sK8) <- {2} [resolution 328,290]
331. main_query_1(sK8) <- {2} [resolution 329,190]
333. $false <- {1, 2} [subsumption resolution 331,281]
334. 1 | ~2 [avatar contradiction clause 333]
335. select_1(sK8) <- {0} [resolution 278,189]
336. select_4(sK8) <- {0} [resolution 335,325]
340. main_query_2(sK8) <- {0} [resolution 336,200]
342. $false <- {0, 3} [subsumption resolution 340,287]
343. ~0 | 3 [avatar contradiction clause 342]
344. $false [avatar sat refutation 288,289,334,343]
% SZS output end Proof for query07_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.013 s
% ------------------------------
% ------------------------------
