% Refutation found. Thanks to Tanya!
% SZS status Theorem for query10_query12_tptp4X
% SZS output start Proof for query10_query12_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1] : (where_6(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [input]
7. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [input]
13. ! [X5,X6] : ((lte(X6,X5) & lte(X5,X6)) <=> eq(X5,X6)) [input]
54. ! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [input]
55. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [negated conjecture 54]
60. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
176. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 55]
177. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
178. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
179. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
180. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 179]
181. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
182. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 180,181]
183. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 4]
184. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
185. ! [X0,X1] : ((where_6(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 6]
186. ! [X0,X1] : ((where_6(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [flattening 185]
187. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 7]
188. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 187]
189. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
190. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 188,189]
191. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 8]
192. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 60]
193. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 192]
194. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 176]
195. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
196. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 194,195]
197. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 177]
198. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 177]
199. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 178]
200. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 178]
201. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 182]
202. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 182]
203. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 182]
204. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 183]
205. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 183]
206. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 184]
207. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 184]
208. ~where_6(X0,X1) | lte(X0,X1) [cnf transformation 186]
209. ~where_6(X0,X1) | lte(X1,X0) [cnf transformation 186]
210. ~lte(X1,X0) | where_6(X0,X1) | ~lte(X0,X1) [cnf transformation 186]
211. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 190]
212. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 190]
213. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 190]
214. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 191]
215. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 191]
220. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 193]
221. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 193]
222. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 193]
259. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 196]
260. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 196]
297. 0 <=> main_query_1(sK6,sK7) [avatar definition]
298. main_query_1(sK6,sK7) <- {0} [avatar component clause 297]
300. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
301. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 300]
303. 2 <=> main_query_2(sK6,sK7) [avatar definition]
304. main_query_2(sK6,sK7) <- {2} [avatar component clause 303]
306. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
307. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 306]
308. ~1 | ~3 [avatar split clause 260,306,300]
309. 0 | 2 [avatar split clause 259,303,297]
312. select_4(sK6,sK7) <- {2} [resolution 214,304]
313. where_6(sK6,sK7) <- {2} [resolution 312,212]
315. lte(sK7,sK6) <- {2} [resolution 313,209]
316. lte(sK6,sK7) <- {2} [resolution 313,208]
329. eq(sK6,sK7) | ~lte(sK6,sK7) <- {2} [resolution 220,315]
331. eq(sK6,sK7) <- {2} [subsumption resolution 329,316]
342. where_3(sK6,sK7) <- {2} [resolution 331,200]
357. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 201,197]
361. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 211,206]
362. from_5(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 357,207]
365. from_2(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 361,198]
366. ~where_6(X0,X1) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 362,213]
368. ~select_4(X0,X1) | ~where_3(X0,X1) | select_1(X0,X1) [resolution 365,203]
386. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 368,312]
387. select_1(sK6,sK7) <- {2} [subsumption resolution 386,342]
388. main_query_1(sK6,sK7) <- {2} [resolution 387,205]
390. $false <- {1, 2} [subsumption resolution 388,301]
391. 1 | ~2 [avatar contradiction clause 390]
392. select_1(sK6,sK7) <- {0} [resolution 298,204]
394. where_3(sK6,sK7) <- {0} [resolution 392,202]
395. eq(sK6,sK7) <- {0} [resolution 394,199]
397. lte(sK7,sK6) <- {0} [resolution 395,222]
398. lte(sK6,sK7) <- {0} [resolution 395,221]
403. where_6(sK6,sK7) | ~lte(sK6,sK7) <- {0} [resolution 397,210]
404. where_6(sK6,sK7) <- {0} [subsumption resolution 403,398]
416. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 404,366]
419. select_4(sK6,sK7) <- {0} [subsumption resolution 416,392]
427. main_query_2(sK6,sK7) <- {0} [resolution 419,215]
429. $false <- {0, 3} [subsumption resolution 427,307]
430. ~0 | 3 [avatar contradiction clause 429]
431. $false [avatar sat refutation 308,309,391,430]
% SZS output end Proof for query10_query12_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.011 s
% ------------------------------
% ------------------------------
