% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query07_tptp4X
% SZS output start Proof for query00_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
6. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X6] : (main_query_2(X6) <=> select_3(X6)) [input]
8. ! [X7] : eq(X7,X7) [input]
14. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
15. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
46. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
47. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 46]
48. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
50. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
51. ! [X0] : eq(X0,X0) [rectify 8]
57. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 14]
58. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 15]
89. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 47]
151. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 89]
152. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
153. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 48]
154. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 153]
155. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
156. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 154,155]
157. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 49]
158. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
159. ! [X0] : ((where_5(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [nnf transformation 5]
160. ! [X0] : ((where_5(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [flattening 159]
161. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 6]
162. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 161]
163. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
164. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 162,163]
165. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 50]
168. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 57]
169. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 168]
170. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 58]
171. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 151]
172. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
173. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 171,172]
174. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 152]
175. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 152]
176. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 156]
177. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 156]
178. ~main_query_1(X0) | select_1(X0) [cnf transformation 157]
179. ~select_1(X0) | main_query_1(X0) [cnf transformation 157]
180. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 158]
181. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 158]
184. gt(X0,X0) | where_5(X0) | lt(X0,X0) [cnf transformation 160]
185. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 164]
187. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 164]
188. ~main_query_2(X0) | select_3(X0) [cnf transformation 165]
189. ~select_3(X0) | main_query_2(X0) [cnf transformation 165]
190. eq(X0,X0) [cnf transformation 51]
199. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 169]
201. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 170]
231. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 173]
232. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 173]
262. 0 <=> main_query_1(sK8) [avatar definition]
263. main_query_1(sK8) <- {0} [avatar component clause 262]
265. 1 <=> ~main_query_1(sK8) [avatar definition]
266. ~main_query_1(sK8) <- {1} [avatar component clause 265]
268. 2 <=> main_query_2(sK8) [avatar definition]
269. main_query_2(sK8) <- {2} [avatar component clause 268]
271. 3 <=> ~main_query_2(sK8) [avatar definition]
272. ~main_query_2(sK8) <- {3} [avatar component clause 271]
273. ~1 | ~3 [avatar split clause 232,271,265]
274. 0 | 2 [avatar split clause 231,268,262]
275. select_3(sK8) <- {2} [resolution 188,269]
281. ~lt(X0,X0) [resolution 199,190]
283. where_5(X1) | lt(X1,X1) | lt(X1,X1) [resolution 184,201]
284. where_5(X1) | lt(X1,X1) [duplicate literal removal 283]
285. where_5(X1) [subsumption resolution 284,281]
290. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 176,174]
292. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 185,180]
294. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 290,181]
297. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 292,175]
299. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 294,187]
300. ~select_1(X1) | select_3(X1) [subsumption resolution 299,285]
302. ~select_3(X1) | select_1(X1) [resolution 297,177]
303. select_1(sK8) <- {2} [resolution 302,275]
305. main_query_1(sK8) <- {2} [resolution 303,179]
306. $false <- {1, 2} [subsumption resolution 305,266]
307. 1 | ~2 [avatar contradiction clause 306]
308. select_1(sK8) <- {0} [resolution 263,178]
309. select_3(sK8) <- {0} [resolution 308,300]
312. main_query_2(sK8) <- {0} [resolution 309,189]
314. $false <- {0, 3} [subsumption resolution 312,272]
315. ~0 | 3 [avatar contradiction clause 314]
316. $false [avatar sat refutation 273,274,307,315]
% SZS output end Proof for query00_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.016 s
% ------------------------------
% ------------------------------
