% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query02_tptp4X
% SZS output start Proof for query00_query02_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_4(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (select_3(X6) <=> ? [X7,X8,X9,X10] : ($true & from_4(X6,X7,X8,X9,X10))) [input]
6. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
41. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
42. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
45. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
46. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [rectify 5]
47. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 46]
48. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 6]
83. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 42]
135. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 83]
136. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 137]
139. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
140. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 138,139]
141. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 44]
142. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 45]
143. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 47]
144. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 143]
145. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
146. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0))) [skolemisation 144,145]
147. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 48]
150. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 135]
151. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
152. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 150,151]
153. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 136]
154. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 136]
155. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 140]
156. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 140]
157. ~main_query_1(X0) | select_1(X0) [cnf transformation 141]
158. ~select_1(X0) | main_query_1(X0) [cnf transformation 141]
159. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 142]
160. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 142]
161. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 146]
162. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 146]
163. ~main_query_2(X0) | select_3(X0) [cnf transformation 147]
164. ~select_3(X0) | main_query_2(X0) [cnf transformation 147]
197. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 152]
198. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 152]
200. 0 <=> main_query_1(sK8) [avatar definition]
201. main_query_1(sK8) <- {0} [avatar component clause 200]
203. 1 <=> ~main_query_1(sK8) [avatar definition]
204. ~main_query_1(sK8) <- {1} [avatar component clause 203]
206. 2 <=> main_query_2(sK8) [avatar definition]
207. main_query_2(sK8) <- {2} [avatar component clause 206]
209. 3 <=> ~main_query_2(sK8) [avatar definition]
210. ~main_query_2(sK8) <- {3} [avatar component clause 209]
211. ~1 | ~3 [avatar split clause 198,209,203]
212. 0 | 2 [avatar split clause 197,206,200]
213. select_3(sK8) <- {2} [resolution 163,207]
244. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 155,153]
246. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 161,159]
249. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 244,160]
253. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 246,154]
259. ~select_1(X7) | select_3(X7) [resolution 249,162]
265. ~select_3(X1) | select_1(X1) [resolution 253,156]
266. select_1(sK8) <- {2} [resolution 265,213]
268. main_query_1(sK8) <- {2} [resolution 266,158]
269. $false <- {1, 2} [subsumption resolution 268,204]
270. 1 | ~2 [avatar contradiction clause 269]
271. select_1(sK8) <- {0} [resolution 201,157]
272. select_3(sK8) <- {0} [resolution 271,259]
275. main_query_2(sK8) <- {0} [resolution 272,164]
276. $false <- {0, 3} [subsumption resolution 275,210]
277. ~0 | 3 [avatar contradiction clause 276]
278. $false [avatar sat refutation 211,212,270,277]
% SZS output end Proof for query00_query02_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.026 s
% ------------------------------
% ------------------------------
