% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_12_query_15_tptp4X
% SZS output start Proof for query_12_query_15_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8,X9,X10,X11] : (from_9(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
7. ! [X7] : (select_8(X7) <=> ? [X8,X9,X10,X11] : ($true & from_9(X7,X8,X9,X10,X11))) [input]
8. ! [X5] : (where_6_subquery_7(X5) <=> select_8(X5)) [input]
9. ! [X7,X8] : (where_6(X7,X8) <=> (where_6_subquery_7(X7) & eq(X7,X8))) [input]
10. ! [X7,X8] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X7,X8) & from_5(X7,X8,X9,X10,X11))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
59. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
60. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 59]
61. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
62. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
63. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 6]
64. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [rectify 7]
65. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 64]
66. ! [X0] : (where_6_subquery_7(X0) <=> select_8(X0)) [rectify 8]
67. ! [X0,X1] : (where_6(X0,X1) <=> (where_6_subquery_7(X0) & eq(X0,X1))) [rectify 9]
68. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 10]
69. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
117. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 60]
197. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 117]
198. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
199. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
200. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
201. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 200]
202. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
203. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 201,202]
204. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 61]
205. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 62]
206. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 63]
207. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 65]
208. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 207]
209. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
210. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 208,209]
211. ! [X0] : ((where_6_subquery_7(X0) | ~select_8(X0)) & (select_8(X0) | ~where_6_subquery_7(X0))) [nnf transformation 66]
212. ! [X0,X1] : ((where_6(X0,X1) | (~where_6_subquery_7(X0) | ~eq(X0,X1))) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 67]
213. ! [X0,X1] : ((where_6(X0,X1) | ~where_6_subquery_7(X0) | ~eq(X0,X1)) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 212]
214. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 68]
215. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 214]
216. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
217. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 215,216]
218. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 69]
221. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 197]
222. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
223. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 221,222]
224. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 198]
225. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 198]
226. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 199]
227. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 199]
228. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 203]
229. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 203]
230. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 203]
231. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 204]
232. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 204]
233. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 205]
234. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 205]
236. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 206]
238. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 210]
240. ~select_8(X0) | where_6_subquery_7(X0) [cnf transformation 211]
241. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 213]
243. ~eq(X0,X1) | ~where_6_subquery_7(X0) | where_6(X0,X1) [cnf transformation 213]
244. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 217]
245. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 217]
246. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 217]
247. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 218]
248. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 218]
294. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 223]
295. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 223]
297. 0 <=> main_query_1(sK10,sK11) [avatar definition]
298. main_query_1(sK10,sK11) <- {0} [avatar component clause 297]
300. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
301. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 300]
303. 2 <=> main_query_2(sK10,sK11) [avatar definition]
304. main_query_2(sK10,sK11) <- {2} [avatar component clause 303]
306. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
307. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 306]
308. ~1 | ~3 [avatar split clause 295,306,300]
309. 0 | 2 [avatar split clause 294,303,297]
380. select_1(sK10,sK11) <- {0} [resolution 298,231]
383. where_3(sK10,sK11) <- {0} [resolution 380,229]
385. eq(sK10,sK11) <- {0} [resolution 383,226]
388. ~where_6_subquery_7(sK10) | where_6(sK10,sK11) <- {0} [resolution 385,243]
399. 8 <=> where_6(sK10,sK11) [avatar definition]
400. where_6(sK10,sK11) <- {8} [avatar component clause 399]
402. 10 <=> where_6_subquery_7(sK10) [avatar definition]
405. 11 <=> ~where_6_subquery_7(sK10) [avatar definition]
407. 8 | ~11 | ~0 [avatar split clause 388,297,405,399]
444. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 228,224]
448. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 244,233]
532. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 444,234]
533. from_9(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 444,236]
539. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 448,225]
571. 14 <=> select_8(sK10) [avatar definition]
572. select_8(sK10) <- {14} [avatar component clause 571]
682. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 532,246]
713. ~select_1(X2,X3) | select_8(X2) [resolution 533,238]
726. select_8(sK10) <- {0} [resolution 713,380]
741. 14 | ~0 [avatar split clause 726,297,571]
762. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 539,230]
768. where_6_subquery_7(sK10) <- {14} [resolution 572,240]
771. 10 | ~14 [avatar split clause 768,571,402]
804. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {8} [resolution 400,682]
809. select_4(sK10,sK11) <- {0, 8} [subsumption resolution 804,380]
834. main_query_2(sK10,sK11) <- {0, 8} [resolution 809,248]
836. $false <- {0, 3, 8} [subsumption resolution 834,307]
837. ~0 | 3 | ~8 [avatar contradiction clause 836]
853. select_4(sK10,sK11) <- {2} [resolution 304,247]
857. where_6(sK10,sK11) <- {2} [resolution 853,245]
858. 8 | ~2 [avatar split clause 857,303,399]
862. eq(sK10,sK11) <- {8} [resolution 400,241]
865. where_3(sK10,sK11) <- {8} [resolution 862,227]
1220. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 762,853]
1224. select_1(sK10,sK11) <- {2, 8} [subsumption resolution 1220,865]
1236. main_query_1(sK10,sK11) <- {2, 8} [resolution 1224,232]
1242. $false <- {1, 2, 8} [subsumption resolution 1236,301]
1243. 1 | ~2 | ~8 [avatar contradiction clause 1242]
1244. $false [avatar sat refutation 308,309,407,741,771,837,858,1243]
% SZS output end Proof for query_12_query_15_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.011 s
% ------------------------------
% ------------------------------
