% Refutation found. Thanks to Tanya!
% SZS status Theorem for query01_query09_tptp4X
% SZS output start Proof for query01_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X5,X6,X7,X8] : (from_7(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
5. ! [X0] : (select_6(X0) <=> ? [X5,X6,X7,X8] : ($true & from_7(X0,X5,X6,X7,X8))) [input]
6. ! [X0] : (subquery_from_5(X0) <=> select_6(X0)) [input]
7. ! [X9] : (from_4(X9) <=> subquery_from_5(X9)) [input]
8. ! [X9] : (select_3(X9) <=> ($true & from_4(X9))) [input]
9. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
44. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
45. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 44]
46. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
47. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
48. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [rectify 5]
49. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 48]
50. ! [X0] : (from_4(X0) <=> subquery_from_5(X0)) [rectify 7]
51. ! [X0] : (select_3(X0) <=> ($true & from_4(X0))) [rectify 8]
52. ! [X0] : (select_3(X0) <=> from_4(X0)) [true and false elimination 51]
138. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 45]
139. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
140. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 46]
141. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 140]
142. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
143. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 141,142]
144. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
145. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 47]
146. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 49]
147. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 146]
148. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
149. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 147,148]
150. ! [X0] : ((subquery_from_5(X0) | ~select_6(X0)) & (select_6(X0) | ~subquery_from_5(X0))) [nnf transformation 6]
151. ! [X0] : ((from_4(X0) | ~subquery_from_5(X0)) & (subquery_from_5(X0) | ~from_4(X0))) [nnf transformation 50]
152. ! [X0] : ((select_3(X0) | ~from_4(X0)) & (from_4(X0) | ~select_3(X0))) [nnf transformation 52]
153. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 9]
156. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 138]
157. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
158. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 156,157]
159. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 139]
160. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 139]
161. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 143]
162. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 143]
163. ~main_query_1(X0) | select_1(X0) [cnf transformation 144]
164. ~select_1(X0) | main_query_1(X0) [cnf transformation 144]
165. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 145]
166. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 145]
167. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 149]
168. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 149]
169. ~subquery_from_5(X0) | select_6(X0) [cnf transformation 150]
170. ~select_6(X0) | subquery_from_5(X0) [cnf transformation 150]
171. ~from_4(X0) | subquery_from_5(X0) [cnf transformation 151]
172. ~subquery_from_5(X0) | from_4(X0) [cnf transformation 151]
173. ~select_3(X0) | from_4(X0) [cnf transformation 152]
174. ~from_4(X0) | select_3(X0) [cnf transformation 152]
175. ~main_query_2(X0) | select_3(X0) [cnf transformation 153]
176. ~select_3(X0) | main_query_2(X0) [cnf transformation 153]
209. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 158]
210. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 158]
212. 0 <=> main_query_1(sK8) [avatar definition]
213. main_query_1(sK8) <- {0} [avatar component clause 212]
215. 1 <=> ~main_query_1(sK8) [avatar definition]
216. ~main_query_1(sK8) <- {1} [avatar component clause 215]
218. 2 <=> main_query_2(sK8) [avatar definition]
219. main_query_2(sK8) <- {2} [avatar component clause 218]
221. 3 <=> ~main_query_2(sK8) [avatar definition]
222. ~main_query_2(sK8) <- {3} [avatar component clause 221]
223. ~1 | ~3 [avatar split clause 210,221,215]
224. 0 | 2 [avatar split clause 209,218,212]
225. select_3(sK8) <- {2} [resolution 175,219]
227. from_4(sK8) <- {2} [resolution 225,173]
230. subquery_from_5(sK8) <- {2} [resolution 227,171]
233. select_6(sK8) <- {2} [resolution 230,169]
262. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 161,159]
264. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 167,165]
266. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 262,166]
270. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 264,160]
273. ~select_1(X1) | select_6(X1) [resolution 266,168]
281. ~select_6(X9) | select_1(X9) [resolution 270,162]
284. select_1(sK8) <- {2} [resolution 281,233]
288. main_query_1(sK8) <- {2} [resolution 284,164]
289. $false <- {1, 2} [subsumption resolution 288,216]
290. 1 | ~2 [avatar contradiction clause 289]
291. select_1(sK8) <- {0} [resolution 213,163]
294. select_6(sK8) <- {0} [resolution 291,273]
299. subquery_from_5(sK8) <- {0} [resolution 294,170]
302. from_4(sK8) <- {0} [resolution 299,172]
306. select_3(sK8) <- {0} [resolution 302,174]
308. main_query_2(sK8) <- {0} [resolution 306,176]
310. $false <- {0, 3} [subsumption resolution 308,222]
311. ~0 | 3 [avatar contradiction clause 310]
312. $false [avatar sat refutation 223,224,290,311]
% SZS output end Proof for query01_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.019 s
% ------------------------------
% ------------------------------
