% Refutation found. Thanks to Tanya!
% SZS status Theorem for query12_query14_tptp4X
% SZS output start Proof for query12_query14_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X1,X0] : (where_3(X1,X0) <=> (lte(X1,X0) & lte(X0,X1))) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X1,X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (query_7(X5) <=> select_8(X5)) [input]
9. ! [X1,X0] : (where_6(X1,X0) <=> (query_7(X0) & eq(X0,X1))) [input]
10. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X1,X0) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
16. ! [X7,X8] : ((lte(X8,X7) & lte(X7,X8)) <=> eq(X7,X8)) [input]
64. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
65. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 64]
66. ! [X0,X1] : (where_3(X0,X1) <=> (lte(X0,X1) & lte(X1,X0))) [rectify 2]
67. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
68. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 7]
69. ! [X0] : (query_7(X0) <=> select_8(X0)) [rectify 8]
70. ! [X0,X1] : (where_6(X0,X1) <=> (query_7(X1) & eq(X1,X0))) [rectify 9]
71. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
76. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 16]
124. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 65]
214. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 124]
215. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
216. ! [X0,X1] : ((where_3(X0,X1) | (~lte(X0,X1) | ~lte(X1,X0))) & ((lte(X0,X1) & lte(X1,X0)) | ~where_3(X0,X1))) [nnf transformation 66]
217. ! [X0,X1] : ((where_3(X0,X1) | ~lte(X0,X1) | ~lte(X1,X0)) & ((lte(X0,X1) & lte(X1,X0)) | ~where_3(X0,X1))) [flattening 216]
218. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X1,X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X1,X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
219. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X1,X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X1,X0) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 218]
220. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X1,X0) & from_2(X0,X1,X5,X6,X7)) => (where_3(X1,X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
221. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X1,X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X1,X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 219,220]
222. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 67]
223. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
224. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
225. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 68]
226. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 225]
227. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
228. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 226,227]
229. ! [X0] : ((query_7(X0) | ~select_8(X0)) & (select_8(X0) | ~query_7(X0))) [nnf transformation 69]
230. ! [X0,X1] : ((where_6(X0,X1) | (~query_7(X1) | ~eq(X1,X0))) & ((query_7(X1) & eq(X1,X0)) | ~where_6(X0,X1))) [nnf transformation 70]
231. ! [X0,X1] : ((where_6(X0,X1) | ~query_7(X1) | ~eq(X1,X0)) & ((query_7(X1) & eq(X1,X0)) | ~where_6(X0,X1))) [flattening 230]
232. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X1,X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X1,X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 10]
233. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X1,X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X1,X0) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 232]
234. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X1,X0) & from_5(X0,X1,X5,X6,X7)) => (where_6(X1,X0) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
235. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X1,X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X1,X0) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 233,234]
236. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 71]
237. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 76]
238. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 237]
239. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 214]
240. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
241. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 239,240]
242. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 215]
243. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 215]
244. ~where_3(X0,X1) | lte(X1,X0) [cnf transformation 217]
245. ~where_3(X0,X1) | lte(X0,X1) [cnf transformation 217]
246. ~lte(X1,X0) | ~lte(X0,X1) | where_3(X0,X1) [cnf transformation 217]
247. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 221]
248. ~select_1(X0,X1) | where_3(X1,X0) [cnf transformation 221]
249. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X1,X0) | select_1(X0,X1) [cnf transformation 221]
250. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 222]
251. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 222]
252. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 223]
253. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 223]
255. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 224]
257. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 228]
259. ~select_8(X0) | query_7(X0) [cnf transformation 229]
260. ~where_6(X0,X1) | eq(X1,X0) [cnf transformation 231]
262. ~eq(X1,X0) | ~query_7(X1) | where_6(X0,X1) [cnf transformation 231]
263. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 235]
264. ~select_4(X0,X1) | where_6(X1,X0) [cnf transformation 235]
265. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X1,X0) | select_4(X0,X1) [cnf transformation 235]
266. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 236]
267. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 236]
272. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 238]
273. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 238]
274. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 238]
318. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 241]
319. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 241]
363. 0 <=> main_query_1(sK10,sK11) [avatar definition]
364. main_query_1(sK10,sK11) <- {0} [avatar component clause 363]
366. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
367. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 366]
369. 2 <=> main_query_2(sK10,sK11) [avatar definition]
370. main_query_2(sK10,sK11) <- {2} [avatar component clause 369]
372. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
373. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 372]
374. ~1 | ~3 [avatar split clause 319,372,366]
375. 0 | 2 [avatar split clause 318,369,363]
376. select_4(sK10,sK11) <- {2} [resolution 266,370]
378. where_6(sK11,sK10) <- {2} [resolution 376,264]
380. eq(sK10,sK11) <- {2} [resolution 378,260]
386. lte(sK11,sK10) <- {2} [resolution 380,274]
387. lte(sK10,sK11) <- {2} [resolution 380,273]
399. ~lte(sK11,sK10) | where_3(sK11,sK10) <- {2} [resolution 387,246]
400. where_3(sK11,sK10) <- {2} [subsumption resolution 399,386]
442. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 247,242]
447. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 263,252]
450. from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 442,255]
451. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 442,253]
455. from_2(X4,X5,sK7(X4,X5),sK8(X4,X5),sK9(X4,X5)) | ~select_4(X4,X5) [resolution 447,243]
459. ~select_1(X2,X3) | select_8(X2) [resolution 450,257]
460. ~where_6(X1,X0) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 451,265]
467. ~select_4(X0,X1) | ~where_3(X1,X0) | select_1(X0,X1) [resolution 455,249]
469. ~where_3(sK11,sK10) | select_1(sK10,sK11) <- {2} [resolution 467,376]
470. select_1(sK10,sK11) <- {2} [subsumption resolution 469,400]
472. main_query_1(sK10,sK11) <- {2} [resolution 470,251]
474. $false <- {1, 2} [subsumption resolution 472,367]
475. 1 | ~2 [avatar contradiction clause 474]
476. select_1(sK10,sK11) <- {0} [resolution 364,250]
477. select_8(sK10) <- {0} [resolution 476,459]
479. where_3(sK11,sK10) <- {0} [resolution 476,248]
480. query_7(sK10) <- {0} [resolution 477,259]
481. lte(sK11,sK10) <- {0} [resolution 479,245]
482. lte(sK10,sK11) <- {0} [resolution 479,244]
485. eq(sK10,sK11) | ~lte(sK10,sK11) <- {0} [resolution 481,272]
487. eq(sK10,sK11) <- {0} [subsumption resolution 485,482]
494. ~query_7(sK10) | where_6(sK11,sK10) <- {0} [resolution 487,262]
498. where_6(sK11,sK10) <- {0} [subsumption resolution 494,480]
508. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {0} [resolution 498,460]
511. select_4(sK10,sK11) <- {0} [subsumption resolution 508,476]
516. main_query_2(sK10,sK11) <- {0} [resolution 511,267]
518. $false <- {0, 3} [subsumption resolution 516,373]
519. ~0 | 3 [avatar contradiction clause 518]
520. $false [avatar sat refutation 374,375,475,519]
% SZS output end Proof for query12_query14_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
