% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_11_query_16_tptp4X
% SZS output start Proof for query_11_query_16_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1] : (where_9(X0,X1) <=> eq(X0,X1)) [input]
7. ! [X0,X1] : (select_7(X0,X1) <=> ? [X2,X3,X4] : (where_9(X0,X1) & from_8(X0,X1,X2,X3,X4))) [input]
8. ! [X5,X6] : (subquery_from_6(X5,X6) <=> select_7(X5,X6)) [input]
9. ! [X7,X8] : (from_5(X7,X8) <=> subquery_from_6(X7,X8)) [input]
10. ! [X7,X8] : (select_4(X7,X8) <=> ($true & from_5(X7,X8))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
60. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
61. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 60]
62. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
63. ! [X0,X1] : (subquery_from_6(X0,X1) <=> select_7(X0,X1)) [rectify 8]
64. ! [X0,X1] : (from_5(X0,X1) <=> subquery_from_6(X0,X1)) [rectify 9]
65. ! [X0,X1] : (select_4(X0,X1) <=> ($true & from_5(X0,X1))) [rectify 10]
66. ! [X0,X1] : (select_4(X0,X1) <=> from_5(X0,X1)) [true and false elimination 65]
67. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
116. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 61]
198. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 116]
199. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
200. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
201. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
202. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 201]
203. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
204. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 202,203]
205. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 62]
206. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 5]
207. ! [X0,X1] : ((where_9(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_9(X0,X1))) [nnf transformation 6]
208. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : (~where_9(X0,X1) | ~from_8(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_9(X0,X1) & from_8(X0,X1,X2,X3,X4)) | ~select_7(X0,X1))) [nnf transformation 7]
209. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : (~where_9(X0,X1) | ~from_8(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_9(X0,X1) & from_8(X0,X1,X5,X6,X7)) | ~select_7(X0,X1))) [rectify 208]
210. ! [X1,X0] : (? [X5,X6,X7] : (where_9(X0,X1) & from_8(X0,X1,X5,X6,X7)) => (where_9(X0,X1) & from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
211. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : (~where_9(X0,X1) | ~from_8(X0,X1,X2,X3,X4))) & ((where_9(X0,X1) & from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_7(X0,X1))) [skolemisation 209,210]
212. ! [X0,X1] : ((subquery_from_6(X0,X1) | ~select_7(X0,X1)) & (select_7(X0,X1) | ~subquery_from_6(X0,X1))) [nnf transformation 63]
213. ! [X0,X1] : ((from_5(X0,X1) | ~subquery_from_6(X0,X1)) & (subquery_from_6(X0,X1) | ~from_5(X0,X1))) [nnf transformation 64]
214. ! [X0,X1] : ((select_4(X0,X1) | ~from_5(X0,X1)) & (from_5(X0,X1) | ~select_4(X0,X1))) [nnf transformation 66]
215. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 67]
218. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 198]
219. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
220. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 218,219]
221. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 199]
222. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 199]
223. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 200]
224. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 200]
225. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 204]
226. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 204]
227. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 204]
228. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 205]
229. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 205]
230. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 206]
231. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 206]
232. ~where_9(X0,X1) | eq(X0,X1) [cnf transformation 207]
233. ~eq(X0,X1) | where_9(X0,X1) [cnf transformation 207]
234. from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1) [cnf transformation 211]
235. ~select_7(X0,X1) | where_9(X0,X1) [cnf transformation 211]
236. ~from_8(X0,X1,X2,X3,X4) | ~where_9(X0,X1) | select_7(X0,X1) [cnf transformation 211]
237. ~subquery_from_6(X0,X1) | select_7(X0,X1) [cnf transformation 212]
238. ~select_7(X0,X1) | subquery_from_6(X0,X1) [cnf transformation 212]
239. ~from_5(X0,X1) | subquery_from_6(X0,X1) [cnf transformation 213]
240. ~subquery_from_6(X0,X1) | from_5(X0,X1) [cnf transformation 213]
241. ~select_4(X0,X1) | from_5(X0,X1) [cnf transformation 214]
242. ~from_5(X0,X1) | select_4(X0,X1) [cnf transformation 214]
243. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 215]
244. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 215]
291. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 220]
292. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 220]
294. 0 <=> main_query_1(sK6,sK7) [avatar definition]
295. main_query_1(sK6,sK7) <- {0} [avatar component clause 294]
297. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
298. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 297]
300. 2 <=> main_query_2(sK6,sK7) [avatar definition]
301. main_query_2(sK6,sK7) <- {2} [avatar component clause 300]
303. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
304. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 303]
305. ~1 | ~3 [avatar split clause 292,303,297]
306. 0 | 2 [avatar split clause 291,300,294]
311. select_4(sK6,sK7) <- {2} [resolution 243,301]
313. from_5(sK6,sK7) <- {2} [resolution 311,241]
316. subquery_from_6(sK6,sK7) <- {2} [resolution 313,239]
319. select_7(sK6,sK7) <- {2} [resolution 316,237]
322. where_9(sK6,sK7) <- {2} [resolution 319,235]
324. eq(sK6,sK7) <- {2} [resolution 322,232]
331. where_3(sK6,sK7) <- {2} [resolution 324,224]
426. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 225,221]
430. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_7(X2,X3) [resolution 234,230]
493. from_8(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 426,231]
501. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1) [resolution 430,222]
601. ~where_9(X2,X3) | ~select_1(X2,X3) | select_7(X2,X3) [resolution 493,236]
610. ~select_7(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 501,227]
745. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 610,319]
749. select_1(sK6,sK7) <- {2} [subsumption resolution 745,331]
764. main_query_1(sK6,sK7) <- {2} [resolution 749,229]
766. $false <- {1, 2} [subsumption resolution 764,298]
767. 1 | ~2 [avatar contradiction clause 766]
770. select_1(sK6,sK7) <- {0} [resolution 295,228]
776. where_3(sK6,sK7) <- {0} [resolution 770,226]
777. eq(sK6,sK7) <- {0} [resolution 776,223]
780. where_9(sK6,sK7) <- {0} [resolution 777,233]
790. ~select_1(sK6,sK7) | select_7(sK6,sK7) <- {0} [resolution 780,601]
794. select_7(sK6,sK7) <- {0} [subsumption resolution 790,770]
872. subquery_from_6(sK6,sK7) <- {0} [resolution 794,238]
917. from_5(sK6,sK7) <- {0} [resolution 872,240]
948. select_4(sK6,sK7) <- {0} [resolution 917,242]
981. main_query_2(sK6,sK7) <- {0} [resolution 948,244]
983. $false <- {0, 3} [subsumption resolution 981,304]
984. ~0 | 3 [avatar contradiction clause 983]
985. $false [avatar sat refutation 305,306,767,984]
% SZS output end Proof for query_11_query_16_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
