% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query09_tptp4X
% SZS output start Proof for query00_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_7(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (select_6(X6) <=> ? [X7,X8,X9,X10] : ($true & from_7(X6,X7,X8,X9,X10))) [input]
6. ! [X6] : (subquery_from_5(X6) <=> select_6(X6)) [input]
7. ! [X11] : (from_4(X11) <=> subquery_from_5(X11)) [input]
8. ! [X11] : (select_3(X11) <=> ($true & from_4(X11))) [input]
9. ! [X6] : (main_query_2(X6) <=> select_3(X6)) [input]
44. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
45. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 44]
46. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
47. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
48. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
49. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [rectify 5]
50. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 49]
51. ! [X0] : (subquery_from_5(X0) <=> select_6(X0)) [rectify 6]
52. ! [X0] : (from_4(X0) <=> subquery_from_5(X0)) [rectify 7]
53. ! [X0] : (select_3(X0) <=> ($true & from_4(X0))) [rectify 8]
54. ! [X0] : (select_3(X0) <=> from_4(X0)) [true and false elimination 53]
55. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 9]
90. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 45]
142. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 90]
143. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
144. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 46]
145. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 144]
146. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
147. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 145,146]
148. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 47]
149. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 48]
150. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 50]
151. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 150]
152. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
153. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 151,152]
154. ! [X0] : ((subquery_from_5(X0) | ~select_6(X0)) & (select_6(X0) | ~subquery_from_5(X0))) [nnf transformation 51]
155. ! [X0] : ((from_4(X0) | ~subquery_from_5(X0)) & (subquery_from_5(X0) | ~from_4(X0))) [nnf transformation 52]
156. ! [X0] : ((select_3(X0) | ~from_4(X0)) & (from_4(X0) | ~select_3(X0))) [nnf transformation 54]
157. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 55]
160. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 142]
161. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
162. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 160,161]
163. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 143]
164. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 143]
165. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 147]
166. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 147]
167. ~main_query_1(X0) | select_1(X0) [cnf transformation 148]
168. ~select_1(X0) | main_query_1(X0) [cnf transformation 148]
169. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 149]
170. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 149]
171. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 153]
172. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 153]
173. ~subquery_from_5(X0) | select_6(X0) [cnf transformation 154]
174. ~select_6(X0) | subquery_from_5(X0) [cnf transformation 154]
175. ~from_4(X0) | subquery_from_5(X0) [cnf transformation 155]
176. ~subquery_from_5(X0) | from_4(X0) [cnf transformation 155]
177. ~select_3(X0) | from_4(X0) [cnf transformation 156]
178. ~from_4(X0) | select_3(X0) [cnf transformation 156]
179. ~main_query_2(X0) | select_3(X0) [cnf transformation 157]
180. ~select_3(X0) | main_query_2(X0) [cnf transformation 157]
213. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 162]
214. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 162]
216. 0 <=> main_query_1(sK8) [avatar definition]
217. main_query_1(sK8) <- {0} [avatar component clause 216]
219. 1 <=> ~main_query_1(sK8) [avatar definition]
220. ~main_query_1(sK8) <- {1} [avatar component clause 219]
222. 2 <=> main_query_2(sK8) [avatar definition]
223. main_query_2(sK8) <- {2} [avatar component clause 222]
225. 3 <=> ~main_query_2(sK8) [avatar definition]
226. ~main_query_2(sK8) <- {3} [avatar component clause 225]
227. ~1 | ~3 [avatar split clause 214,225,219]
228. 0 | 2 [avatar split clause 213,222,216]
229. select_3(sK8) <- {2} [resolution 179,223]
231. from_4(sK8) <- {2} [resolution 229,177]
234. subquery_from_5(sK8) <- {2} [resolution 231,175]
237. select_6(sK8) <- {2} [resolution 234,173]
266. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 165,163]
268. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 171,169]
270. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 266,170]
274. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 268,164]
279. ~select_1(X5) | select_6(X5) [resolution 270,172]
285. ~select_6(X1) | select_1(X1) [resolution 274,166]
288. select_1(sK8) <- {2} [resolution 285,237]
292. main_query_1(sK8) <- {2} [resolution 288,168]
293. $false <- {1, 2} [subsumption resolution 292,220]
294. 1 | ~2 [avatar contradiction clause 293]
295. select_1(sK8) <- {0} [resolution 217,167]
298. select_6(sK8) <- {0} [resolution 295,279]
303. subquery_from_5(sK8) <- {0} [resolution 298,174]
306. from_4(sK8) <- {0} [resolution 303,176]
310. select_3(sK8) <- {0} [resolution 306,178]
312. main_query_2(sK8) <- {0} [resolution 310,180]
314. $false <- {0, 3} [subsumption resolution 312,226]
315. ~0 | 3 [avatar contradiction clause 314]
316. $false [avatar sat refutation 227,228,294,315]
% SZS output end Proof for query00_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.030 s
% ------------------------------
% ------------------------------
