% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query07_tptp4X
% SZS output start Proof for query00_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
6. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X6] : (main_query_2(X6) <=> select_3(X6)) [input]
8. ! [X7] : eq(X7,X7) [input]
14. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
15. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
41. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
42. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
45. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
46. ! [X0] : eq(X0,X0) [rectify 8]
52. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 14]
53. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 15]
79. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 42]
131. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 79]
132. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
133. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 133]
135. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 134,135]
137. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 44]
138. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
139. ! [X0] : ((where_5(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [nnf transformation 5]
140. ! [X0] : ((where_5(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [flattening 139]
141. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 6]
142. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 141]
143. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
144. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 142,143]
145. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 45]
148. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 52]
149. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 148]
150. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 53]
151. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 131]
152. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
153. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 151,152]
154. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 132]
155. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 132]
156. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 136]
157. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 136]
158. ~main_query_1(X0) | select_1(X0) [cnf transformation 137]
159. ~select_1(X0) | main_query_1(X0) [cnf transformation 137]
160. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 138]
161. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 138]
164. gt(X0,X0) | where_5(X0) | lt(X0,X0) [cnf transformation 140]
165. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 144]
167. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 144]
168. ~main_query_2(X0) | select_3(X0) [cnf transformation 145]
169. ~select_3(X0) | main_query_2(X0) [cnf transformation 145]
170. eq(X0,X0) [cnf transformation 46]
179. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 149]
181. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 150]
206. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 153]
207. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 153]
232. 0 <=> main_query_1(sK8) [avatar definition]
233. main_query_1(sK8) <- {0} [avatar component clause 232]
235. 1 <=> ~main_query_1(sK8) [avatar definition]
236. ~main_query_1(sK8) <- {1} [avatar component clause 235]
238. 2 <=> main_query_2(sK8) [avatar definition]
239. main_query_2(sK8) <- {2} [avatar component clause 238]
241. 3 <=> ~main_query_2(sK8) [avatar definition]
242. ~main_query_2(sK8) <- {3} [avatar component clause 241]
243. ~1 | ~3 [avatar split clause 207,241,235]
244. 0 | 2 [avatar split clause 206,238,232]
245. select_3(sK8) <- {2} [resolution 168,239]
251. ~lt(X0,X0) [resolution 179,170]
253. where_5(X1) | lt(X1,X1) | lt(X1,X1) [resolution 164,181]
254. where_5(X1) | lt(X1,X1) [duplicate literal removal 253]
255. where_5(X1) [subsumption resolution 254,251]
260. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 156,154]
262. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 165,160]
264. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 260,161]
267. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 262,155]
269. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 264,167]
270. ~select_1(X1) | select_3(X1) [subsumption resolution 269,255]
272. ~select_3(X1) | select_1(X1) [resolution 267,157]
273. select_1(sK8) <- {2} [resolution 272,245]
275. main_query_1(sK8) <- {2} [resolution 273,159]
276. $false <- {1, 2} [subsumption resolution 275,236]
277. 1 | ~2 [avatar contradiction clause 276]
278. select_1(sK8) <- {0} [resolution 233,158]
279. select_3(sK8) <- {0} [resolution 278,270]
282. main_query_2(sK8) <- {0} [resolution 279,169]
284. $false <- {0, 3} [subsumption resolution 282,242]
285. ~0 | 3 [avatar contradiction clause 284]
286. $false [avatar sat refutation 243,244,277,285]
% SZS output end Proof for query00_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.010 s
% ------------------------------
% ------------------------------
