% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_16_query_17_tptp4X
% SZS output start Proof for query_16_query_17_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
3. ! [X0,X1] : (where_10(X0,X1) <=> eq(X0,X1)) [input]
4. ! [X0,X1] : (select_8(X0,X1) <=> ? [X2,X3,X4] : (where_10(X0,X1) & from_9(X0,X1,X2,X3,X4))) [input]
5. ! [X5,X6] : (subquery_from_7(X5,X6) <=> select_8(X5,X6)) [input]
6. ! [X7,X8] : (from_6(X7,X8) <=> subquery_from_7(X7,X8)) [input]
7. ! [X7] : (select_5(X7) <=> ? [X8] : ($true & from_6(X7,X8))) [input]
8. ! [X5] : (where_3_subquery_4(X5) <=> select_5(X5)) [input]
9. ! [X0] : (where_3(X0) <=> where_3_subquery_4(X0)) [input]
10. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
11. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
12. ! [X0,X1,X2,X3,X4] : (from_15(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
13. ! [X0,X1] : (where_16(X0,X1) <=> eq(X0,X1)) [input]
14. ! [X0,X1] : (select_14(X0,X1) <=> ? [X2,X3,X4] : (where_16(X0,X1) & from_15(X0,X1,X2,X3,X4))) [input]
15. ! [X5,X6] : (subquery_from_13(X5,X6) <=> select_14(X5,X6)) [input]
16. ! [X7,X8] : (from_12(X7,X8) <=> subquery_from_13(X7,X8)) [input]
17. ! [X7] : (select_11(X7) <=> ? [X8] : ($true & from_12(X7,X8))) [input]
18. ! [X5] : (main_query_2(X5) <=> select_11(X5)) [input]
19. ! [X9] : eq(X9,X9) [input]
20. ! [X9,X10] : (eq(X9,X10) => eq(X10,X9)) [input]
31. ! [X9,X7,X8] : ((from_12(X7,X8) & eq(X8,X9)) => from_12(X7,X9)) [input]
33. ! [X9,X7,X8] : ((from_6(X7,X8) & eq(X8,X9)) => from_6(X7,X9)) [input]
71. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
72. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 71]
73. ! [X0,X1] : (subquery_from_7(X0,X1) <=> select_8(X0,X1)) [rectify 5]
74. ! [X0,X1] : (from_6(X0,X1) <=> subquery_from_7(X0,X1)) [rectify 6]
75. ! [X0] : (select_5(X0) <=> ? [X1] : ($true & from_6(X0,X1))) [rectify 7]
76. ! [X0] : (select_5(X0) <=> ? [X1] : from_6(X0,X1)) [true and false elimination 75]
77. ! [X0] : (where_3_subquery_4(X0) <=> select_5(X0)) [rectify 8]
78. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 11]
79. ! [X0,X1] : (subquery_from_13(X0,X1) <=> select_14(X0,X1)) [rectify 15]
80. ! [X0,X1] : (from_12(X0,X1) <=> subquery_from_13(X0,X1)) [rectify 16]
81. ! [X0] : (select_11(X0) <=> ? [X1] : ($true & from_12(X0,X1))) [rectify 17]
82. ! [X0] : (select_11(X0) <=> ? [X1] : from_12(X0,X1)) [true and false elimination 81]
83. ! [X0] : (main_query_2(X0) <=> select_11(X0)) [rectify 18]
84. ! [X0] : eq(X0,X0) [rectify 19]
85. ! [X0,X1] : (eq(X0,X1) => eq(X1,X0)) [rectify 20]
96. ! [X0,X1,X2] : ((from_12(X1,X2) & eq(X2,X0)) => from_12(X1,X0)) [rectify 31]
98. ! [X0,X1,X2] : ((from_6(X1,X2) & eq(X2,X0)) => from_6(X1,X0)) [rectify 33]
136. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 72]
137. ! [X0,X1] : (eq(X1,X0) | ~eq(X0,X1)) [ennf transformation 85]
144. ! [X0,X1,X2] : (from_12(X1,X0) | (~from_12(X1,X2) | ~eq(X2,X0))) [ennf transformation 96]
145. ! [X0,X1,X2] : (from_12(X1,X0) | ~from_12(X1,X2) | ~eq(X2,X0)) [flattening 144]
148. ! [X0,X1,X2] : (from_6(X1,X0) | (~from_6(X1,X2) | ~eq(X2,X0))) [ennf transformation 98]
149. ! [X0,X1,X2] : (from_6(X1,X0) | ~from_6(X1,X2) | ~eq(X2,X0)) [flattening 148]
224. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 136]
225. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
226. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 2]
227. ! [X0,X1] : ((where_10(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_10(X0,X1))) [nnf transformation 3]
228. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4] : (~where_10(X0,X1) | ~from_9(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_10(X0,X1) & from_9(X0,X1,X2,X3,X4)) | ~select_8(X0,X1))) [nnf transformation 4]
229. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4] : (~where_10(X0,X1) | ~from_9(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_10(X0,X1) & from_9(X0,X1,X5,X6,X7)) | ~select_8(X0,X1))) [rectify 228]
230. ! [X1,X0] : (? [X5,X6,X7] : (where_10(X0,X1) & from_9(X0,X1,X5,X6,X7)) => (where_10(X0,X1) & from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
231. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4] : (~where_10(X0,X1) | ~from_9(X0,X1,X2,X3,X4))) & ((where_10(X0,X1) & from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_8(X0,X1))) [skolemisation 229,230]
232. ! [X0,X1] : ((subquery_from_7(X0,X1) | ~select_8(X0,X1)) & (select_8(X0,X1) | ~subquery_from_7(X0,X1))) [nnf transformation 73]
233. ! [X0,X1] : ((from_6(X0,X1) | ~subquery_from_7(X0,X1)) & (subquery_from_7(X0,X1) | ~from_6(X0,X1))) [nnf transformation 74]
234. ! [X0] : ((select_5(X0) | ! [X1] : ~from_6(X0,X1)) & (? [X1] : from_6(X0,X1) | ~select_5(X0))) [nnf transformation 76]
235. ! [X0] : ((select_5(X0) | ! [X1] : ~from_6(X0,X1)) & (? [X2] : from_6(X0,X2) | ~select_5(X0))) [rectify 234]
236. ! [X0] : (? [X2] : from_6(X0,X2) => from_6(X0,sK3(X0))) [choice axiom]
237. ! [X0] : ((select_5(X0) | ! [X1] : ~from_6(X0,X1)) & (from_6(X0,sK3(X0)) | ~select_5(X0))) [skolemisation 235,236]
238. ! [X0] : ((where_3_subquery_4(X0) | ~select_5(X0)) & (select_5(X0) | ~where_3_subquery_4(X0))) [nnf transformation 77]
239. ! [X0] : ((where_3(X0) | ~where_3_subquery_4(X0)) & (where_3_subquery_4(X0) | ~where_3(X0))) [nnf transformation 9]
240. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 10]
241. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 240]
242. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
243. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_1(X0))) [skolemisation 241,242]
244. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 78]
245. ! [X0,X1,X2,X3,X4] : ((from_15(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_15(X0,X1,X2,X3,X4))) [nnf transformation 12]
246. ! [X0,X1] : ((where_16(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_16(X0,X1))) [nnf transformation 13]
247. ! [X0,X1] : ((select_14(X0,X1) | ! [X2,X3,X4] : (~where_16(X0,X1) | ~from_15(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_16(X0,X1) & from_15(X0,X1,X2,X3,X4)) | ~select_14(X0,X1))) [nnf transformation 14]
248. ! [X0,X1] : ((select_14(X0,X1) | ! [X2,X3,X4] : (~where_16(X0,X1) | ~from_15(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_16(X0,X1) & from_15(X0,X1,X5,X6,X7)) | ~select_14(X0,X1))) [rectify 247]
249. ! [X1,X0] : (? [X5,X6,X7] : (where_16(X0,X1) & from_15(X0,X1,X5,X6,X7)) => (where_16(X0,X1) & from_15(X0,X1,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)))) [choice axiom]
250. ! [X0,X1] : ((select_14(X0,X1) | ! [X2,X3,X4] : (~where_16(X0,X1) | ~from_15(X0,X1,X2,X3,X4))) & ((where_16(X0,X1) & from_15(X0,X1,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1))) | ~select_14(X0,X1))) [skolemisation 248,249]
251. ! [X0,X1] : ((subquery_from_13(X0,X1) | ~select_14(X0,X1)) & (select_14(X0,X1) | ~subquery_from_13(X0,X1))) [nnf transformation 79]
252. ! [X0,X1] : ((from_12(X0,X1) | ~subquery_from_13(X0,X1)) & (subquery_from_13(X0,X1) | ~from_12(X0,X1))) [nnf transformation 80]
253. ! [X0] : ((select_11(X0) | ! [X1] : ~from_12(X0,X1)) & (? [X1] : from_12(X0,X1) | ~select_11(X0))) [nnf transformation 82]
254. ! [X0] : ((select_11(X0) | ! [X1] : ~from_12(X0,X1)) & (? [X2] : from_12(X0,X2) | ~select_11(X0))) [rectify 253]
255. ! [X0] : (? [X2] : from_12(X0,X2) => from_12(X0,sK11(X0))) [choice axiom]
256. ! [X0] : ((select_11(X0) | ! [X1] : ~from_12(X0,X1)) & (from_12(X0,sK11(X0)) | ~select_11(X0))) [skolemisation 254,255]
257. ! [X0] : ((main_query_2(X0) | ~select_11(X0)) & (select_11(X0) | ~main_query_2(X0))) [nnf transformation 83]
260. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 224]
261. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK12) | ~main_query_1(sK12)) & (main_query_2(sK12) | main_query_1(sK12))) [choice axiom]
262. (~main_query_2(sK12) | ~main_query_1(sK12)) & (main_query_2(sK12) | main_query_1(sK12)) [skolemisation 260,261]
264. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 225]
265. ~from_9(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 226]
266. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 226]
267. ~where_10(X0,X1) | eq(X0,X1) [cnf transformation 227]
268. ~eq(X0,X1) | where_10(X0,X1) [cnf transformation 227]
269. from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_8(X0,X1) [cnf transformation 231]
270. ~select_8(X0,X1) | where_10(X0,X1) [cnf transformation 231]
271. ~from_9(X0,X1,X2,X3,X4) | ~where_10(X0,X1) | select_8(X0,X1) [cnf transformation 231]
272. ~subquery_from_7(X0,X1) | select_8(X0,X1) [cnf transformation 232]
273. ~select_8(X0,X1) | subquery_from_7(X0,X1) [cnf transformation 232]
274. ~from_6(X0,X1) | subquery_from_7(X0,X1) [cnf transformation 233]
275. ~subquery_from_7(X0,X1) | from_6(X0,X1) [cnf transformation 233]
276. from_6(X0,sK3(X0)) | ~select_5(X0) [cnf transformation 237]
277. ~from_6(X0,X1) | select_5(X0) [cnf transformation 237]
278. ~where_3_subquery_4(X0) | select_5(X0) [cnf transformation 238]
279. ~select_5(X0) | where_3_subquery_4(X0) [cnf transformation 238]
280. ~where_3(X0) | where_3_subquery_4(X0) [cnf transformation 239]
281. ~where_3_subquery_4(X0) | where_3(X0) [cnf transformation 239]
283. ~select_1(X0) | where_3(X0) [cnf transformation 243]
284. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 243]
285. ~main_query_1(X0) | select_1(X0) [cnf transformation 244]
286. ~select_1(X0) | main_query_1(X0) [cnf transformation 244]
287. ~from_15(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 245]
288. ~tab1(X0,X1,X2,X3,X4) | from_15(X0,X1,X2,X3,X4) [cnf transformation 245]
289. ~where_16(X0,X1) | eq(X0,X1) [cnf transformation 246]
290. ~eq(X0,X1) | where_16(X0,X1) [cnf transformation 246]
291. from_15(X0,X1,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) | ~select_14(X0,X1) [cnf transformation 250]
292. ~select_14(X0,X1) | where_16(X0,X1) [cnf transformation 250]
293. ~from_15(X0,X1,X2,X3,X4) | ~where_16(X0,X1) | select_14(X0,X1) [cnf transformation 250]
294. ~subquery_from_13(X0,X1) | select_14(X0,X1) [cnf transformation 251]
295. ~select_14(X0,X1) | subquery_from_13(X0,X1) [cnf transformation 251]
296. ~from_12(X0,X1) | subquery_from_13(X0,X1) [cnf transformation 252]
297. ~subquery_from_13(X0,X1) | from_12(X0,X1) [cnf transformation 252]
298. from_12(X0,sK11(X0)) | ~select_11(X0) [cnf transformation 256]
299. ~from_12(X0,X1) | select_11(X0) [cnf transformation 256]
300. ~main_query_2(X0) | select_11(X0) [cnf transformation 257]
301. ~select_11(X0) | main_query_2(X0) [cnf transformation 257]
302. eq(X0,X0) [cnf transformation 84]
303. ~eq(X0,X1) | eq(X1,X0) [cnf transformation 137]
312. ~from_12(X1,X2) | from_12(X1,X0) | ~eq(X2,X0) [cnf transformation 145]
314. ~from_6(X1,X2) | from_6(X1,X0) | ~eq(X2,X0) [cnf transformation 149]
352. main_query_2(sK12) | main_query_1(sK12) [cnf transformation 262]
353. ~main_query_2(sK12) | ~main_query_1(sK12) [cnf transformation 262]
355. 0 <=> main_query_1(sK12) [avatar definition]
356. main_query_1(sK12) <- {0} [avatar component clause 355]
358. 1 <=> ~main_query_1(sK12) [avatar definition]
361. 2 <=> main_query_2(sK12) [avatar definition]
362. main_query_2(sK12) <- {2} [avatar component clause 361]
364. 3 <=> ~main_query_2(sK12) [avatar definition]
365. ~main_query_2(sK12) <- {3} [avatar component clause 364]
366. ~1 | ~3 [avatar split clause 353,364,358]
367. 0 | 2 [avatar split clause 352,361,355]
368. select_11(sK12) <- {2} [resolution 300,362]
370. where_10(X0,X0) [resolution 268,302]
372. subquery_from_7(X0,sK3(X0)) | ~select_5(X0) [resolution 276,274]
374. where_16(X0,X0) [resolution 290,302]
376. subquery_from_13(X0,sK11(X0)) | ~select_11(X0) [resolution 298,296]
383. select_8(X1,sK3(X1)) | ~select_5(X1) [resolution 372,272]
386. select_14(X1,sK11(X1)) | ~select_11(X1) [resolution 376,294]
391. where_10(X1,sK3(X1)) | ~select_5(X1) [resolution 383,270]
396. where_16(X1,sK11(X1)) | ~select_11(X1) [resolution 386,292]
398. eq(X0,sK3(X0)) | ~select_5(X0) [resolution 391,267]
399. ~eq(sK11(X0),X1) | from_12(X0,X1) | ~select_11(X0) [resolution 312,298]
400. eq(X0,sK11(X0)) | ~select_11(X0) [resolution 396,289]
405. eq(sK3(X4),X4) | ~select_5(X4) [resolution 398,303]
408. ~eq(sK3(X0),X1) | from_6(X0,X1) | ~select_5(X0) [resolution 314,276]
412. eq(sK11(X4),X4) | ~select_11(X4) [resolution 400,303]
499. from_12(X0,X0) | ~select_11(X0) | ~select_11(X0) [resolution 399,412]
503. from_12(X0,X0) | ~select_11(X0) [duplicate literal removal 499]
506. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_8(X2,X3) [resolution 269,265]
509. subquery_from_13(X4,X4) | ~select_11(X4) [resolution 503,296]
512. tab1(X2,X3,sK8(X2,X3),sK9(X2,X3),sK10(X2,X3)) | ~select_14(X2,X3) [resolution 291,287]
516. select_14(X5,X5) | ~select_11(X5) [resolution 509,294]
666. from_6(X0,X0) | ~select_5(X0) | ~select_5(X0) [resolution 408,405]
671. from_6(X0,X0) | ~select_5(X0) [duplicate literal removal 666]
684. subquery_from_7(X4,X4) | ~select_5(X4) [resolution 671,274]
688. from_15(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_8(X4,X5) [resolution 506,288]
696. from_2(X0,X1,sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) | ~select_14(X0,X1) [resolution 512,264]
697. from_9(X2,X3,sK8(X2,X3),sK9(X2,X3),sK10(X2,X3)) | ~select_14(X2,X3) [resolution 512,266]
707. select_8(X5,X5) | ~select_5(X5) [resolution 684,272]
1112. ~where_16(X2,X3) | ~select_8(X2,X3) | select_14(X2,X3) [resolution 688,293]
1121. ~select_14(X2,X3) | ~where_3(X2) | select_1(X2) [resolution 696,284]
1129. ~select_14(X2,X3) | ~where_10(X2,X3) | select_8(X2,X3) [resolution 697,271]
1166. ~select_11(X0) | select_1(X0) | ~where_3(X0) [resolution 1121,386]
1179. select_1(sK12) | ~where_3(sK12) <- {2} [resolution 1166,368]
1184. 4 <=> where_3(sK12) [avatar definition]
1185. where_3(sK12) <- {4} [avatar component clause 1184]
1187. 5 <=> ~where_3(sK12) [avatar definition]
1188. ~where_3(sK12) <- {5} [avatar component clause 1187]
1193. 6 <=> select_1(sK12) [avatar definition]
1194. select_1(sK12) <- {6} [avatar component clause 1193]
1195. ~5 | 6 | ~2 [avatar split clause 1179,361,1193,1187]
1917. ~select_8(X0,X0) | select_14(X0,X0) [resolution 1112,374]
1933. select_14(X1,X1) | ~select_5(X1) [resolution 1917,707]
1939. ~where_10(X5,X5) | select_8(X5,X5) | ~select_11(X5) [resolution 1129,516]
1948. select_8(X5,X5) | ~select_11(X5) [subsumption resolution 1939,370]
1953. subquery_from_13(X1,X1) | ~select_5(X1) [resolution 1933,295]
1961. subquery_from_7(X1,X1) | ~select_11(X1) [resolution 1948,273]
1970. from_12(X4,X4) | ~select_5(X4) [resolution 1953,297]
1986. from_6(X4,X4) | ~select_11(X4) [resolution 1961,275]
2003. ~select_5(X5) | select_11(X5) [resolution 1970,299]
2035. ~select_11(X5) | select_5(X5) [resolution 1986,277]
2056. select_5(sK12) <- {2} [resolution 2035,368]
2083. where_3_subquery_4(sK12) <- {2} [resolution 2056,279]
2091. where_3(sK12) <- {2} [resolution 2083,281]
2093. $false <- {2, 5} [subsumption resolution 2091,1188]
2094. ~2 | 5 [avatar contradiction clause 2093]
2095. select_1(sK12) <- {0} [resolution 356,285]
2096. 6 | ~0 [avatar split clause 2095,355,1193]
2109. where_3_subquery_4(sK12) <- {4} [resolution 1185,280]
2117. main_query_1(sK12) <- {6} [resolution 1194,286]
2118. where_3(sK12) <- {6} [resolution 1194,283]
2121. 0 | ~6 [avatar split clause 2117,1193,355]
2131. select_5(sK12) <- {4} [resolution 2109,278]
2139. select_11(sK12) <- {4} [resolution 2131,2003]
2150. main_query_2(sK12) <- {4} [resolution 2139,301]
2151. $false <- {3, 4} [subsumption resolution 2150,365]
2152. 3 | ~4 [avatar contradiction clause 2151]
2153. 4 | ~6 [avatar split clause 2118,1193,1184]
2154. $false [avatar sat refutation 366,367,1195,2094,2096,2121,2152,2153]
% SZS output end Proof for query_16_query_17_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5756
% Time elapsed: 0.156 s
% ------------------------------
% ------------------------------
