% Refutation found. Thanks to Tanya!
% SZS status Theorem for query01_query05_tptp4X
% SZS output start Proof for query01_query05_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X5,X1,X2,X3,X4] : (from_4(X5,X1,X2,X3,X4) <=> tab1(X5,X1,X2,X3,X4)) [input]
5. ! [X5] : (where_5(X5) <=> eq(X5,X5)) [input]
6. ! [X5] : (select_3(X5) <=> ? [X1,X2,X3,X4] : (where_5(X5) & from_4(X5,X1,X2,X3,X4))) [input]
7. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
8. ! [X6] : eq(X6,X6) [input]
41. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
42. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
45. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [rectify 5]
46. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
47. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
48. ! [X0] : eq(X0,X0) [rectify 8]
132. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 42]
133. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 134]
136. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 135,136]
138. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
139. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 44]
140. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 45]
141. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 46]
142. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 141]
143. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
144. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 142,143]
145. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 47]
148. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 132]
149. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
150. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 148,149]
151. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 133]
152. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 133]
153. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 137]
154. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 137]
155. ~main_query_1(X0) | select_1(X0) [cnf transformation 138]
156. ~select_1(X0) | main_query_1(X0) [cnf transformation 138]
157. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 139]
158. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 139]
160. ~eq(X0,X0) | where_5(X0) [cnf transformation 140]
161. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 144]
163. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 144]
164. ~main_query_2(X0) | select_3(X0) [cnf transformation 145]
165. ~select_3(X0) | main_query_2(X0) [cnf transformation 145]
166. eq(X0,X0) [cnf transformation 48]
197. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 150]
198. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 150]
223. 0 <=> main_query_1(sK8) [avatar definition]
224. main_query_1(sK8) <- {0} [avatar component clause 223]
226. 1 <=> ~main_query_1(sK8) [avatar definition]
227. ~main_query_1(sK8) <- {1} [avatar component clause 226]
229. 2 <=> main_query_2(sK8) [avatar definition]
230. main_query_2(sK8) <- {2} [avatar component clause 229]
232. 3 <=> ~main_query_2(sK8) [avatar definition]
233. ~main_query_2(sK8) <- {3} [avatar component clause 232]
234. ~1 | ~3 [avatar split clause 198,232,226]
235. 0 | 2 [avatar split clause 197,229,223]
236. select_3(sK8) <- {2} [resolution 164,230]
240. where_5(X1) [resolution 160,166]
251. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 153,151]
253. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 161,157]
255. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 251,158]
258. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 253,152]
260. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 255,163]
261. ~select_1(X1) | select_3(X1) [subsumption resolution 260,240]
263. ~select_3(X1) | select_1(X1) [resolution 258,154]
264. select_1(sK8) <- {2} [resolution 263,236]
266. main_query_1(sK8) <- {2} [resolution 264,156]
267. $false <- {1, 2} [subsumption resolution 266,227]
268. 1 | ~2 [avatar contradiction clause 267]
269. select_1(sK8) <- {0} [resolution 224,155]
270. select_3(sK8) <- {0} [resolution 269,261]
273. main_query_2(sK8) <- {0} [resolution 270,165]
275. $false <- {0, 3} [subsumption resolution 273,233]
276. ~0 | 3 [avatar contradiction clause 275]
277. $false [avatar sat refutation 234,235,268,276]
% SZS output end Proof for query01_query05_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
