% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_06_query_07_tptp4X
% SZS output start Proof for query_06_query_07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
7. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
12. ! [X7] : lte(X7,X7) [input]
15. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
16. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
18. ! [X7,X8] : (gte(X7,X8) <=> lte(X8,X7)) [input]
47. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
48. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 47]
49. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
50. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
51. ! [X0] : eq(X0,X0) [rectify 9]
54. ! [X0] : lte(X0,X0) [rectify 12]
57. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 15]
58. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 16]
60. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 18]
89. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 48]
151. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 89]
152. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
153. ! [X0] : ((where_3(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
154. ! [X0] : ((where_3(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_3(X0))) [flattening 153]
155. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
156. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 155]
157. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
158. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 156,157]
159. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 49]
160. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
161. ! [X0] : ((where_6(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [nnf transformation 6]
162. ! [X0] : ((where_6(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [flattening 161]
163. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 7]
164. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 163]
165. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
166. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 164,165]
167. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 50]
170. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 57]
171. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 170]
172. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 58]
173. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 60]
174. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 151]
175. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
176. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 174,175]
177. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 152]
178. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 152]
181. ~gte(X0,X0) | where_3(X0) | ~lte(X0,X0) [cnf transformation 154]
182. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 158]
184. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0,X1) [cnf transformation 158]
185. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 159]
186. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 159]
187. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 160]
188. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 160]
191. gt(X0,X0) | where_6(X0) | lt(X0,X0) [cnf transformation 162]
192. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 166]
194. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0,X1) [cnf transformation 166]
195. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 167]
196. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 167]
197. eq(X0,X0) [cnf transformation 51]
200. lte(X0,X0) [cnf transformation 54]
206. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 171]
208. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 172]
211. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 173]
240. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 176]
241. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 176]
243. 0 <=> main_query_1(sK6,sK7) [avatar definition]
244. main_query_1(sK6,sK7) <- {0} [avatar component clause 243]
246. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
247. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 246]
249. 2 <=> main_query_2(sK6,sK7) [avatar definition]
250. main_query_2(sK6,sK7) <- {2} [avatar component clause 249]
252. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
253. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 252]
254. ~1 | ~3 [avatar split clause 241,252,246]
255. 0 | 2 [avatar split clause 240,249,243]
256. select_4(sK6,sK7) <- {2} [resolution 195,250]
262. ~lt(X0,X0) [resolution 206,197]
265. gte(X1,X1) [resolution 211,200]
267. where_3(X0) | ~lte(X0,X0) [resolution 181,265]
269. where_3(X0) [subsumption resolution 267,200]
271. where_6(X1) | lt(X1,X1) | lt(X1,X1) [resolution 191,208]
272. where_6(X1) | lt(X1,X1) [duplicate literal removal 271]
273. where_6(X1) [subsumption resolution 272,262]
289. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 182,177]
291. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 192,187]
304. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 289,188]
310. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 291,178]
325. ~select_1(X2,X3) | ~where_6(X2) | select_4(X2,X3) [resolution 304,194]
331. ~select_1(X2,X3) | select_4(X2,X3) [subsumption resolution 325,273]
347. ~select_4(X2,X3) | ~where_3(X2) | select_1(X2,X3) [resolution 310,184]
353. ~select_4(X2,X3) | select_1(X2,X3) [subsumption resolution 347,269]
361. select_1(sK6,sK7) <- {2} [resolution 353,256]
370. main_query_1(sK6,sK7) <- {2} [resolution 361,186]
372. $false <- {1, 2} [subsumption resolution 370,247]
373. 1 | ~2 [avatar contradiction clause 372]
374. select_1(sK6,sK7) <- {0} [resolution 244,185]
375. select_4(sK6,sK7) <- {0} [resolution 374,331]
379. main_query_2(sK6,sK7) <- {0} [resolution 375,196]
381. $false <- {0, 3} [subsumption resolution 379,253]
382. ~0 | 3 [avatar contradiction clause 381]
383. $false [avatar sat refutation 254,255,373,382]
% SZS output end Proof for query_06_query_07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
