% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_11_query_28_tptp4X
% SZS output start Proof for query_11_query_28_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X7,X2,X3,X4] : (from_5(X0,X7,X2,X3,X4) <=> tab1(X0,X7,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0,X7] : (where_10(X0,X7) <=> eq(X0,X7)) [input]
8. ! [X0,X7] : (select_8(X0,X7) <=> ? [X1,X2,X3,X4] : (where_10(X0,X7) & from_9(X0,X1,X2,X3,X4))) [input]
9. ! [X5,X7] : (where_6_subquery_7(X5,X7) <=> select_8(X5,X7)) [input]
10. ! [X0,X7] : (where_6(X0,X7) <=> where_6_subquery_7(X0,X7)) [input]
11. ! [X0,X7] : (select_4(X0,X7) <=> ? [X2,X3,X4] : (where_6(X0,X7) & from_5(X0,X7,X2,X3,X4))) [input]
12. ! [X5,X7] : (main_query_2(X5,X7) <=> select_4(X5,X7)) [input]
14. ! [X8,X9] : (eq(X8,X9) => eq(X9,X8)) [input]
44. ! [X8,X5,X6] : ((select_1(X5,X6) & eq(X6,X8)) => select_1(X5,X8)) [input]
66. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
67. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 66]
68. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
69. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
70. ! [X0,X1] : (where_10(X0,X1) <=> eq(X0,X1)) [rectify 7]
71. ! [X0,X1] : (select_8(X0,X1) <=> ? [X2,X3,X4,X5] : (where_10(X0,X1) & from_9(X0,X2,X3,X4,X5))) [rectify 8]
72. ! [X0,X1] : (where_6_subquery_7(X0,X1) <=> select_8(X0,X1)) [rectify 9]
73. ! [X0,X1] : (where_6(X0,X1) <=> where_6_subquery_7(X0,X1)) [rectify 10]
74. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 11]
75. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 12]
77. ! [X0,X1] : (eq(X0,X1) => eq(X1,X0)) [rectify 14]
107. ! [X0,X1,X2] : ((select_1(X1,X2) & eq(X2,X0)) => select_1(X1,X0)) [rectify 44]
129. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 67]
130. ! [X0,X1] : (eq(X1,X0) | ~eq(X0,X1)) [ennf transformation 77]
177. ! [X0,X1,X2] : (select_1(X1,X0) | (~select_1(X1,X2) | ~eq(X2,X0))) [ennf transformation 107]
178. ! [X0,X1,X2] : (select_1(X1,X0) | ~select_1(X1,X2) | ~eq(X2,X0)) [flattening 177]
221. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 129]
222. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
223. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
224. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
225. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 224]
226. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
227. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 225,226]
228. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 68]
229. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 69]
230. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
231. ! [X0,X1] : ((where_10(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_10(X0,X1))) [nnf transformation 70]
232. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5] : (~where_10(X0,X1) | ~from_9(X0,X2,X3,X4,X5))) & (? [X2,X3,X4,X5] : (where_10(X0,X1) & from_9(X0,X2,X3,X4,X5)) | ~select_8(X0,X1))) [nnf transformation 71]
233. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5] : (~where_10(X0,X1) | ~from_9(X0,X2,X3,X4,X5))) & (? [X6,X7,X8,X9] : (where_10(X0,X1) & from_9(X0,X6,X7,X8,X9)) | ~select_8(X0,X1))) [rectify 232]
234. ! [X1,X0] : (? [X6,X7,X8,X9] : (where_10(X0,X1) & from_9(X0,X6,X7,X8,X9)) => (where_10(X0,X1) & from_9(X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1),sK6(X0,X1)))) [choice axiom]
235. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5] : (~where_10(X0,X1) | ~from_9(X0,X2,X3,X4,X5))) & ((where_10(X0,X1) & from_9(X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1),sK6(X0,X1))) | ~select_8(X0,X1))) [skolemisation 233,234]
236. ! [X0,X1] : ((where_6_subquery_7(X0,X1) | ~select_8(X0,X1)) & (select_8(X0,X1) | ~where_6_subquery_7(X0,X1))) [nnf transformation 72]
237. ! [X0,X1] : ((where_6(X0,X1) | ~where_6_subquery_7(X0,X1)) & (where_6_subquery_7(X0,X1) | ~where_6(X0,X1))) [nnf transformation 73]
238. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 74]
239. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 238]
240. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
241. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 239,240]
242. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 75]
245. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 221]
246. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
247. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 245,246]
248. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 222]
249. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 222]
250. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 223]
251. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 223]
252. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 227]
253. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 227]
254. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 227]
255. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 228]
256. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 228]
257. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 229]
258. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 229]
260. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 230]
261. ~where_10(X0,X1) | eq(X0,X1) [cnf transformation 231]
262. ~eq(X0,X1) | where_10(X0,X1) [cnf transformation 231]
264. ~select_8(X0,X1) | where_10(X0,X1) [cnf transformation 235]
265. ~from_9(X0,X2,X3,X4,X5) | ~where_10(X0,X1) | select_8(X0,X1) [cnf transformation 235]
266. ~where_6_subquery_7(X0,X1) | select_8(X0,X1) [cnf transformation 236]
267. ~select_8(X0,X1) | where_6_subquery_7(X0,X1) [cnf transformation 236]
268. ~where_6(X0,X1) | where_6_subquery_7(X0,X1) [cnf transformation 237]
269. ~where_6_subquery_7(X0,X1) | where_6(X0,X1) [cnf transformation 237]
270. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 241]
271. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 241]
272. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 241]
273. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 242]
274. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 242]
276. ~eq(X0,X1) | eq(X1,X0) [cnf transformation 130]
304. ~select_1(X1,X2) | select_1(X1,X0) | ~eq(X2,X0) [cnf transformation 178]
326. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 247]
327. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 247]
329. 0 <=> main_query_1(sK10,sK11) [avatar definition]
330. main_query_1(sK10,sK11) <- {0} [avatar component clause 329]
332. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
333. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 332]
335. 2 <=> main_query_2(sK10,sK11) [avatar definition]
336. main_query_2(sK10,sK11) <- {2} [avatar component clause 335]
338. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
339. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 338]
340. ~1 | ~3 [avatar split clause 327,338,332]
341. 0 | 2 [avatar split clause 326,335,329]
346. select_4(sK10,sK11) <- {2} [resolution 273,336]
347. where_6(sK10,sK11) <- {2} [resolution 346,271]
349. where_6_subquery_7(sK10,sK11) <- {2} [resolution 347,268]
352. select_8(sK10,sK11) <- {2} [resolution 349,266]
355. where_10(sK10,sK11) <- {2} [resolution 352,264]
357. eq(sK10,sK11) <- {2} [resolution 355,261]
364. where_3(sK10,sK11) <- {2} [resolution 357,251]
466. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 252,248]
470. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 270,257]
557. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 466,258]
558. from_9(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 466,260]
566. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 470,249]
729. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 557,272]
738. ~where_10(X2,X4) | ~select_1(X2,X3) | select_8(X2,X4) [resolution 558,265]
749. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 566,254]
772. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 749,346]
776. select_1(sK10,sK11) <- {2} [subsumption resolution 772,364]
790. main_query_1(sK10,sK11) <- {2} [resolution 776,256]
792. $false <- {1, 2} [subsumption resolution 790,333]
793. 1 | ~2 [avatar contradiction clause 792]
796. select_1(sK10,sK11) <- {0} [resolution 330,255]
798. ~eq(sK11,X0) | select_1(sK10,X0) <- {0} [resolution 796,304]
801. where_3(sK10,sK11) <- {0} [resolution 796,253]
802. eq(sK10,sK11) <- {0} [resolution 801,250]
805. where_10(sK10,sK11) <- {0} [resolution 802,262]
806. eq(sK11,sK10) <- {0} [resolution 802,276]
815. ~select_1(sK10,X0) | select_8(sK10,sK11) <- {0} [resolution 805,738]
823. 4 <=> select_8(sK10,sK11) [avatar definition]
824. select_8(sK10,sK11) <- {4} [avatar component clause 823]
826. 6 <=> ! [X0] : ~select_1(sK10,X0) [avatar definition]
827. ~select_1(sK10,X0) <- {6} [avatar component clause 826]
828. 4 | 6 | ~0 [avatar split clause 815,329,826,823]
830. select_1(sK10,sK10) <- {0} [resolution 806,798]
841. $false <- {0, 6} [subsumption resolution 830,827]
842. ~0 | ~6 [avatar contradiction clause 841]
930. where_6_subquery_7(sK10,sK11) <- {4} [resolution 824,267]
1025. where_6(sK10,sK11) <- {4} [resolution 930,269]
1065. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {4} [resolution 1025,729]
1069. select_4(sK10,sK11) <- {0, 4} [subsumption resolution 1065,796]
1124. main_query_2(sK10,sK11) <- {0, 4} [resolution 1069,274]
1126. $false <- {0, 3, 4} [subsumption resolution 1124,339]
1127. ~0 | 3 | ~4 [avatar contradiction clause 1126]
1128. $false [avatar sat refutation 340,341,793,828,842,1127]
% SZS output end Proof for query_11_query_28_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
