% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query08_tptp4X
% SZS output start Proof for query00_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
6. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
11. ! [X6] : lte(X6,X6) [input]
18. ! [X6,X7] : (gte(X6,X7) <=> lte(X7,X6)) [input]
42. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
43. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 42]
44. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
45. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
46. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
50. ! [X0] : lte(X0,X0) [rectify 11]
57. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 18]
81. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 43]
133. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 81]
134. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 44]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 135]
137. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 136,137]
139. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 45]
140. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
141. ! [X0] : ((where_5(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [nnf transformation 5]
142. ! [X0] : ((where_5(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [flattening 141]
143. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 6]
144. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 143]
145. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
146. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 144,145]
147. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 46]
150. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 57]
151. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 133]
152. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
153. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 151,152]
154. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 134]
155. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 134]
156. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 138]
157. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 138]
158. ~main_query_1(X0) | select_1(X0) [cnf transformation 139]
159. ~select_1(X0) | main_query_1(X0) [cnf transformation 139]
160. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 140]
161. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 140]
164. ~gte(X0,X0) | where_5(X0) | ~lte(X0,X0) [cnf transformation 142]
165. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 146]
167. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 146]
168. ~main_query_2(X0) | select_3(X0) [cnf transformation 147]
169. ~select_3(X0) | main_query_2(X0) [cnf transformation 147]
173. lte(X0,X0) [cnf transformation 50]
180. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 150]
204. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 153]
205. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 153]
207. 0 <=> main_query_1(sK8) [avatar definition]
208. main_query_1(sK8) <- {0} [avatar component clause 207]
210. 1 <=> ~main_query_1(sK8) [avatar definition]
211. ~main_query_1(sK8) <- {1} [avatar component clause 210]
213. 2 <=> main_query_2(sK8) [avatar definition]
214. main_query_2(sK8) <- {2} [avatar component clause 213]
216. 3 <=> ~main_query_2(sK8) [avatar definition]
217. ~main_query_2(sK8) <- {3} [avatar component clause 216]
218. ~1 | ~3 [avatar split clause 205,216,210]
219. 0 | 2 [avatar split clause 204,213,207]
220. select_3(sK8) <- {2} [resolution 168,214]
229. gte(X1,X1) [resolution 180,173]
233. where_5(X0) | ~lte(X0,X0) [resolution 164,229]
235. where_5(X0) [subsumption resolution 233,173]
301. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 156,154]
303. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 165,160]
305. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 301,161]
308. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 303,155]
310. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 305,167]
311. ~select_1(X1) | select_3(X1) [subsumption resolution 310,235]
313. ~select_3(X1) | select_1(X1) [resolution 308,157]
314. select_1(sK8) <- {2} [resolution 313,220]
318. main_query_1(sK8) <- {2} [resolution 314,159]
319. $false <- {1, 2} [subsumption resolution 318,211]
320. 1 | ~2 [avatar contradiction clause 319]
321. select_1(sK8) <- {0} [resolution 208,158]
324. select_3(sK8) <- {0} [resolution 321,311]
329. main_query_2(sK8) <- {0} [resolution 324,169]
331. $false <- {0, 3} [subsumption resolution 329,217]
332. ~0 | 3 [avatar contradiction clause 331]
333. $false [avatar sat refutation 218,219,320,332]
% SZS output end Proof for query00_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.029 s
% ------------------------------
% ------------------------------
