% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_11_query_15_tptp4X
% SZS output start Proof for query_11_query_15_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (where_6_subquery_7(X5) <=> select_8(X5)) [input]
9. ! [X0,X1] : (where_6(X0,X1) <=> (where_6_subquery_7(X0) & eq(X0,X1))) [input]
10. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
60. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
61. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 60]
62. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
63. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 7]
64. ! [X0] : (where_6_subquery_7(X0) <=> select_8(X0)) [rectify 8]
65. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
114. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 61]
194. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 114]
195. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
196. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
197. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
198. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 197]
199. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
200. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 198,199]
201. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 62]
202. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
203. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
204. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 63]
205. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 204]
206. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
207. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 205,206]
208. ! [X0] : ((where_6_subquery_7(X0) | ~select_8(X0)) & (select_8(X0) | ~where_6_subquery_7(X0))) [nnf transformation 64]
209. ! [X0,X1] : ((where_6(X0,X1) | (~where_6_subquery_7(X0) | ~eq(X0,X1))) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 9]
210. ! [X0,X1] : ((where_6(X0,X1) | ~where_6_subquery_7(X0) | ~eq(X0,X1)) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 209]
211. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 10]
212. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 211]
213. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
214. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 212,213]
215. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 65]
218. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 194]
219. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
220. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 218,219]
221. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 195]
222. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 195]
223. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 196]
224. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 196]
225. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 200]
226. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 200]
227. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 200]
228. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 201]
229. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 201]
230. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 202]
231. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 202]
233. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 203]
235. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 207]
237. ~select_8(X0) | where_6_subquery_7(X0) [cnf transformation 208]
238. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 210]
240. ~eq(X0,X1) | ~where_6_subquery_7(X0) | where_6(X0,X1) [cnf transformation 210]
241. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 214]
242. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 214]
243. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 214]
244. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 215]
245. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 215]
292. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 220]
293. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 220]
295. 0 <=> main_query_1(sK10,sK11) [avatar definition]
296. main_query_1(sK10,sK11) <- {0} [avatar component clause 295]
298. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
299. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 298]
301. 2 <=> main_query_2(sK10,sK11) [avatar definition]
302. main_query_2(sK10,sK11) <- {2} [avatar component clause 301]
304. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
305. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 304]
306. ~1 | ~3 [avatar split clause 293,304,298]
307. 0 | 2 [avatar split clause 292,301,295]
378. select_1(sK10,sK11) <- {0} [resolution 296,228]
380. where_3(sK10,sK11) <- {0} [resolution 378,226]
381. eq(sK10,sK11) <- {0} [resolution 380,223]
384. ~where_6_subquery_7(sK10) | where_6(sK10,sK11) <- {0} [resolution 381,240]
394. 8 <=> where_6(sK10,sK11) [avatar definition]
395. where_6(sK10,sK11) <- {8} [avatar component clause 394]
397. 10 <=> where_6_subquery_7(sK10) [avatar definition]
400. 11 <=> ~where_6_subquery_7(sK10) [avatar definition]
402. 8 | ~11 | ~0 [avatar split clause 384,295,400,394]
446. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 225,221]
450. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 241,230]
586. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 446,231]
587. from_9(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 446,233]
604. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 450,222]
639. 14 <=> select_8(sK10) [avatar definition]
640. select_8(sK10) <- {14} [avatar component clause 639]
1187. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 586,243]
1195. ~select_1(X2,X3) | select_8(X2) [resolution 587,235]
1201. select_8(sK10) <- {0} [resolution 1195,378]
1216. 14 | ~0 [avatar split clause 1201,295,639]
1218. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 604,227]
1236. where_6_subquery_7(sK10) <- {14} [resolution 640,237]
1239. 10 | ~14 [avatar split clause 1236,639,397]
1266. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {8} [resolution 395,1187]
1271. select_4(sK10,sK11) <- {0, 8} [subsumption resolution 1266,378]
1296. main_query_2(sK10,sK11) <- {0, 8} [resolution 1271,245]
1298. $false <- {0, 3, 8} [subsumption resolution 1296,305]
1299. ~0 | 3 | ~8 [avatar contradiction clause 1298]
1315. select_4(sK10,sK11) <- {2} [resolution 302,244]
1319. where_6(sK10,sK11) <- {2} [resolution 1315,242]
1320. 8 | ~2 [avatar split clause 1319,301,394]
1324. eq(sK10,sK11) <- {8} [resolution 395,238]
1327. where_3(sK10,sK11) <- {8} [resolution 1324,224]
1745. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 1218,1315]
1749. select_1(sK10,sK11) <- {2, 8} [subsumption resolution 1745,1327]
1758. main_query_1(sK10,sK11) <- {2, 8} [resolution 1749,229]
1764. $false <- {1, 2, 8} [subsumption resolution 1758,299]
1765. 1 | ~2 | ~8 [avatar contradiction clause 1764]
1766. $false [avatar sat refutation 306,307,402,1216,1239,1299,1320,1765]
% SZS output end Proof for query_11_query_15_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5500
% Time elapsed: 0.035 s
% ------------------------------
% ------------------------------
