% Refutation found. Thanks to Tanya!
% SZS status Theorem for query03_query07_tptp4X
% SZS output start Proof for query03_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_4(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (where_5(X6) <=> (~gt(X6,X6) & ~lt(X6,X6))) [input]
6. ! [X6] : (select_3(X6) <=> ? [X7,X8,X9,X10] : (where_5(X6) & from_4(X6,X7,X8,X9,X10))) [input]
7. ! [X5] : (main_query_2(X5) <=> select_3(X5)) [input]
8. ! [X11] : eq(X11,X11) [input]
14. ! [X11,X12] : (lt(X11,X12) <=> (~eq(X11,X12) & lte(X11,X12))) [input]
15. ! [X11,X12] : (gt(X11,X12) <=> lt(X12,X11)) [input]
46. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
47. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 46]
48. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
50. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
51. ! [X0] : (where_5(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [rectify 5]
52. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
53. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
54. ! [X0] : eq(X0,X0) [rectify 8]
60. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 14]
61. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 15]
92. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 47]
154. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 92]
155. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
156. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 48]
157. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 156]
158. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
159. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 157,158]
160. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 49]
161. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 50]
162. ! [X0] : ((where_5(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [nnf transformation 51]
163. ! [X0] : ((where_5(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [flattening 162]
164. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 52]
165. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 164]
166. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
167. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 165,166]
168. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 53]
171. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 60]
172. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 171]
173. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 61]
174. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 154]
175. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
176. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 174,175]
177. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 155]
178. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 155]
179. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 159]
180. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 159]
181. ~main_query_1(X0) | select_1(X0) [cnf transformation 160]
182. ~select_1(X0) | main_query_1(X0) [cnf transformation 160]
183. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 161]
184. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 161]
187. gt(X0,X0) | where_5(X0) | lt(X0,X0) [cnf transformation 163]
188. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 167]
190. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 167]
191. ~main_query_2(X0) | select_3(X0) [cnf transformation 168]
192. ~select_3(X0) | main_query_2(X0) [cnf transformation 168]
193. eq(X0,X0) [cnf transformation 54]
202. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 172]
204. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 173]
234. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 176]
235. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 176]
265. 0 <=> main_query_1(sK8) [avatar definition]
266. main_query_1(sK8) <- {0} [avatar component clause 265]
268. 1 <=> ~main_query_1(sK8) [avatar definition]
269. ~main_query_1(sK8) <- {1} [avatar component clause 268]
271. 2 <=> main_query_2(sK8) [avatar definition]
272. main_query_2(sK8) <- {2} [avatar component clause 271]
274. 3 <=> ~main_query_2(sK8) [avatar definition]
275. ~main_query_2(sK8) <- {3} [avatar component clause 274]
276. ~1 | ~3 [avatar split clause 235,274,268]
277. 0 | 2 [avatar split clause 234,271,265]
278. select_3(sK8) <- {2} [resolution 191,272]
284. ~lt(X0,X0) [resolution 202,193]
286. where_5(X1) | lt(X1,X1) | lt(X1,X1) [resolution 187,204]
287. where_5(X1) | lt(X1,X1) [duplicate literal removal 286]
288. where_5(X1) [subsumption resolution 287,284]
293. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 179,177]
295. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 188,183]
297. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 293,184]
300. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 295,178]
302. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 297,190]
303. ~select_1(X1) | select_3(X1) [subsumption resolution 302,288]
305. ~select_3(X1) | select_1(X1) [resolution 300,180]
306. select_1(sK8) <- {2} [resolution 305,278]
308. main_query_1(sK8) <- {2} [resolution 306,182]
309. $false <- {1, 2} [subsumption resolution 308,269]
310. 1 | ~2 [avatar contradiction clause 309]
311. select_1(sK8) <- {0} [resolution 266,181]
312. select_3(sK8) <- {0} [resolution 311,303]
315. main_query_2(sK8) <- {0} [resolution 312,192]
317. $false <- {0, 3} [subsumption resolution 315,275]
318. ~0 | 3 [avatar contradiction clause 317]
319. $false [avatar sat refutation 276,277,310,318]
% SZS output end Proof for query03_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
