% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_29_query_30_tptp4X
% SZS output start Proof for query_29_query_30_tptp4X
1. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) <=> (tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4))) [input]
2. ! [X5,X6,X7,X8,X9] : (from_6(X5,X6,X7,X8,X9) <=> tab2(X5,X6,X7,X8,X9)) [input]
3. ! [X5,X0] : (where_7(X5,X0) <=> eq(X5,X0)) [input]
4. ! [X5,X0] : (select_5(X5,X0) <=> ? [X6,X7,X8,X9] : (where_7(X5,X0) & from_6(X5,X6,X7,X8,X9))) [input]
5. ! [X10,X0] : (where_3_subquery_4(X10,X0) <=> select_5(X10,X0)) [input]
6. ! [X0] : (where_3(X0) <=> where_3_subquery_4(X0,X0)) [input]
7. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_3(X0) & from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [input]
8. ! [X11,X12] : (main_query_1(X11,X12) <=> select_1(X11,X12)) [input]
9. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (from_9(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) <=> (tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X0] : (where_14(X5,X0) <=> neq(X5,X0)) [input]
12. ! [X5,X0] : (select_12(X5,X0) <=> ? [X6,X7,X8,X9] : (where_14(X5,X0) & from_13(X5,X6,X7,X8,X9))) [input]
13. ! [X10,X0] : (where_10_subquery_11(X10,X0) <=> select_12(X10,X0)) [input]
14. ! [X5,X6,X7,X8,X9] : (from_17(X5,X6,X7,X8,X9) <=> tab2(X5,X6,X7,X8,X9)) [input]
15. ! [X5,X0] : (where_18(X5,X0) <=> eq(X5,X0)) [input]
16. ! [X5,X0] : (select_16(X5,X0) <=> ? [X6,X7,X8,X9] : (where_18(X5,X0) & from_17(X5,X6,X7,X8,X9))) [input]
17. ! [X10,X0] : (where_10_subquery_15(X10,X0) <=> select_16(X10,X0)) [input]
18. ! [X0] : (where_10(X0) <=> (where_10_subquery_15(X0,X0) & ~where_10_subquery_11(X0,X0))) [input]
19. ! [X0,X1] : (select_8(X0,X1) <=> ? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_10(X0) & from_9(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [input]
20. ! [X11,X12] : (main_query_2(X11,X12) <=> select_8(X11,X12)) [input]
21. ! [X13] : eq(X13,X13) [input]
29. ! [X13,X14] : (neq(X13,X14) <=> ~eq(X13,X14)) [input]
102. ! [X11,X12] : (main_query_1(X11,X12) <=> main_query_2(X11,X12)) [input]
103. ~! [X11,X12] : (main_query_1(X11,X12) <=> main_query_2(X11,X12)) [negated conjecture 102]
104. ! [X0,X1,X2,X3,X4] : (from_6(X0,X1,X2,X3,X4) <=> tab2(X0,X1,X2,X3,X4)) [rectify 2]
105. ! [X0,X1] : (where_7(X0,X1) <=> eq(X0,X1)) [rectify 3]
106. ! [X0,X1] : (select_5(X0,X1) <=> ? [X2,X3,X4,X5] : (where_7(X0,X1) & from_6(X0,X2,X3,X4,X5))) [rectify 4]
107. ! [X0,X1] : (where_3_subquery_4(X0,X1) <=> select_5(X0,X1)) [rectify 5]
108. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 8]
110. ! [X0,X1] : (where_14(X0,X1) <=> neq(X0,X1)) [rectify 11]
111. ! [X0,X1] : (select_12(X0,X1) <=> ? [X2,X3,X4,X5] : (where_14(X0,X1) & from_13(X0,X2,X3,X4,X5))) [rectify 12]
112. ! [X0,X1] : (where_10_subquery_11(X0,X1) <=> select_12(X0,X1)) [rectify 13]
113. ! [X0,X1,X2,X3,X4] : (from_17(X0,X1,X2,X3,X4) <=> tab2(X0,X1,X2,X3,X4)) [rectify 14]
114. ! [X0,X1] : (where_18(X0,X1) <=> eq(X0,X1)) [rectify 15]
115. ! [X0,X1] : (select_16(X0,X1) <=> ? [X2,X3,X4,X5] : (where_18(X0,X1) & from_17(X0,X2,X3,X4,X5))) [rectify 16]
116. ! [X0,X1] : (where_10_subquery_15(X0,X1) <=> select_16(X0,X1)) [rectify 17]
117. ! [X0,X1] : (main_query_2(X0,X1) <=> select_8(X0,X1)) [rectify 20]
118. ! [X0] : eq(X0,X0) [rectify 21]
126. ! [X0,X1] : (neq(X0,X1) <=> ~eq(X0,X1)) [rectify 29]
199. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 103]
347. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 199]
348. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | (~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4))) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [nnf transformation 1]
349. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4)) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [flattening 348]
350. ! [X0,X1,X2,X3,X4] : ((from_6(X0,X1,X2,X3,X4) | ~tab2(X0,X1,X2,X3,X4)) & (tab2(X0,X1,X2,X3,X4) | ~from_6(X0,X1,X2,X3,X4))) [nnf transformation 104]
351. ! [X0,X1] : ((where_7(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_7(X0,X1))) [nnf transformation 105]
352. ! [X0,X1] : ((select_5(X0,X1) | ! [X2,X3,X4,X5] : (~where_7(X0,X1) | ~from_6(X0,X2,X3,X4,X5))) & (? [X2,X3,X4,X5] : (where_7(X0,X1) & from_6(X0,X2,X3,X4,X5)) | ~select_5(X0,X1))) [nnf transformation 106]
353. ! [X0,X1] : ((select_5(X0,X1) | ! [X2,X3,X4,X5] : (~where_7(X0,X1) | ~from_6(X0,X2,X3,X4,X5))) & (? [X6,X7,X8,X9] : (where_7(X0,X1) & from_6(X0,X6,X7,X8,X9)) | ~select_5(X0,X1))) [rectify 352]
354. ! [X1,X0] : (? [X6,X7,X8,X9] : (where_7(X0,X1) & from_6(X0,X6,X7,X8,X9)) => (where_7(X0,X1) & from_6(X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1)))) [choice axiom]
355. ! [X0,X1] : ((select_5(X0,X1) | ! [X2,X3,X4,X5] : (~where_7(X0,X1) | ~from_6(X0,X2,X3,X4,X5))) & ((where_7(X0,X1) & from_6(X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1))) | ~select_5(X0,X1))) [skolemisation 353,354]
356. ! [X0,X1] : ((where_3_subquery_4(X0,X1) | ~select_5(X0,X1)) & (select_5(X0,X1) | ~where_3_subquery_4(X0,X1))) [nnf transformation 107]
357. ! [X0] : ((where_3(X0) | ~where_3_subquery_4(X0,X0)) & (where_3_subquery_4(X0,X0) | ~where_3(X0))) [nnf transformation 6]
358. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & (? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_3(X0) & from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9)) | ~select_1(X0,X1))) [nnf transformation 7]
359. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_3(X0) & from_2(X0,X1,X10,X11,X12,X13,X14,X15,X16,X17)) | ~select_1(X0,X1))) [rectify 358]
360. ! [X1,X0] : (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_3(X0) & from_2(X0,X1,X10,X11,X12,X13,X14,X15,X16,X17)) => (where_3(X0) & from_2(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1),sK11(X0,X1)))) [choice axiom]
361. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & ((where_3(X0) & from_2(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1),sK11(X0,X1))) | ~select_1(X0,X1))) [skolemisation 359,360]
362. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 108]
363. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_9(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | (~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4))) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_9(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [nnf transformation 9]
364. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_9(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4)) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_9(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [flattening 363]
366. ! [X0,X1] : ((where_14(X0,X1) | ~neq(X0,X1)) & (neq(X0,X1) | ~where_14(X0,X1))) [nnf transformation 110]
367. ! [X0,X1] : ((select_12(X0,X1) | ! [X2,X3,X4,X5] : (~where_14(X0,X1) | ~from_13(X0,X2,X3,X4,X5))) & (? [X2,X3,X4,X5] : (where_14(X0,X1) & from_13(X0,X2,X3,X4,X5)) | ~select_12(X0,X1))) [nnf transformation 111]
368. ! [X0,X1] : ((select_12(X0,X1) | ! [X2,X3,X4,X5] : (~where_14(X0,X1) | ~from_13(X0,X2,X3,X4,X5))) & (? [X6,X7,X8,X9] : (where_14(X0,X1) & from_13(X0,X6,X7,X8,X9)) | ~select_12(X0,X1))) [rectify 367]
369. ! [X1,X0] : (? [X6,X7,X8,X9] : (where_14(X0,X1) & from_13(X0,X6,X7,X8,X9)) => (where_14(X0,X1) & from_13(X0,sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1)))) [choice axiom]
370. ! [X0,X1] : ((select_12(X0,X1) | ! [X2,X3,X4,X5] : (~where_14(X0,X1) | ~from_13(X0,X2,X3,X4,X5))) & ((where_14(X0,X1) & from_13(X0,sK12(X0,X1),sK13(X0,X1),sK14(X0,X1),sK15(X0,X1))) | ~select_12(X0,X1))) [skolemisation 368,369]
371. ! [X0,X1] : ((where_10_subquery_11(X0,X1) | ~select_12(X0,X1)) & (select_12(X0,X1) | ~where_10_subquery_11(X0,X1))) [nnf transformation 112]
372. ! [X0,X1,X2,X3,X4] : ((from_17(X0,X1,X2,X3,X4) | ~tab2(X0,X1,X2,X3,X4)) & (tab2(X0,X1,X2,X3,X4) | ~from_17(X0,X1,X2,X3,X4))) [nnf transformation 113]
373. ! [X0,X1] : ((where_18(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_18(X0,X1))) [nnf transformation 114]
374. ! [X0,X1] : ((select_16(X0,X1) | ! [X2,X3,X4,X5] : (~where_18(X0,X1) | ~from_17(X0,X2,X3,X4,X5))) & (? [X2,X3,X4,X5] : (where_18(X0,X1) & from_17(X0,X2,X3,X4,X5)) | ~select_16(X0,X1))) [nnf transformation 115]
375. ! [X0,X1] : ((select_16(X0,X1) | ! [X2,X3,X4,X5] : (~where_18(X0,X1) | ~from_17(X0,X2,X3,X4,X5))) & (? [X6,X7,X8,X9] : (where_18(X0,X1) & from_17(X0,X6,X7,X8,X9)) | ~select_16(X0,X1))) [rectify 374]
376. ! [X1,X0] : (? [X6,X7,X8,X9] : (where_18(X0,X1) & from_17(X0,X6,X7,X8,X9)) => (where_18(X0,X1) & from_17(X0,sK16(X0,X1),sK17(X0,X1),sK18(X0,X1),sK19(X0,X1)))) [choice axiom]
377. ! [X0,X1] : ((select_16(X0,X1) | ! [X2,X3,X4,X5] : (~where_18(X0,X1) | ~from_17(X0,X2,X3,X4,X5))) & ((where_18(X0,X1) & from_17(X0,sK16(X0,X1),sK17(X0,X1),sK18(X0,X1),sK19(X0,X1))) | ~select_16(X0,X1))) [skolemisation 375,376]
378. ! [X0,X1] : ((where_10_subquery_15(X0,X1) | ~select_16(X0,X1)) & (select_16(X0,X1) | ~where_10_subquery_15(X0,X1))) [nnf transformation 116]
379. ! [X0] : ((where_10(X0) | (~where_10_subquery_15(X0,X0) | where_10_subquery_11(X0,X0))) & ((where_10_subquery_15(X0,X0) & ~where_10_subquery_11(X0,X0)) | ~where_10(X0))) [nnf transformation 18]
380. ! [X0] : ((where_10(X0) | ~where_10_subquery_15(X0,X0) | where_10_subquery_11(X0,X0)) & ((where_10_subquery_15(X0,X0) & ~where_10_subquery_11(X0,X0)) | ~where_10(X0))) [flattening 379]
381. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_10(X0) | ~from_9(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & (? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_10(X0) & from_9(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9)) | ~select_8(X0,X1))) [nnf transformation 19]
382. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_10(X0) | ~from_9(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_10(X0) & from_9(X0,X1,X10,X11,X12,X13,X14,X15,X16,X17)) | ~select_8(X0,X1))) [rectify 381]
383. ! [X1,X0] : (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_10(X0) & from_9(X0,X1,X10,X11,X12,X13,X14,X15,X16,X17)) => (where_10(X0) & from_9(X0,X1,sK20(X0,X1),sK21(X0,X1),sK22(X0,X1),sK23(X0,X1),sK24(X0,X1),sK25(X0,X1),sK26(X0,X1),sK27(X0,X1)))) [choice axiom]
384. ! [X0,X1] : ((select_8(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_10(X0) | ~from_9(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & ((where_10(X0) & from_9(X0,X1,sK20(X0,X1),sK21(X0,X1),sK22(X0,X1),sK23(X0,X1),sK24(X0,X1),sK25(X0,X1),sK26(X0,X1),sK27(X0,X1))) | ~select_8(X0,X1))) [skolemisation 382,383]
385. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_8(X0,X1)) & (select_8(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 117]
388. ! [X0,X1] : ((neq(X0,X1) | eq(X0,X1)) & (~eq(X0,X1) | ~neq(X0,X1))) [nnf transformation 126]
389. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 347]
390. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK28,sK29) | ~main_query_1(sK28,sK29)) & (main_query_2(sK28,sK29) | main_query_1(sK28,sK29))) [choice axiom]
391. (~main_query_2(sK28,sK29) | ~main_query_1(sK28,sK29)) & (main_query_2(sK28,sK29) | main_query_1(sK28,sK29)) [skolemisation 389,390]
392. ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | tab1(X0,X1,X2,X3,X4) [cnf transformation 349]
394. from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4) [cnf transformation 349]
395. ~from_6(X0,X1,X2,X3,X4) | tab2(X0,X1,X2,X3,X4) [cnf transformation 350]
396. ~tab2(X0,X1,X2,X3,X4) | from_6(X0,X1,X2,X3,X4) [cnf transformation 350]
398. ~eq(X0,X1) | where_7(X0,X1) [cnf transformation 351]
399. from_6(X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1)) | ~select_5(X0,X1) [cnf transformation 355]
401. ~from_6(X0,X2,X3,X4,X5) | ~where_7(X0,X1) | select_5(X0,X1) [cnf transformation 355]
402. ~where_3_subquery_4(X0,X1) | select_5(X0,X1) [cnf transformation 356]
403. ~select_5(X0,X1) | where_3_subquery_4(X0,X1) [cnf transformation 356]
404. where_3_subquery_4(X0,X0) | ~where_3(X0) [cnf transformation 357]
405. ~where_3_subquery_4(X0,X0) | where_3(X0) [cnf transformation 357]
406. from_2(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1),sK11(X0,X1)) | ~select_1(X0,X1) [cnf transformation 361]
407. ~select_1(X0,X1) | where_3(X0) [cnf transformation 361]
408. ~from_2(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~where_3(X0) | select_1(X0,X1) [cnf transformation 361]
409. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 362]
410. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 362]
411. ~from_9(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | tab1(X0,X1,X2,X3,X4) [cnf transformation 364]
413. from_9(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4) [cnf transformation 364]
416. ~where_14(X0,X1) | neq(X0,X1) [cnf transformation 366]
419. ~select_12(X0,X1) | where_14(X0,X1) [cnf transformation 370]
421. ~where_10_subquery_11(X0,X1) | select_12(X0,X1) [cnf transformation 371]
423. ~from_17(X0,X1,X2,X3,X4) | tab2(X0,X1,X2,X3,X4) [cnf transformation 372]
424. ~tab2(X0,X1,X2,X3,X4) | from_17(X0,X1,X2,X3,X4) [cnf transformation 372]
426. ~eq(X0,X1) | where_18(X0,X1) [cnf transformation 373]
427. from_17(X0,sK16(X0,X1),sK17(X0,X1),sK18(X0,X1),sK19(X0,X1)) | ~select_16(X0,X1) [cnf transformation 377]
429. ~from_17(X0,X2,X3,X4,X5) | ~where_18(X0,X1) | select_16(X0,X1) [cnf transformation 377]
430. ~where_10_subquery_15(X0,X1) | select_16(X0,X1) [cnf transformation 378]
431. ~select_16(X0,X1) | where_10_subquery_15(X0,X1) [cnf transformation 378]
433. where_10_subquery_15(X0,X0) | ~where_10(X0) [cnf transformation 380]
434. ~where_10_subquery_15(X0,X0) | where_10(X0) | where_10_subquery_11(X0,X0) [cnf transformation 380]
435. from_9(X0,X1,sK20(X0,X1),sK21(X0,X1),sK22(X0,X1),sK23(X0,X1),sK24(X0,X1),sK25(X0,X1),sK26(X0,X1),sK27(X0,X1)) | ~select_8(X0,X1) [cnf transformation 384]
436. ~select_8(X0,X1) | where_10(X0) [cnf transformation 384]
437. ~from_9(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~where_10(X0) | select_8(X0,X1) [cnf transformation 384]
438. ~main_query_2(X0,X1) | select_8(X0,X1) [cnf transformation 385]
439. ~select_8(X0,X1) | main_query_2(X0,X1) [cnf transformation 385]
440. eq(X0,X0) [cnf transformation 118]
448. ~neq(X0,X1) | ~eq(X0,X1) [cnf transformation 388]
521. main_query_2(sK28,sK29) | main_query_1(sK28,sK29) [cnf transformation 391]
522. ~main_query_2(sK28,sK29) | ~main_query_1(sK28,sK29) [cnf transformation 391]
524. 0 <=> main_query_1(sK28,sK29) [avatar definition]
525. main_query_1(sK28,sK29) <- {0} [avatar component clause 524]
527. 1 <=> ~main_query_1(sK28,sK29) [avatar definition]
530. 2 <=> main_query_2(sK28,sK29) [avatar definition]
531. main_query_2(sK28,sK29) <- {2} [avatar component clause 530]
533. 3 <=> ~main_query_2(sK28,sK29) [avatar definition]
534. ~main_query_2(sK28,sK29) <- {3} [avatar component clause 533]
535. ~1 | ~3 [avatar split clause 522,533,527]
536. 0 | 2 [avatar split clause 521,530,524]
538. where_7(X0,X0) [resolution 398,440]
540. select_5(X0,X0) | ~where_3(X0) [resolution 402,404]
543. where_18(X0,X0) [resolution 426,440]
545. select_16(X0,X0) | ~where_10(X0) [resolution 430,433]
548. select_8(sK28,sK29) <- {2} [resolution 438,531]
549. where_10(sK28) <- {2} [resolution 548,436]
594. tab2(X0,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1)) | ~select_5(X0,X1) [resolution 399,395]
610. tab2(X0,sK16(X0,X1),sK17(X0,X1),sK18(X0,X1),sK19(X0,X1)) | ~select_16(X0,X1) [resolution 427,423]
620. ~tab2(X20,X21,X22,X23,X24) | ~tab1(X25,X26,X27,X28,X29) | ~where_3(X25) | select_1(X25,X26) [resolution 394,408]
622. 4 <=> ! [X25,X27,X29,X26,X28] : (~tab1(X25,X26,X27,X28,X29) | select_1(X25,X26) | ~where_3(X25)) [avatar definition]
623. ~tab1(X25,X26,X27,X28,X29) | select_1(X25,X26) | ~where_3(X25) <- {4} [avatar component clause 622]
625. 6 <=> ! [X20,X22,X21,X23,X24] : ~tab2(X20,X21,X22,X23,X24) [avatar definition]
626. ~tab2(X20,X21,X22,X23,X24) <- {6} [avatar component clause 625]
627. 4 | 6 [avatar split clause 620,625,622]
631. ~tab2(X20,X21,X22,X23,X24) | ~tab1(X25,X26,X27,X28,X29) | ~where_10(X25) | select_8(X25,X26) [resolution 413,437]
633. 8 <=> ! [X25,X27,X29,X26,X28] : (~tab1(X25,X26,X27,X28,X29) | select_8(X25,X26) | ~where_10(X25)) [avatar definition]
634. ~tab1(X25,X26,X27,X28,X29) | select_8(X25,X26) | ~where_10(X25) <- {8} [avatar component clause 633]
635. 8 | 6 [avatar split clause 631,625,633]
640. from_17(X4,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5),sK3(X4,X5)) | ~select_5(X4,X5) [resolution 594,424]
691. from_6(X0,sK16(X0,X1),sK17(X0,X1),sK18(X0,X1),sK19(X0,X1)) | ~select_16(X0,X1) [resolution 610,396]
699. ~select_16(X21,X22) <- {6} [resolution 610,626]
706. ~where_10(X0) <- {6} [resolution 699,545]
708. $false <- {2, 6} [resolution 706,549]
709. ~2 | ~6 [avatar contradiction clause 708]
711. tab1(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1)) | ~select_1(X0,X1) [resolution 406,392]
730. ~select_1(X17,X18) | select_8(X17,X18) | ~where_10(X17) <- {8} [resolution 711,634]
731. tab1(X0,X1,sK20(X0,X1),sK21(X0,X1),sK22(X0,X1)) | ~select_8(X0,X1) [resolution 435,411]
749. ~select_8(X15,X16) | select_1(X15,X16) | ~where_3(X15) <- {4} [resolution 731,623]
751. select_1(sK28,sK29) | ~where_3(sK28) <- {2, 4} [resolution 749,548]
753. 10 <=> where_3(sK28) [avatar definition]
754. where_3(sK28) <- {10} [avatar component clause 753]
756. 11 <=> ~where_3(sK28) [avatar definition]
757. ~where_3(sK28) <- {11} [avatar component clause 756]
762. 12 <=> select_1(sK28,sK29) [avatar definition]
763. select_1(sK28,sK29) <- {12} [avatar component clause 762]
764. ~11 | 12 | ~2 | ~4 [avatar split clause 751,622,530,762,756]
933. ~where_18(X2,X4) | ~select_5(X2,X3) | select_16(X2,X4) [resolution 640,429]
944. ~select_5(X0,X1) | select_16(X0,X0) [resolution 933,543]
950. select_16(X0,X0) | ~where_3(X0) [resolution 944,540]
960. where_10_subquery_15(X9,X9) | ~where_3(X9) [resolution 950,431]
970. where_10_subquery_11(X5,X5) | where_10(X5) | ~where_3(X5) [resolution 960,434]
979. select_12(X5,X5) | ~where_3(X5) | where_10(X5) [resolution 970,421]
986. where_14(X0,X0) | where_10(X0) | ~where_3(X0) [resolution 979,419]
1012. neq(X11,X11) | ~where_3(X11) | where_10(X11) [resolution 986,416]
1031. ~where_3(X0) | where_10(X0) | ~eq(X0,X0) [resolution 1012,448]
1033. ~where_3(X0) | where_10(X0) [subsumption resolution 1031,440]
1295. ~select_16(X2,X3) | ~where_7(X2,X4) | select_5(X2,X4) [resolution 691,401]
1315. ~where_7(X2,X3) | select_5(X2,X3) | ~where_10(X2) [resolution 1295,545]
1329. select_5(X0,X0) | ~where_10(X0) [resolution 1315,538]
1344. where_3_subquery_4(X1,X1) | ~where_10(X1) [resolution 1329,403]
1364. ~where_10(X0) | where_3(X0) [resolution 1344,405]
1381. where_3(sK28) <- {2} [resolution 1364,549]
1382. $false <- {2, 11} [subsumption resolution 1381,757]
1383. ~2 | 11 [avatar contradiction clause 1382]
1386. select_1(sK28,sK29) <- {0} [resolution 525,409]
1387. 12 | ~0 [avatar split clause 1386,524,762]
1388. where_10(sK28) <- {10} [resolution 754,1033]
1399. select_8(sK28,sK29) | ~where_10(sK28) <- {8, 12} [resolution 763,730]
1402. main_query_1(sK28,sK29) <- {12} [resolution 763,410]
1403. where_3(sK28) <- {12} [resolution 763,407]
1406. 0 | ~12 [avatar split clause 1402,762,524]
1407. select_8(sK28,sK29) <- {8, 10, 12} [subsumption resolution 1399,1388]
1414. main_query_2(sK28,sK29) <- {8, 10, 12} [resolution 1407,439]
1416. $false <- {3, 8, 10, 12} [subsumption resolution 1414,534]
1417. 3 | ~8 | ~10 | ~12 [avatar contradiction clause 1416]
1419. ~select_5(X0,X1) <- {6} [resolution 626,594]
1438. ~where_3(X1) <- {6} [resolution 1419,540]
1444. $false <- {6, 12} [subsumption resolution 1403,1438]
1445. ~6 | ~12 [avatar contradiction clause 1444]
1446. $false <- {11, 12} [subsumption resolution 1403,757]
1447. 11 | ~12 [avatar contradiction clause 1446]
1448. $false [avatar sat refutation 535,536,627,635,709,764,1383,1387,1406,1417,1445,1447]
% SZS output end Proof for query_29_query_30_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5756
% Time elapsed: 0.198 s
% ------------------------------
% ------------------------------
