% Refutation found. Thanks to Tanya!
% SZS status Theorem for query05_query07_tptp4X
% SZS output start Proof for query05_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X5,X1,X2,X3,X4] : (from_5(X5,X1,X2,X3,X4) <=> tab1(X5,X1,X2,X3,X4)) [input]
6. ! [X5] : (where_6(X5) <=> (~gt(X5,X5) & ~lt(X5,X5))) [input]
7. ! [X5] : (select_4(X5) <=> ? [X1,X2,X3,X4] : (where_6(X5) & from_5(X5,X1,X2,X3,X4))) [input]
8. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
9. ! [X6] : eq(X6,X6) [input]
15. ! [X6,X7] : (lt(X6,X7) <=> (~eq(X6,X7) & lte(X6,X7))) [input]
16. ! [X6,X7] : (gt(X6,X7) <=> lt(X7,X6)) [input]
42. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
43. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 42]
44. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
45. ! [X0] : (where_6(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [rectify 6]
46. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
47. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
48. ! [X0] : eq(X0,X0) [rectify 9]
54. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 15]
55. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 16]
132. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 43]
133. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
134. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 135]
137. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 136,137]
139. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
140. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 44]
141. ! [X0] : ((where_6(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [nnf transformation 45]
142. ! [X0] : ((where_6(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [flattening 141]
143. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 46]
144. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 143]
145. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
146. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 144,145]
147. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 47]
150. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 54]
151. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 150]
152. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 55]
153. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 132]
154. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
155. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 153,154]
156. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 133]
157. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 133]
159. ~eq(X0,X0) | where_3(X0) [cnf transformation 134]
160. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 138]
162. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 138]
163. ~main_query_1(X0) | select_1(X0) [cnf transformation 139]
164. ~select_1(X0) | main_query_1(X0) [cnf transformation 139]
165. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 140]
166. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 140]
169. gt(X0,X0) | where_6(X0) | lt(X0,X0) [cnf transformation 142]
170. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 146]
172. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 146]
173. ~main_query_2(X0) | select_4(X0) [cnf transformation 147]
174. ~select_4(X0) | main_query_2(X0) [cnf transformation 147]
175. eq(X0,X0) [cnf transformation 48]
184. ~lt(X0,X1) | ~eq(X0,X1) [cnf transformation 151]
186. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 152]
211. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 155]
212. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 155]
237. 0 <=> main_query_1(sK8) [avatar definition]
238. main_query_1(sK8) <- {0} [avatar component clause 237]
240. 1 <=> ~main_query_1(sK8) [avatar definition]
241. ~main_query_1(sK8) <- {1} [avatar component clause 240]
243. 2 <=> main_query_2(sK8) [avatar definition]
244. main_query_2(sK8) <- {2} [avatar component clause 243]
246. 3 <=> ~main_query_2(sK8) [avatar definition]
247. ~main_query_2(sK8) <- {3} [avatar component clause 246]
248. ~1 | ~3 [avatar split clause 212,246,240]
249. 0 | 2 [avatar split clause 211,243,237]
250. select_4(sK8) <- {2} [resolution 173,244]
254. where_3(X1) [resolution 159,175]
262. where_6(X1) | lt(X1,X1) | lt(X1,X1) [resolution 169,186]
263. lt(X1,X1) | where_6(X1) [duplicate literal removal 262]
268. where_6(X2) | ~eq(X2,X2) [resolution 263,184]
270. where_6(X2) [subsumption resolution 268,175]
276. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 160,156]
278. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 170,165]
280. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 276,166]
283. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 278,157]
285. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 280,172]
286. ~select_1(X1) | select_4(X1) [subsumption resolution 285,270]
288. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 283,162]
289. ~select_4(X1) | select_1(X1) [subsumption resolution 288,254]
290. select_1(sK8) <- {2} [resolution 289,250]
292. main_query_1(sK8) <- {2} [resolution 290,164]
294. $false <- {1, 2} [subsumption resolution 292,241]
295. 1 | ~2 [avatar contradiction clause 294]
296. select_1(sK8) <- {0} [resolution 238,163]
297. select_4(sK8) <- {0} [resolution 296,286]
301. main_query_2(sK8) <- {0} [resolution 297,174]
303. $false <- {0, 3} [subsumption resolution 301,247]
304. ~0 | 3 [avatar contradiction clause 303]
305. $false [avatar sat refutation 248,249,295,304]
% SZS output end Proof for query05_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
