% Refutation found. Thanks to Tanya!
% SZS status Theorem for query04_query07_tptp4X
% SZS output start Proof for query04_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
7. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
9. ! [X6] : eq(X6,X6) [input]
15. ! [X6,X7] : (lt(X6,X7) <=> (~eq(X6,X7) & lte(X6,X7))) [input]
16. ! [X6,X7] : (gt(X6,X7) <=> lt(X7,X6)) [input]
42. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
43. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 42]
44. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
45. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
46. ! [X0] : eq(X0,X0) [rectify 9]
52. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 15]
53. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 16]
79. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 43]
131. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 79]
132. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
133. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 134]
136. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 135,136]
138. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 44]
139. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
140. ! [X0] : ((where_6(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [nnf transformation 6]
141. ! [X0] : ((where_6(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [flattening 140]
142. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 7]
143. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 142]
144. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
145. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 143,144]
146. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 45]
149. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 52]
150. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 149]
151. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 53]
152. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 131]
153. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
154. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 152,153]
155. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 132]
156. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 132]
158. ~eq(X0,X0) | where_3(X0) [cnf transformation 133]
159. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 137]
161. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 137]
162. ~main_query_1(X0) | select_1(X0) [cnf transformation 138]
163. ~select_1(X0) | main_query_1(X0) [cnf transformation 138]
164. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 139]
165. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 139]
168. gt(X0,X0) | where_6(X0) | lt(X0,X0) [cnf transformation 141]
169. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 145]
171. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 145]
172. ~main_query_2(X0) | select_4(X0) [cnf transformation 146]
173. ~select_4(X0) | main_query_2(X0) [cnf transformation 146]
174. eq(X0,X0) [cnf transformation 46]
183. ~lt(X0,X1) | ~eq(X0,X1) [cnf transformation 150]
185. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 151]
210. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 154]
211. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 154]
236. 0 <=> main_query_1(sK8) [avatar definition]
237. main_query_1(sK8) <- {0} [avatar component clause 236]
239. 1 <=> ~main_query_1(sK8) [avatar definition]
240. ~main_query_1(sK8) <- {1} [avatar component clause 239]
242. 2 <=> main_query_2(sK8) [avatar definition]
243. main_query_2(sK8) <- {2} [avatar component clause 242]
245. 3 <=> ~main_query_2(sK8) [avatar definition]
246. ~main_query_2(sK8) <- {3} [avatar component clause 245]
247. ~1 | ~3 [avatar split clause 211,245,239]
248. 0 | 2 [avatar split clause 210,242,236]
249. select_4(sK8) <- {2} [resolution 172,243]
253. where_3(X1) [resolution 158,174]
261. where_6(X1) | lt(X1,X1) | lt(X1,X1) [resolution 168,185]
262. lt(X1,X1) | where_6(X1) [duplicate literal removal 261]
267. where_6(X2) | ~eq(X2,X2) [resolution 262,183]
269. where_6(X2) [subsumption resolution 267,174]
275. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 159,155]
277. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 169,164]
279. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 275,165]
282. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 277,156]
284. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 279,171]
285. ~select_1(X1) | select_4(X1) [subsumption resolution 284,269]
287. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 282,161]
288. ~select_4(X1) | select_1(X1) [subsumption resolution 287,253]
289. select_1(sK8) <- {2} [resolution 288,249]
291. main_query_1(sK8) <- {2} [resolution 289,163]
293. $false <- {1, 2} [subsumption resolution 291,240]
294. 1 | ~2 [avatar contradiction clause 293]
295. select_1(sK8) <- {0} [resolution 237,162]
296. select_4(sK8) <- {0} [resolution 295,285]
300. main_query_2(sK8) <- {0} [resolution 296,173]
302. $false <- {0, 3} [subsumption resolution 300,246]
303. ~0 | 3 [avatar contradiction clause 302]
304. $false [avatar sat refutation 247,248,294,303]
% SZS output end Proof for query04_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
