% Refutation found. Thanks to Tanya!
% SZS status Theorem for query07_query08_tptp4X
% SZS output start Proof for query07_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X5,X1,X2,X3,X4] : (from_5(X5,X1,X2,X3,X4) <=> tab1(X5,X1,X2,X3,X4)) [input]
6. ! [X5] : (where_6(X5) <=> (gte(X5,X5) & lte(X5,X5))) [input]
7. ! [X5] : (select_4(X5) <=> ? [X1,X2,X3,X4] : (where_6(X5) & from_5(X5,X1,X2,X3,X4))) [input]
8. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
9. ! [X6] : eq(X6,X6) [input]
12. ! [X6] : lte(X6,X6) [input]
15. ! [X6,X7] : (lt(X6,X7) <=> (~eq(X6,X7) & lte(X6,X7))) [input]
16. ! [X6,X7] : (gt(X6,X7) <=> lt(X7,X6)) [input]
18. ! [X6,X7] : (gte(X6,X7) <=> lte(X7,X6)) [input]
42. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
43. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 42]
44. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
45. ! [X0] : (where_6(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 6]
46. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
47. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
48. ! [X0] : eq(X0,X0) [rectify 9]
51. ! [X0] : lte(X0,X0) [rectify 12]
54. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 15]
55. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 16]
57. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 18]
132. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 43]
133. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
134. ! [X0] : ((where_3(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
135. ! [X0] : ((where_3(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [flattening 134]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 136]
138. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
139. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 137,138]
140. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
141. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 44]
142. ! [X0] : ((where_6(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [nnf transformation 45]
143. ! [X0] : ((where_6(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [flattening 142]
144. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 46]
145. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 144]
146. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
147. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 145,146]
148. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 47]
151. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 54]
152. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 151]
153. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 55]
154. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 57]
155. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 132]
156. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
157. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 155,156]
158. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 133]
159. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 133]
162. gt(X0,X0) | where_3(X0) | lt(X0,X0) [cnf transformation 135]
163. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 139]
165. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 139]
166. ~main_query_1(X0) | select_1(X0) [cnf transformation 140]
167. ~select_1(X0) | main_query_1(X0) [cnf transformation 140]
168. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 141]
169. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 141]
172. ~gte(X0,X0) | where_6(X0) | ~lte(X0,X0) [cnf transformation 143]
173. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 147]
175. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 147]
176. ~main_query_2(X0) | select_4(X0) [cnf transformation 148]
177. ~select_4(X0) | main_query_2(X0) [cnf transformation 148]
178. eq(X0,X0) [cnf transformation 48]
181. lte(X0,X0) [cnf transformation 51]
187. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 152]
189. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 153]
192. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 154]
216. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 157]
217. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 157]
242. 0 <=> main_query_1(sK8) [avatar definition]
243. main_query_1(sK8) <- {0} [avatar component clause 242]
245. 1 <=> ~main_query_1(sK8) [avatar definition]
246. ~main_query_1(sK8) <- {1} [avatar component clause 245]
248. 2 <=> main_query_2(sK8) [avatar definition]
249. main_query_2(sK8) <- {2} [avatar component clause 248]
251. 3 <=> ~main_query_2(sK8) [avatar definition]
252. ~main_query_2(sK8) <- {3} [avatar component clause 251]
253. ~1 | ~3 [avatar split clause 217,251,245]
254. 0 | 2 [avatar split clause 216,248,242]
255. select_4(sK8) <- {2} [resolution 176,249]
261. ~lt(X0,X0) [resolution 187,178]
264. gte(X1,X1) [resolution 192,181]
267. where_3(X1) | lt(X1,X1) | lt(X1,X1) [resolution 162,189]
268. where_3(X1) | lt(X1,X1) [duplicate literal removal 267]
269. where_3(X1) [subsumption resolution 268,261]
270. where_6(X0) | ~lte(X0,X0) [resolution 172,264]
272. where_6(X0) [subsumption resolution 270,181]
280. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 163,158]
282. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 173,168]
284. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 280,169]
287. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 282,159]
289. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 284,175]
290. ~select_1(X1) | select_4(X1) [subsumption resolution 289,272]
292. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 287,165]
293. ~select_4(X1) | select_1(X1) [subsumption resolution 292,269]
294. select_1(sK8) <- {2} [resolution 293,255]
296. main_query_1(sK8) <- {2} [resolution 294,167]
298. $false <- {1, 2} [subsumption resolution 296,246]
299. 1 | ~2 [avatar contradiction clause 298]
300. select_1(sK8) <- {0} [resolution 243,166]
301. select_4(sK8) <- {0} [resolution 300,290]
305. main_query_2(sK8) <- {0} [resolution 301,177]
307. $false <- {0, 3} [subsumption resolution 305,252]
308. ~0 | 3 [avatar contradiction clause 307]
309. $false [avatar sat refutation 253,254,299,308]
% SZS output end Proof for query07_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.021 s
% ------------------------------
% ------------------------------
