% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_02_query_06_tptp4X
% SZS output start Proof for query_02_query_06_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0,X5] : (main_query_1(X0,X5) <=> select_1(X0,X5)) [input]
4. ! [X6,X1,X2,X3,X4] : (from_4(X6,X1,X2,X3,X4) <=> tab1(X6,X1,X2,X3,X4)) [input]
5. ! [X6] : (where_5(X6) <=> (gte(X6,X6) & lte(X6,X6))) [input]
6. ! [X6,X1] : (select_3(X6,X1) <=> ? [X2,X3,X4] : (where_5(X6) & from_4(X6,X1,X2,X3,X4))) [input]
7. ! [X7,X5] : (main_query_2(X7,X5) <=> select_3(X7,X5)) [input]
11. ! [X8] : lte(X8,X8) [input]
17. ! [X8,X9] : (gte(X8,X9) <=> lte(X9,X8)) [input]
46. ! [X0,X5] : (main_query_1(X0,X5) <=> main_query_2(X0,X5)) [input]
47. ~! [X0,X5] : (main_query_1(X0,X5) <=> main_query_2(X0,X5)) [negated conjecture 46]
48. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 3]
50. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
51. ! [X0] : (where_5(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 5]
52. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
53. ! [X0,X1] : (main_query_2(X0,X1) <=> select_3(X0,X1)) [rectify 7]
57. ! [X0] : lte(X0,X0) [rectify 11]
63. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 17]
92. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 47]
154. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 92]
155. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
156. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0,X1))) [nnf transformation 48]
157. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) | ~select_1(X0,X1))) [rectify 156]
158. ! [X1,X0] : (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) => from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) [choice axiom]
159. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1))) [skolemisation 157,158]
160. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 49]
161. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 50]
162. ! [X0] : ((where_5(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [nnf transformation 51]
163. ! [X0] : ((where_5(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_5(X0))) [flattening 162]
164. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0,X1))) [nnf transformation 52]
165. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) | ~select_3(X0,X1))) [rectify 164]
166. ! [X1,X0] : (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) => (where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
167. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_3(X0,X1))) [skolemisation 165,166]
168. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 53]
171. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 63]
172. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 154]
173. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
174. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 172,173]
175. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 155]
176. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 155]
177. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 159]
178. ~from_2(X0,X1,X2,X3,X4) | select_1(X0,X1) [cnf transformation 159]
179. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 160]
180. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 160]
181. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 161]
182. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 161]
185. ~gte(X0,X0) | where_5(X0) | ~lte(X0,X0) [cnf transformation 163]
186. from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [cnf transformation 167]
188. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0,X1) [cnf transformation 167]
189. ~main_query_2(X0,X1) | select_3(X0,X1) [cnf transformation 168]
190. ~select_3(X0,X1) | main_query_2(X0,X1) [cnf transformation 168]
194. lte(X0,X0) [cnf transformation 57]
200. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 171]
229. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 174]
230. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 174]
232. 0 <=> main_query_1(sK6,sK7) [avatar definition]
233. main_query_1(sK6,sK7) <- {0} [avatar component clause 232]
235. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
236. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 235]
238. 2 <=> main_query_2(sK6,sK7) [avatar definition]
239. main_query_2(sK6,sK7) <- {2} [avatar component clause 238]
241. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
242. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 241]
243. ~1 | ~3 [avatar split clause 230,241,235]
244. 0 | 2 [avatar split clause 229,238,232]
245. select_3(sK6,sK7) <- {2} [resolution 189,239]
253. gte(X1,X1) [resolution 200,194]
255. where_5(X0) | ~lte(X0,X0) [resolution 185,253]
257. where_5(X0) [subsumption resolution 255,194]
271. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 177,175]
273. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 186,181]
286. from_4(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 271,182]
292. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 273,176]
300. ~select_1(X2,X3) | ~where_5(X2) | select_3(X2,X3) [resolution 286,188]
306. ~select_1(X2,X3) | select_3(X2,X3) [subsumption resolution 300,257]
322. ~select_3(X2,X3) | select_1(X2,X3) [resolution 292,178]
335. select_1(sK6,sK7) <- {2} [resolution 322,245]
344. main_query_1(sK6,sK7) <- {2} [resolution 335,180]
345. $false <- {1, 2} [subsumption resolution 344,236]
346. 1 | ~2 [avatar contradiction clause 345]
347. select_1(sK6,sK7) <- {0} [resolution 233,179]
348. select_3(sK6,sK7) <- {0} [resolution 347,306]
351. main_query_2(sK6,sK7) <- {0} [resolution 348,190]
353. $false <- {0, 3} [subsumption resolution 351,242]
354. ~0 | 3 [avatar contradiction clause 353]
355. $false [avatar sat refutation 243,244,346,354]
% SZS output end Proof for query_02_query_06_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.010 s
% ------------------------------
% ------------------------------
