% Refutation found. Thanks to Tanya!
% SZS status Theorem for query11_query12_tptp4X
% SZS output start Proof for query11_query12_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8] : (where_6(X7,X8) <=> (lte(X8,X7) & lte(X7,X8))) [input]
7. ! [X7,X8] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X7,X8) & from_5(X7,X8,X9,X10,X11))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
13. ! [X12,X13] : ((lte(X13,X12) & lte(X12,X13)) <=> eq(X12,X13)) [input]
54. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
55. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 54]
56. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
57. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
58. ! [X0,X1] : (where_6(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [rectify 6]
59. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
60. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
65. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
106. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 55]
182. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 106]
183. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
184. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
185. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
186. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 185]
187. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
188. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 186,187]
189. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 56]
190. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 57]
191. ! [X0,X1] : ((where_6(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 58]
192. ! [X0,X1] : ((where_6(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [flattening 191]
193. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 59]
194. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 193]
195. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
196. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 194,195]
197. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 60]
198. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 65]
199. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 198]
200. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 182]
201. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
202. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 200,201]
203. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 183]
204. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 183]
205. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 184]
206. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 184]
207. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 188]
208. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 188]
209. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 188]
210. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 189]
211. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 189]
212. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 190]
213. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 190]
214. ~where_6(X0,X1) | lte(X0,X1) [cnf transformation 192]
215. ~where_6(X0,X1) | lte(X1,X0) [cnf transformation 192]
216. ~lte(X1,X0) | where_6(X0,X1) | ~lte(X0,X1) [cnf transformation 192]
217. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 196]
218. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 196]
219. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 196]
220. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 197]
221. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 197]
226. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 199]
227. ~eq(X0,X1) | lte(X0,X1) [cnf transformation 199]
228. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 199]
265. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 202]
266. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 202]
303. 0 <=> main_query_1(sK6,sK7) [avatar definition]
304. main_query_1(sK6,sK7) <- {0} [avatar component clause 303]
306. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
307. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 306]
309. 2 <=> main_query_2(sK6,sK7) [avatar definition]
310. main_query_2(sK6,sK7) <- {2} [avatar component clause 309]
312. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
313. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 312]
314. ~1 | ~3 [avatar split clause 266,312,306]
315. 0 | 2 [avatar split clause 265,309,303]
318. select_4(sK6,sK7) <- {2} [resolution 220,310]
319. where_6(sK6,sK7) <- {2} [resolution 318,218]
321. lte(sK7,sK6) <- {2} [resolution 319,215]
322. lte(sK6,sK7) <- {2} [resolution 319,214]
335. eq(sK6,sK7) | ~lte(sK6,sK7) <- {2} [resolution 226,321]
337. eq(sK6,sK7) <- {2} [subsumption resolution 335,322]
348. where_3(sK6,sK7) <- {2} [resolution 337,206]
363. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 207,203]
367. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 217,212]
368. from_5(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 363,213]
371. from_2(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 367,204]
372. ~where_6(X0,X1) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 368,219]
374. ~select_4(X0,X1) | ~where_3(X0,X1) | select_1(X0,X1) [resolution 371,209]
392. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 374,318]
393. select_1(sK6,sK7) <- {2} [subsumption resolution 392,348]
394. main_query_1(sK6,sK7) <- {2} [resolution 393,211]
396. $false <- {1, 2} [subsumption resolution 394,307]
397. 1 | ~2 [avatar contradiction clause 396]
398. select_1(sK6,sK7) <- {0} [resolution 304,210]
400. where_3(sK6,sK7) <- {0} [resolution 398,208]
401. eq(sK6,sK7) <- {0} [resolution 400,205]
403. lte(sK7,sK6) <- {0} [resolution 401,228]
404. lte(sK6,sK7) <- {0} [resolution 401,227]
409. where_6(sK6,sK7) | ~lte(sK6,sK7) <- {0} [resolution 403,216]
410. where_6(sK6,sK7) <- {0} [subsumption resolution 409,404]
422. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 410,372]
425. select_4(sK6,sK7) <- {0} [subsumption resolution 422,398]
433. main_query_2(sK6,sK7) <- {0} [resolution 425,221]
435. $false <- {0, 3} [subsumption resolution 433,313]
436. ~0 | 3 [avatar contradiction clause 435]
437. $false [avatar sat refutation 314,315,397,436]
% SZS output end Proof for query11_query12_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
