% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_07_query_09_tptp4X
% SZS output start Proof for query_07_query_09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : ($true & from_5(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (subquery_from_3(X5) <=> select_4(X5)) [input]
4. ! [X6] : (from_2(X6) <=> subquery_from_3(X6)) [input]
5. ! [X6] : (select_1(X6) <=> ($true & from_2(X6))) [input]
6. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
7. ! [X7,X8,X9,X10,X11] : (from_7(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
8. ! [X7] : (where_8(X7) <=> (gte(X7,X7) & lte(X7,X7))) [input]
9. ! [X7] : (select_6(X7) <=> ? [X8,X9,X10,X11] : (where_8(X7) & from_7(X7,X8,X9,X10,X11))) [input]
10. ! [X5] : (main_query_2(X5) <=> select_6(X5)) [input]
14. ! [X12] : lte(X12,X12) [input]
21. ! [X12,X13] : (gte(X12,X13) <=> lte(X13,X12)) [input]
42. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
43. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 42]
44. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4)) [true and false elimination 2]
45. ! [X0] : (subquery_from_3(X0) <=> select_4(X0)) [rectify 3]
46. ! [X0] : (from_2(X0) <=> subquery_from_3(X0)) [rectify 4]
47. ! [X0] : (select_1(X0) <=> ($true & from_2(X0))) [rectify 5]
48. ! [X0] : (select_1(X0) <=> from_2(X0)) [true and false elimination 47]
49. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 6]
50. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 7]
51. ! [X0] : (where_8(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 8]
52. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4))) [rectify 9]
53. ! [X0] : (main_query_2(X0) <=> select_6(X0)) [rectify 10]
57. ! [X0] : lte(X0,X0) [rectify 14]
64. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 21]
85. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 43]
131. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 85]
132. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 1]
133. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_5(X0,X1,X2,X3,X4) | ~select_4(X0))) [nnf transformation 44]
134. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) | ~select_4(X0))) [rectify 133]
135. ! [X0] : (? [X5,X6,X7,X8] : from_5(X0,X5,X6,X7,X8) => from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
136. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : ~from_5(X0,X1,X2,X3,X4)) & (from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0))) [skolemisation 134,135]
137. ! [X0] : ((subquery_from_3(X0) | ~select_4(X0)) & (select_4(X0) | ~subquery_from_3(X0))) [nnf transformation 45]
138. ! [X0] : ((from_2(X0) | ~subquery_from_3(X0)) & (subquery_from_3(X0) | ~from_2(X0))) [nnf transformation 46]
139. ! [X0] : ((select_1(X0) | ~from_2(X0)) & (from_2(X0) | ~select_1(X0))) [nnf transformation 48]
140. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 49]
141. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 50]
142. ! [X0] : ((where_8(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_8(X0))) [nnf transformation 51]
143. ! [X0] : ((where_8(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_8(X0))) [flattening 142]
144. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4)) | ~select_6(X0))) [nnf transformation 52]
145. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) | ~select_6(X0))) [rectify 144]
146. ! [X0] : (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) => (where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
147. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & ((where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_6(X0))) [skolemisation 145,146]
148. ! [X0] : ((main_query_2(X0) | ~select_6(X0)) & (select_6(X0) | ~main_query_2(X0))) [nnf transformation 53]
151. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 64]
152. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 131]
153. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
154. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 152,153]
155. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 132]
156. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 132]
157. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [cnf transformation 136]
158. ~from_5(X0,X1,X2,X3,X4) | select_4(X0) [cnf transformation 136]
159. ~subquery_from_3(X0) | select_4(X0) [cnf transformation 137]
160. ~select_4(X0) | subquery_from_3(X0) [cnf transformation 137]
161. ~from_2(X0) | subquery_from_3(X0) [cnf transformation 138]
162. ~subquery_from_3(X0) | from_2(X0) [cnf transformation 138]
163. ~select_1(X0) | from_2(X0) [cnf transformation 139]
164. ~from_2(X0) | select_1(X0) [cnf transformation 139]
165. ~main_query_1(X0) | select_1(X0) [cnf transformation 140]
166. ~select_1(X0) | main_query_1(X0) [cnf transformation 140]
167. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 141]
168. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 141]
171. ~gte(X0,X0) | where_8(X0) | ~lte(X0,X0) [cnf transformation 143]
172. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 147]
174. ~from_7(X0,X1,X2,X3,X4) | ~where_8(X0) | select_6(X0) [cnf transformation 147]
175. ~main_query_2(X0) | select_6(X0) [cnf transformation 148]
176. ~select_6(X0) | main_query_2(X0) [cnf transformation 148]
180. lte(X0,X0) [cnf transformation 57]
187. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 151]
208. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 154]
209. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 154]
211. 0 <=> main_query_1(sK8) [avatar definition]
212. main_query_1(sK8) <- {0} [avatar component clause 211]
214. 1 <=> ~main_query_1(sK8) [avatar definition]
215. ~main_query_1(sK8) <- {1} [avatar component clause 214]
217. 2 <=> main_query_2(sK8) [avatar definition]
218. main_query_2(sK8) <- {2} [avatar component clause 217]
220. 3 <=> ~main_query_2(sK8) [avatar definition]
221. ~main_query_2(sK8) <- {3} [avatar component clause 220]
222. ~1 | ~3 [avatar split clause 209,220,214]
223. 0 | 2 [avatar split clause 208,217,211]
224. select_6(sK8) <- {2} [resolution 175,218]
233. gte(X1,X1) [resolution 187,180]
237. where_8(X0) | ~lte(X0,X0) [resolution 171,233]
239. where_8(X0) [subsumption resolution 237,180]
284. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [resolution 157,155]
307. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 172,167]
310. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_4(X0) [resolution 284,168]
314. from_5(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 307,156]
317. ~select_4(X1) | ~where_8(X1) | select_6(X1) [resolution 310,174]
318. ~select_4(X1) | select_6(X1) [subsumption resolution 317,239]
322. ~select_6(X1) | select_4(X1) [resolution 314,158]
330. select_4(sK8) <- {2} [resolution 322,224]
334. subquery_from_3(sK8) <- {2} [resolution 330,160]
337. from_2(sK8) <- {2} [resolution 334,162]
341. select_1(sK8) <- {2} [resolution 337,164]
345. main_query_1(sK8) <- {2} [resolution 341,166]
347. $false <- {1, 2} [subsumption resolution 345,215]
348. 1 | ~2 [avatar contradiction clause 347]
349. select_1(sK8) <- {0} [resolution 212,165]
353. from_2(sK8) <- {0} [resolution 349,163]
357. subquery_from_3(sK8) <- {0} [resolution 353,161]
361. select_4(sK8) <- {0} [resolution 357,159]
364. select_6(sK8) <- {0} [resolution 361,318]
369. main_query_2(sK8) <- {0} [resolution 364,176]
371. $false <- {0, 3} [subsumption resolution 369,221]
372. ~0 | 3 [avatar contradiction clause 371]
373. $false [avatar sat refutation 222,223,348,372]
% SZS output end Proof for query_07_query_09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.037 s
% ------------------------------
% ------------------------------
