% Refutation found. Thanks to Tanya!
% SZS status Theorem for query06_query08_tptp4X
% SZS output start Proof for query06_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (query_2(X5) <=> select_3(X5)) [input]
4. ! [X6] : (from_5(X6) <=> query_2(X6)) [input]
5. ! [X6] : (select_1(X6) <=> ($true & from_5(X6))) [input]
6. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
7. ! [X7,X8,X9,X10,X11] : (from_7(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
8. ! [X7] : (where_8(X7) <=> (gte(X7,X7) & lte(X7,X7))) [input]
9. ! [X7] : (select_6(X7) <=> ? [X8,X9,X10,X11] : (where_8(X7) & from_7(X7,X8,X9,X10,X11))) [input]
10. ! [X5] : (main_query_2(X5) <=> select_6(X5)) [input]
14. ! [X12] : lte(X12,X12) [input]
20. ! [X12,X13] : (gte(X12,X13) <=> lte(X13,X12)) [input]
44. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
45. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 44]
46. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 2]
47. ! [X0] : (query_2(X0) <=> select_3(X0)) [rectify 3]
48. ! [X0] : (from_5(X0) <=> query_2(X0)) [rectify 4]
49. ! [X0] : (select_1(X0) <=> ($true & from_5(X0))) [rectify 5]
50. ! [X0] : (select_1(X0) <=> from_5(X0)) [true and false elimination 49]
51. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 6]
52. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 7]
53. ! [X0] : (where_8(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 8]
54. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4))) [rectify 9]
55. ! [X0] : (main_query_2(X0) <=> select_6(X0)) [rectify 10]
59. ! [X0] : lte(X0,X0) [rectify 14]
65. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 20]
89. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 45]
141. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 89]
142. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 1]
143. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 46]
144. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 143]
145. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
146. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0))) [skolemisation 144,145]
147. ! [X0] : ((query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~query_2(X0))) [nnf transformation 47]
148. ! [X0] : ((from_5(X0) | ~query_2(X0)) & (query_2(X0) | ~from_5(X0))) [nnf transformation 48]
149. ! [X0] : ((select_1(X0) | ~from_5(X0)) & (from_5(X0) | ~select_1(X0))) [nnf transformation 50]
150. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 51]
151. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 52]
152. ! [X0] : ((where_8(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_8(X0))) [nnf transformation 53]
153. ! [X0] : ((where_8(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_8(X0))) [flattening 152]
154. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4)) | ~select_6(X0))) [nnf transformation 54]
155. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) | ~select_6(X0))) [rectify 154]
156. ! [X0] : (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) => (where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
157. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & ((where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_6(X0))) [skolemisation 155,156]
158. ! [X0] : ((main_query_2(X0) | ~select_6(X0)) & (select_6(X0) | ~main_query_2(X0))) [nnf transformation 55]
161. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 65]
162. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 141]
163. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
164. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 162,163]
165. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 142]
166. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 142]
167. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [cnf transformation 146]
168. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 146]
169. ~query_2(X0) | select_3(X0) [cnf transformation 147]
170. ~select_3(X0) | query_2(X0) [cnf transformation 147]
171. ~from_5(X0) | query_2(X0) [cnf transformation 148]
172. ~query_2(X0) | from_5(X0) [cnf transformation 148]
173. ~select_1(X0) | from_5(X0) [cnf transformation 149]
174. ~from_5(X0) | select_1(X0) [cnf transformation 149]
175. ~main_query_1(X0) | select_1(X0) [cnf transformation 150]
176. ~select_1(X0) | main_query_1(X0) [cnf transformation 150]
177. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 151]
178. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 151]
181. ~gte(X0,X0) | where_8(X0) | ~lte(X0,X0) [cnf transformation 153]
182. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 157]
184. ~from_7(X0,X1,X2,X3,X4) | ~where_8(X0) | select_6(X0) [cnf transformation 157]
185. ~main_query_2(X0) | select_6(X0) [cnf transformation 158]
186. ~select_6(X0) | main_query_2(X0) [cnf transformation 158]
190. lte(X0,X0) [cnf transformation 59]
196. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 161]
220. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 164]
221. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 164]
246. 0 <=> main_query_1(sK8) [avatar definition]
247. main_query_1(sK8) <- {0} [avatar component clause 246]
249. 1 <=> ~main_query_1(sK8) [avatar definition]
250. ~main_query_1(sK8) <- {1} [avatar component clause 249]
252. 2 <=> main_query_2(sK8) [avatar definition]
253. main_query_2(sK8) <- {2} [avatar component clause 252]
255. 3 <=> ~main_query_2(sK8) [avatar definition]
256. ~main_query_2(sK8) <- {3} [avatar component clause 255]
257. ~1 | ~3 [avatar split clause 221,255,249]
258. 0 | 2 [avatar split clause 220,252,246]
259. select_6(sK8) <- {2} [resolution 185,253]
267. gte(X1,X1) [resolution 196,190]
269. where_8(X0) | ~lte(X0,X0) [resolution 181,267]
271. where_8(X0) [subsumption resolution 269,190]
277. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [resolution 167,165]
279. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 182,177]
281. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [resolution 277,178]
284. from_4(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 279,166]
286. ~select_3(X1) | ~where_8(X1) | select_6(X1) [resolution 281,184]
287. ~select_3(X1) | select_6(X1) [subsumption resolution 286,271]
289. ~select_6(X1) | select_3(X1) [resolution 284,168]
290. select_3(sK8) <- {2} [resolution 289,259]
292. query_2(sK8) <- {2} [resolution 290,170]
293. from_5(sK8) <- {2} [resolution 292,172]
295. select_1(sK8) <- {2} [resolution 293,174]
297. main_query_1(sK8) <- {2} [resolution 295,176]
299. $false <- {1, 2} [subsumption resolution 297,250]
300. 1 | ~2 [avatar contradiction clause 299]
301. select_1(sK8) <- {0} [resolution 247,175]
303. from_5(sK8) <- {0} [resolution 301,173]
305. query_2(sK8) <- {0} [resolution 303,171]
307. select_3(sK8) <- {0} [resolution 305,169]
308. select_6(sK8) <- {0} [resolution 307,287]
311. main_query_2(sK8) <- {0} [resolution 308,186]
313. $false <- {0, 3} [subsumption resolution 311,256]
314. ~0 | 3 [avatar contradiction clause 313]
315. $false [avatar sat refutation 257,258,300,314]
% SZS output end Proof for query06_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.012 s
% ------------------------------
% ------------------------------
