% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_05_query_06_tptp4X
% SZS output start Proof for query_05_query_06_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X5,X1,X2,X3,X4] : (from_5(X5,X1,X2,X3,X4) <=> tab1(X5,X1,X2,X3,X4)) [input]
6. ! [X5] : (where_6(X5) <=> eq(X5,X5)) [input]
7. ! [X5] : (select_4(X5) <=> ? [X1,X2,X3,X4] : (where_6(X5) & from_5(X5,X1,X2,X3,X4))) [input]
8. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
40. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
41. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 40]
42. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
43. ! [X0] : (where_6(X0) <=> eq(X0,X0)) [rectify 6]
44. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
45. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
46. ! [X0] : eq(X0,X0) [rectify 9]
122. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 41]
123. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
124. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
125. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
126. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 125]
127. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
128. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 126,127]
129. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
130. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 42]
131. ! [X0] : ((where_6(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_6(X0))) [nnf transformation 43]
132. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 44]
133. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 132]
134. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
135. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 133,134]
136. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 45]
139. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 122]
140. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
141. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 139,140]
142. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 123]
143. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 123]
144. eq(X0,X0) | ~where_3(X0) [cnf transformation 124]
145. ~eq(X0,X0) | where_3(X0) [cnf transformation 124]
146. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 128]
148. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 128]
149. ~main_query_1(X0) | select_1(X0) [cnf transformation 129]
150. ~select_1(X0) | main_query_1(X0) [cnf transformation 129]
151. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 130]
152. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 130]
154. ~eq(X0,X0) | where_6(X0) [cnf transformation 131]
155. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 135]
157. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 135]
158. ~main_query_2(X0) | select_4(X0) [cnf transformation 136]
159. ~select_4(X0) | main_query_2(X0) [cnf transformation 136]
160. eq(X0,X0) [cnf transformation 46]
189. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 141]
190. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 141]
192. 0 <=> main_query_1(sK8) [avatar definition]
193. main_query_1(sK8) <- {0} [avatar component clause 192]
195. 1 <=> ~main_query_1(sK8) [avatar definition]
196. ~main_query_1(sK8) <- {1} [avatar component clause 195]
198. 2 <=> main_query_2(sK8) [avatar definition]
199. main_query_2(sK8) <- {2} [avatar component clause 198]
201. 3 <=> ~main_query_2(sK8) [avatar definition]
202. ~main_query_2(sK8) <- {3} [avatar component clause 201]
203. ~1 | ~3 [avatar split clause 190,201,195]
204. 0 | 2 [avatar split clause 189,198,192]
205. select_4(sK8) <- {2} [resolution 158,199]
209. where_3(X1) [resolution 145,160]
212. where_6(X1) | ~where_3(X1) [resolution 154,144]
214. where_6(X1) [subsumption resolution 212,209]
252. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 146,142]
254. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 155,151]
257. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 252,152]
261. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 254,143]
264. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 257,157]
265. ~select_1(X1) | select_4(X1) [subsumption resolution 264,214]
269. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 261,148]
275. ~select_4(X1) | select_1(X1) [subsumption resolution 269,209]
278. select_1(sK8) <- {2} [resolution 275,205]
282. main_query_1(sK8) <- {2} [resolution 278,150]
284. $false <- {1, 2} [subsumption resolution 282,196]
285. 1 | ~2 [avatar contradiction clause 284]
286. select_1(sK8) <- {0} [resolution 193,149]
289. select_4(sK8) <- {0} [resolution 286,265]
295. main_query_2(sK8) <- {0} [resolution 289,159]
297. $false <- {0, 3} [subsumption resolution 295,202]
298. ~0 | 3 [avatar contradiction clause 297]
299. $false [avatar sat refutation 203,204,285,298]
% SZS output end Proof for query_05_query_06_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.014 s
% ------------------------------
% ------------------------------
