% Refutation found. Thanks to Tanya!
% SZS status Theorem for query11_query14_tptp4X
% SZS output start Proof for query11_query14_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8,X9,X10,X11] : (from_9(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
7. ! [X7] : (select_8(X7) <=> ? [X8,X9,X10,X11] : ($true & from_9(X7,X8,X9,X10,X11))) [input]
8. ! [X5] : (query_7(X5) <=> select_8(X5)) [input]
9. ! [X7,X8] : (where_6(X7,X8) <=> (query_7(X7) & eq(X7,X8))) [input]
10. ! [X7,X8] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X7,X8) & from_5(X7,X8,X9,X10,X11))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
64. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
65. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 64]
66. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
67. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
68. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 6]
69. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [rectify 7]
70. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 69]
71. ! [X0] : (query_7(X0) <=> select_8(X0)) [rectify 8]
72. ! [X0,X1] : (where_6(X0,X1) <=> (query_7(X0) & eq(X0,X1))) [rectify 9]
73. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 10]
74. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
127. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 65]
217. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 127]
218. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
219. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
220. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
221. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 220]
222. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
223. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 221,222]
224. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 66]
225. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 67]
226. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 68]
227. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 70]
228. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 227]
229. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
230. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 228,229]
231. ! [X0] : ((query_7(X0) | ~select_8(X0)) & (select_8(X0) | ~query_7(X0))) [nnf transformation 71]
232. ! [X0,X1] : ((where_6(X0,X1) | (~query_7(X0) | ~eq(X0,X1))) & ((query_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 72]
233. ! [X0,X1] : ((where_6(X0,X1) | ~query_7(X0) | ~eq(X0,X1)) & ((query_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 232]
234. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 73]
235. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 234]
236. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
237. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 235,236]
238. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 74]
241. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 217]
242. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
243. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 241,242]
244. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 218]
245. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 218]
246. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 219]
247. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 219]
248. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 223]
249. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 223]
250. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 223]
251. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 224]
252. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 224]
253. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 225]
254. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 225]
256. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 226]
258. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 230]
260. ~select_8(X0) | query_7(X0) [cnf transformation 231]
261. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 233]
263. ~eq(X0,X1) | ~query_7(X0) | where_6(X0,X1) [cnf transformation 233]
264. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 237]
265. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 237]
266. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 237]
267. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 238]
268. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 238]
319. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 243]
320. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 243]
364. 0 <=> main_query_1(sK10,sK11) [avatar definition]
365. main_query_1(sK10,sK11) <- {0} [avatar component clause 364]
367. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
368. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 367]
370. 2 <=> main_query_2(sK10,sK11) [avatar definition]
371. main_query_2(sK10,sK11) <- {2} [avatar component clause 370]
373. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
374. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 373]
375. ~1 | ~3 [avatar split clause 320,373,367]
376. 0 | 2 [avatar split clause 319,370,364]
379. select_4(sK10,sK11) <- {2} [resolution 267,371]
380. where_6(sK10,sK11) <- {2} [resolution 379,265]
382. eq(sK10,sK11) <- {2} [resolution 380,261]
392. where_3(sK10,sK11) <- {2} [resolution 382,247]
440. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 248,244]
444. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 264,253]
445. from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 440,256]
446. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 440,254]
452. from_2(X4,X5,sK7(X4,X5),sK8(X4,X5),sK9(X4,X5)) | ~select_4(X4,X5) [resolution 444,245]
454. ~select_1(X2,X3) | select_8(X2) [resolution 445,258]
455. ~where_6(X0,X1) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 446,266]
462. ~select_4(X0,X1) | ~where_3(X0,X1) | select_1(X0,X1) [resolution 452,250]
464. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 462,379]
465. select_1(sK10,sK11) <- {2} [subsumption resolution 464,392]
467. main_query_1(sK10,sK11) <- {2} [resolution 465,252]
469. $false <- {1, 2} [subsumption resolution 467,368]
470. 1 | ~2 [avatar contradiction clause 469]
471. select_1(sK10,sK11) <- {0} [resolution 365,251]
472. select_8(sK10) <- {0} [resolution 471,454]
474. where_3(sK10,sK11) <- {0} [resolution 471,249]
475. query_7(sK10) <- {0} [resolution 472,260]
476. eq(sK10,sK11) <- {0} [resolution 474,246]
479. ~query_7(sK10) | where_6(sK10,sK11) <- {0} [resolution 476,263]
484. where_6(sK10,sK11) <- {0} [subsumption resolution 479,475]
495. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {0} [resolution 484,455]
498. select_4(sK10,sK11) <- {0} [subsumption resolution 495,471]
504. main_query_2(sK10,sK11) <- {0} [resolution 498,268]
506. $false <- {0, 3} [subsumption resolution 504,374]
507. ~0 | 3 [avatar contradiction clause 506]
508. $false [avatar sat refutation 375,376,470,507]
% SZS output end Proof for query11_query14_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
