% Refutation found. Thanks to Tanya!
% SZS status Theorem for query05_query07_tptp4X
% SZS output start Proof for query05_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X6,X1,X2,X3,X4] : (from_5(X6,X1,X2,X3,X4) <=> tab1(X6,X1,X2,X3,X4)) [input]
6. ! [X6] : (where_6(X6) <=> (~gt(X6,X6) & ~lt(X6,X6))) [input]
7. ! [X6] : (select_4(X6) <=> ? [X1,X2,X3,X4] : (where_6(X6) & from_5(X6,X1,X2,X3,X4))) [input]
8. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
16. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
17. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
43. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
44. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 43]
45. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
46. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
47. ! [X0] : (where_6(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [rectify 6]
48. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
49. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
50. ! [X0] : eq(X0,X0) [rectify 9]
57. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 16]
58. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 17]
84. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 44]
136. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 84]
137. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
138. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
139. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
140. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 139]
141. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
142. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 140,141]
143. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 45]
144. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 46]
145. ! [X0] : ((where_6(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [nnf transformation 47]
146. ! [X0] : ((where_6(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_6(X0))) [flattening 145]
147. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 48]
148. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 147]
149. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
150. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 148,149]
151. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 49]
154. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 57]
155. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 154]
156. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 58]
157. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 136]
158. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
159. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 157,158]
160. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 137]
161. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 137]
163. ~eq(X0,X0) | where_3(X0) [cnf transformation 138]
164. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 142]
166. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 142]
167. ~main_query_1(X0) | select_1(X0) [cnf transformation 143]
168. ~select_1(X0) | main_query_1(X0) [cnf transformation 143]
169. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 144]
170. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 144]
173. gt(X0,X0) | where_6(X0) | lt(X0,X0) [cnf transformation 146]
174. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 150]
176. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 150]
177. ~main_query_2(X0) | select_4(X0) [cnf transformation 151]
178. ~select_4(X0) | main_query_2(X0) [cnf transformation 151]
179. eq(X0,X0) [cnf transformation 50]
189. ~lt(X0,X1) | ~eq(X0,X1) [cnf transformation 155]
191. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 156]
216. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 159]
217. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 159]
219. 0 <=> main_query_1(sK8) [avatar definition]
220. main_query_1(sK8) <- {0} [avatar component clause 219]
222. 1 <=> ~main_query_1(sK8) [avatar definition]
223. ~main_query_1(sK8) <- {1} [avatar component clause 222]
225. 2 <=> main_query_2(sK8) [avatar definition]
226. main_query_2(sK8) <- {2} [avatar component clause 225]
228. 3 <=> ~main_query_2(sK8) [avatar definition]
229. ~main_query_2(sK8) <- {3} [avatar component clause 228]
230. ~1 | ~3 [avatar split clause 217,228,222]
231. 0 | 2 [avatar split clause 216,225,219]
232. select_4(sK8) <- {2} [resolution 177,226]
236. where_3(X1) [resolution 163,179]
245. where_6(X1) | lt(X1,X1) | lt(X1,X1) [resolution 173,191]
246. lt(X1,X1) | where_6(X1) [duplicate literal removal 245]
249. where_6(X2) | ~eq(X2,X2) [resolution 246,189]
251. where_6(X2) [subsumption resolution 249,179]
324. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 164,160]
355. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 174,169]
449. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 324,170]
453. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 355,161]
456. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 449,176]
462. ~select_1(X1) | select_4(X1) [subsumption resolution 456,251]
467. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 453,166]
473. ~select_4(X1) | select_1(X1) [subsumption resolution 467,236]
476. select_1(sK8) <- {2} [resolution 473,232]
480. main_query_1(sK8) <- {2} [resolution 476,168]
482. $false <- {1, 2} [subsumption resolution 480,223]
483. 1 | ~2 [avatar contradiction clause 482]
484. select_1(sK8) <- {0} [resolution 220,167]
487. select_4(sK8) <- {0} [resolution 484,462]
493. main_query_2(sK8) <- {0} [resolution 487,178]
495. $false <- {0, 3} [subsumption resolution 493,229]
496. ~0 | 3 [avatar contradiction clause 495]
497. $false [avatar sat refutation 230,231,483,496]
% SZS output end Proof for query05_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
