% Refutation found. Thanks to Tanya!
% SZS status Theorem for query01_query09_tptp4X
% SZS output start Proof for query01_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X5,X6,X7,X8] : (from_7(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
5. ! [X0] : (select_6(X0) <=> ? [X5,X6,X7,X8] : ($true & from_7(X0,X5,X6,X7,X8))) [input]
6. ! [X0] : (subquery_from_5(X0) <=> select_6(X0)) [input]
7. ! [X9] : (from_4(X9) <=> subquery_from_5(X9)) [input]
8. ! [X9] : (select_3(X9) <=> ($true & from_4(X9))) [input]
9. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
41. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
42. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 41]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
44. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
45. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [rectify 5]
46. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 45]
47. ! [X0] : (from_4(X0) <=> subquery_from_5(X0)) [rectify 7]
48. ! [X0] : (select_3(X0) <=> ($true & from_4(X0))) [rectify 8]
49. ! [X0] : (select_3(X0) <=> from_4(X0)) [true and false elimination 48]
126. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 42]
127. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
128. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 43]
129. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 128]
130. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
131. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 129,130]
132. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
133. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 44]
134. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 46]
135. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 134]
136. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
137. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 135,136]
138. ! [X0] : ((subquery_from_5(X0) | ~select_6(X0)) & (select_6(X0) | ~subquery_from_5(X0))) [nnf transformation 6]
139. ! [X0] : ((from_4(X0) | ~subquery_from_5(X0)) & (subquery_from_5(X0) | ~from_4(X0))) [nnf transformation 47]
140. ! [X0] : ((select_3(X0) | ~from_4(X0)) & (from_4(X0) | ~select_3(X0))) [nnf transformation 49]
141. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 9]
144. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 126]
145. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
146. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 144,145]
147. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 127]
148. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 127]
149. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 131]
150. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 131]
151. ~main_query_1(X0) | select_1(X0) [cnf transformation 132]
152. ~select_1(X0) | main_query_1(X0) [cnf transformation 132]
153. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 133]
154. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 133]
155. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 137]
156. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 137]
157. ~subquery_from_5(X0) | select_6(X0) [cnf transformation 138]
158. ~select_6(X0) | subquery_from_5(X0) [cnf transformation 138]
159. ~from_4(X0) | subquery_from_5(X0) [cnf transformation 139]
160. ~subquery_from_5(X0) | from_4(X0) [cnf transformation 139]
161. ~select_3(X0) | from_4(X0) [cnf transformation 140]
162. ~from_4(X0) | select_3(X0) [cnf transformation 140]
163. ~main_query_2(X0) | select_3(X0) [cnf transformation 141]
164. ~select_3(X0) | main_query_2(X0) [cnf transformation 141]
194. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 146]
195. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 146]
197. 0 <=> main_query_1(sK8) [avatar definition]
198. main_query_1(sK8) <- {0} [avatar component clause 197]
200. 1 <=> ~main_query_1(sK8) [avatar definition]
201. ~main_query_1(sK8) <- {1} [avatar component clause 200]
203. 2 <=> main_query_2(sK8) [avatar definition]
204. main_query_2(sK8) <- {2} [avatar component clause 203]
206. 3 <=> ~main_query_2(sK8) [avatar definition]
207. ~main_query_2(sK8) <- {3} [avatar component clause 206]
208. ~1 | ~3 [avatar split clause 195,206,200]
209. 0 | 2 [avatar split clause 194,203,197]
210. select_3(sK8) <- {2} [resolution 163,204]
212. from_4(sK8) <- {2} [resolution 210,161]
215. subquery_from_5(sK8) <- {2} [resolution 212,159]
218. select_6(sK8) <- {2} [resolution 215,157]
247. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 149,147]
249. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 155,153]
252. from_7(X2,sK0(X2),sK1(X2),sK2(X2),sK3(X2)) | ~select_1(X2) [resolution 247,154]
258. from_2(X5,sK4(X5),sK5(X5),sK6(X5),sK7(X5)) | ~select_6(X5) [resolution 249,148]
262. ~select_1(X1) | select_6(X1) [resolution 252,156]
268. ~select_6(X1) | select_1(X1) [resolution 258,150]
271. select_1(sK8) <- {2} [resolution 268,218]
275. main_query_1(sK8) <- {2} [resolution 271,152]
276. $false <- {1, 2} [subsumption resolution 275,201]
277. 1 | ~2 [avatar contradiction clause 276]
278. select_1(sK8) <- {0} [resolution 198,151]
281. select_6(sK8) <- {0} [resolution 278,262]
286. subquery_from_5(sK8) <- {0} [resolution 281,158]
289. from_4(sK8) <- {0} [resolution 286,160]
293. select_3(sK8) <- {0} [resolution 289,162]
297. main_query_2(sK8) <- {0} [resolution 293,164]
299. $false <- {0, 3} [subsumption resolution 297,207]
300. ~0 | 3 [avatar contradiction clause 299]
301. $false [avatar sat refutation 208,209,277,300]
% SZS output end Proof for query01_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
