% Refutation found. Thanks to Tanya!
% SZS status Theorem for query01_query07_tptp4X
% SZS output start Proof for query01_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
6. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
8. ! [X5] : eq(X5,X5) [input]
15. ! [X5,X6] : (lt(X5,X6) <=> (~eq(X5,X6) & lte(X5,X6))) [input]
16. ! [X5,X6] : (gt(X5,X6) <=> lt(X6,X5)) [input]
39. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
40. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 39]
41. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
42. ! [X0] : eq(X0,X0) [rectify 8]
49. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 15]
50. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 16]
118. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 40]
119. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
120. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 41]
121. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 120]
122. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
123. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 121,122]
124. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
125. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
126. ! [X0] : ((where_5(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [nnf transformation 5]
127. ! [X0] : ((where_5(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [flattening 126]
128. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 6]
129. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 128]
130. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
131. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 129,130]
132. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 7]
135. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 49]
136. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 135]
137. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 50]
138. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 118]
139. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
140. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 138,139]
141. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 119]
142. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 119]
143. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 123]
144. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 123]
145. ~main_query_1(X0) | select_1(X0) [cnf transformation 124]
146. ~select_1(X0) | main_query_1(X0) [cnf transformation 124]
147. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 125]
148. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 125]
151. gt(X0,X0) | where_5(X0) | lt(X0,X0) [cnf transformation 127]
152. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 131]
154. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 131]
155. ~main_query_2(X0) | select_3(X0) [cnf transformation 132]
156. ~select_3(X0) | main_query_2(X0) [cnf transformation 132]
157. eq(X0,X0) [cnf transformation 42]
167. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 136]
169. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 137]
191. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 140]
192. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 140]
194. 0 <=> main_query_1(sK8) [avatar definition]
195. main_query_1(sK8) <- {0} [avatar component clause 194]
197. 1 <=> ~main_query_1(sK8) [avatar definition]
198. ~main_query_1(sK8) <- {1} [avatar component clause 197]
200. 2 <=> main_query_2(sK8) [avatar definition]
201. main_query_2(sK8) <- {2} [avatar component clause 200]
203. 3 <=> ~main_query_2(sK8) [avatar definition]
204. ~main_query_2(sK8) <- {3} [avatar component clause 203]
205. ~1 | ~3 [avatar split clause 192,203,197]
206. 0 | 2 [avatar split clause 191,200,194]
207. select_3(sK8) <- {2} [resolution 155,201]
214. ~lt(X0,X0) [resolution 167,157]
216. where_5(X1) | lt(X1,X1) | lt(X1,X1) [resolution 151,169]
217. where_5(X1) | lt(X1,X1) [duplicate literal removal 216]
218. where_5(X1) [subsumption resolution 217,214]
242. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 143,141]
284. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 152,147]
408. from_4(X2,sK0(X2),sK1(X2),sK2(X2),sK3(X2)) | ~select_1(X2) [resolution 242,148]
414. from_2(X5,sK4(X5),sK5(X5),sK6(X5),sK7(X5)) | ~select_3(X5) [resolution 284,142]
418. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 408,154]
424. ~select_1(X1) | select_3(X1) [subsumption resolution 418,218]
430. ~select_3(X1) | select_1(X1) [resolution 414,144]
436. select_1(sK8) <- {2} [resolution 430,207]
438. main_query_1(sK8) <- {2} [resolution 436,146]
439. $false <- {1, 2} [subsumption resolution 438,198]
440. 1 | ~2 [avatar contradiction clause 439]
441. select_1(sK8) <- {0} [resolution 195,145]
442. select_3(sK8) <- {0} [resolution 441,424]
445. main_query_2(sK8) <- {0} [resolution 442,156]
447. $false <- {0, 3} [subsumption resolution 445,204]
448. ~0 | 3 [avatar contradiction clause 447]
449. $false [avatar sat refutation 205,206,440,448]
% SZS output end Proof for query01_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
