% Refutation found. Thanks to Tanya!
% SZS status Theorem for query04_query05_tptp4X
% SZS output start Proof for query04_query05_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> eq(X0,X0)) [input]
7. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
42. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
43. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 42]
44. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
45. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
46. ! [X0] : eq(X0,X0) [rectify 9]
79. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 43]
131. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 79]
132. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
133. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 134]
136. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 135,136]
138. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 44]
139. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
140. ! [X0] : ((where_6(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_6(X0))) [nnf transformation 6]
141. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 7]
142. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 141]
143. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
144. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 142,143]
145. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 45]
148. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 131]
149. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
150. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 148,149]
151. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 132]
152. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 132]
153. eq(X0,X0) | ~where_3(X0) [cnf transformation 133]
154. ~eq(X0,X0) | where_3(X0) [cnf transformation 133]
155. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 137]
157. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 137]
158. ~main_query_1(X0) | select_1(X0) [cnf transformation 138]
159. ~select_1(X0) | main_query_1(X0) [cnf transformation 138]
160. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 139]
161. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 139]
163. ~eq(X0,X0) | where_6(X0) [cnf transformation 140]
164. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 144]
166. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 144]
167. ~main_query_2(X0) | select_4(X0) [cnf transformation 145]
168. ~select_4(X0) | main_query_2(X0) [cnf transformation 145]
169. eq(X0,X0) [cnf transformation 46]
200. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 150]
201. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 150]
226. 0 <=> main_query_1(sK8) [avatar definition]
227. main_query_1(sK8) <- {0} [avatar component clause 226]
229. 1 <=> ~main_query_1(sK8) [avatar definition]
230. ~main_query_1(sK8) <- {1} [avatar component clause 229]
232. 2 <=> main_query_2(sK8) [avatar definition]
233. main_query_2(sK8) <- {2} [avatar component clause 232]
235. 3 <=> ~main_query_2(sK8) [avatar definition]
236. ~main_query_2(sK8) <- {3} [avatar component clause 235]
237. ~1 | ~3 [avatar split clause 201,235,229]
238. 0 | 2 [avatar split clause 200,232,226]
239. select_4(sK8) <- {2} [resolution 167,233]
243. where_3(X1) [resolution 154,169]
246. where_6(X1) | ~where_3(X1) [resolution 163,153]
248. where_6(X1) [subsumption resolution 246,243]
263. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 155,151]
265. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 164,160]
267. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 263,161]
270. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 265,152]
272. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 267,166]
273. ~select_1(X1) | select_4(X1) [subsumption resolution 272,248]
275. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 270,157]
276. ~select_4(X1) | select_1(X1) [subsumption resolution 275,243]
277. select_1(sK8) <- {2} [resolution 276,239]
279. main_query_1(sK8) <- {2} [resolution 277,159]
281. $false <- {1, 2} [subsumption resolution 279,230]
282. 1 | ~2 [avatar contradiction clause 281]
283. select_1(sK8) <- {0} [resolution 227,158]
284. select_4(sK8) <- {0} [resolution 283,273]
288. main_query_2(sK8) <- {0} [resolution 284,168]
290. $false <- {0, 3} [subsumption resolution 288,236]
291. ~0 | 3 [avatar contradiction clause 290]
292. $false [avatar sat refutation 237,238,282,291]
% SZS output end Proof for query04_query05_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.010 s
% ------------------------------
% ------------------------------
