% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_04_query_07_tptp4X
% SZS output start Proof for query_04_query_07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [input]
4. ! [X5,X6,X7,X8,X9] : (from_4(X5,X6,X7,X8,X9) <=> tab1(X5,X6,X7,X8,X9)) [input]
5. ! [X5] : (where_5(X5) <=> (~gt(X5,X5) & ~lt(X5,X5))) [input]
6. ! [X5,X6] : (select_3(X5,X6) <=> ? [X7,X8,X9] : (where_5(X5) & from_4(X5,X6,X7,X8,X9))) [input]
7. ! [X10,X11] : (main_query_2(X10,X11) <=> select_3(X10,X11)) [input]
8. ! [X12] : eq(X12,X12) [input]
14. ! [X12,X13] : (lt(X12,X13) <=> (~eq(X12,X13) & lte(X12,X13))) [input]
15. ! [X12,X13] : (gt(X12,X13) <=> lt(X13,X12)) [input]
46. ! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [input]
47. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [negated conjecture 46]
48. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
49. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
50. ! [X0] : (where_5(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [rectify 5]
51. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
52. ! [X0,X1] : (main_query_2(X0,X1) <=> select_3(X0,X1)) [rectify 7]
53. ! [X0] : eq(X0,X0) [rectify 8]
59. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 14]
60. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 15]
152. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 47]
153. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
154. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0,X1))) [nnf transformation 48]
155. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) | ~select_1(X0,X1))) [rectify 154]
156. ! [X1,X0] : (? [X5,X6,X7] : from_2(X0,X1,X5,X6,X7) => from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) [choice axiom]
157. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1))) [skolemisation 155,156]
158. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 3]
159. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 49]
160. ! [X0] : ((where_5(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [nnf transformation 50]
161. ! [X0] : ((where_5(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [flattening 160]
162. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0,X1))) [nnf transformation 51]
163. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) | ~select_3(X0,X1))) [rectify 162]
164. ! [X1,X0] : (? [X5,X6,X7] : (where_5(X0) & from_4(X0,X1,X5,X6,X7)) => (where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
165. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_3(X0,X1))) [skolemisation 163,164]
166. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 52]
169. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 59]
170. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 169]
171. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 60]
172. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 152]
173. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
174. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 172,173]
175. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 153]
176. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 153]
177. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 157]
178. ~from_2(X0,X1,X2,X3,X4) | select_1(X0,X1) [cnf transformation 157]
179. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 158]
180. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 158]
181. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 159]
182. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 159]
185. gt(X0,X0) | where_5(X0) | lt(X0,X0) [cnf transformation 161]
186. from_4(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [cnf transformation 165]
188. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0,X1) [cnf transformation 165]
189. ~main_query_2(X0,X1) | select_3(X0,X1) [cnf transformation 166]
190. ~select_3(X0,X1) | main_query_2(X0,X1) [cnf transformation 166]
191. eq(X0,X0) [cnf transformation 53]
200. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 170]
202. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 171]
232. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 174]
233. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 174]
235. 0 <=> main_query_1(sK6,sK7) [avatar definition]
236. main_query_1(sK6,sK7) <- {0} [avatar component clause 235]
238. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
239. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 238]
241. 2 <=> main_query_2(sK6,sK7) [avatar definition]
242. main_query_2(sK6,sK7) <- {2} [avatar component clause 241]
244. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
245. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 244]
246. ~1 | ~3 [avatar split clause 233,244,238]
247. 0 | 2 [avatar split clause 232,241,235]
248. select_3(sK6,sK7) <- {2} [resolution 189,242]
254. ~lt(X0,X0) [resolution 200,191]
256. where_5(X1) | lt(X1,X1) | lt(X1,X1) [resolution 185,202]
257. where_5(X1) | lt(X1,X1) [duplicate literal removal 256]
258. where_5(X1) [subsumption resolution 257,254]
271. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 177,175]
273. tab1(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 186,181]
286. from_4(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 271,182]
292. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_3(X0,X1) [resolution 273,176]
307. ~select_1(X2,X3) | ~where_5(X2) | select_3(X2,X3) [resolution 286,188]
313. ~select_1(X2,X3) | select_3(X2,X3) [subsumption resolution 307,258]
329. ~select_3(X2,X3) | select_1(X2,X3) [resolution 292,178]
342. select_1(sK6,sK7) <- {2} [resolution 329,248]
351. main_query_1(sK6,sK7) <- {2} [resolution 342,180]
352. $false <- {1, 2} [subsumption resolution 351,239]
353. 1 | ~2 [avatar contradiction clause 352]
354. select_1(sK6,sK7) <- {0} [resolution 236,179]
355. select_3(sK6,sK7) <- {0} [resolution 354,313]
358. main_query_2(sK6,sK7) <- {0} [resolution 355,190]
360. $false <- {0, 3} [subsumption resolution 358,245]
361. ~0 | 3 [avatar contradiction clause 360]
362. $false [avatar sat refutation 246,247,353,361]
% SZS output end Proof for query_04_query_07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
