% Refutation found. Thanks to Tanya!
% SZS status Theorem for query00_query09_tptp4X
% SZS output start Proof for query00_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_6(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (select_5(X6) <=> ? [X7,X8,X9,X10] : ($true & from_6(X6,X7,X8,X9,X10))) [input]
6. ! [X11] : (query_4(X11) <=> select_5(X11)) [input]
7. ! [X12] : (from_7(X12) <=> query_4(X12)) [input]
8. ! [X12] : (select_3(X12) <=> ($true & from_7(X12))) [input]
9. ! [X11] : (main_query_2(X11) <=> select_3(X11)) [input]
43. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
44. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 43]
45. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
46. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
47. ! [X0,X1,X2,X3,X4] : (from_6(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
48. ! [X0] : (select_5(X0) <=> ? [X1,X2,X3,X4] : ($true & from_6(X0,X1,X2,X3,X4))) [rectify 5]
49. ! [X0] : (select_5(X0) <=> ? [X1,X2,X3,X4] : from_6(X0,X1,X2,X3,X4)) [true and false elimination 48]
50. ! [X0] : (query_4(X0) <=> select_5(X0)) [rectify 6]
51. ! [X0] : (from_7(X0) <=> query_4(X0)) [rectify 7]
52. ! [X0] : (select_3(X0) <=> ($true & from_7(X0))) [rectify 8]
53. ! [X0] : (select_3(X0) <=> from_7(X0)) [true and false elimination 52]
54. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 9]
88. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 44]
140. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 88]
141. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
142. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 45]
143. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 142]
144. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
145. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 143,144]
146. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 46]
147. ! [X0,X1,X2,X3,X4] : ((from_6(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_6(X0,X1,X2,X3,X4))) [nnf transformation 47]
148. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_6(X0,X1,X2,X3,X4) | ~select_5(X0))) [nnf transformation 49]
149. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_6(X0,X5,X6,X7,X8) | ~select_5(X0))) [rectify 148]
150. ! [X0] : (? [X5,X6,X7,X8] : from_6(X0,X5,X6,X7,X8) => from_6(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
151. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (from_6(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_5(X0))) [skolemisation 149,150]
152. ! [X0] : ((query_4(X0) | ~select_5(X0)) & (select_5(X0) | ~query_4(X0))) [nnf transformation 50]
153. ! [X0] : ((from_7(X0) | ~query_4(X0)) & (query_4(X0) | ~from_7(X0))) [nnf transformation 51]
154. ! [X0] : ((select_3(X0) | ~from_7(X0)) & (from_7(X0) | ~select_3(X0))) [nnf transformation 53]
155. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 54]
158. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 140]
159. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
160. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 158,159]
161. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 141]
162. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 141]
163. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 145]
164. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 145]
165. ~main_query_1(X0) | select_1(X0) [cnf transformation 146]
166. ~select_1(X0) | main_query_1(X0) [cnf transformation 146]
167. ~from_6(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 147]
168. ~tab1(X0,X1,X2,X3,X4) | from_6(X0,X1,X2,X3,X4) [cnf transformation 147]
169. from_6(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_5(X0) [cnf transformation 151]
170. ~from_6(X0,X1,X2,X3,X4) | select_5(X0) [cnf transformation 151]
171. ~query_4(X0) | select_5(X0) [cnf transformation 152]
172. ~select_5(X0) | query_4(X0) [cnf transformation 152]
173. ~from_7(X0) | query_4(X0) [cnf transformation 153]
174. ~query_4(X0) | from_7(X0) [cnf transformation 153]
175. ~select_3(X0) | from_7(X0) [cnf transformation 154]
176. ~from_7(X0) | select_3(X0) [cnf transformation 154]
177. ~main_query_2(X0) | select_3(X0) [cnf transformation 155]
178. ~select_3(X0) | main_query_2(X0) [cnf transformation 155]
210. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 160]
211. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 160]
236. 0 <=> main_query_1(sK8) [avatar definition]
237. main_query_1(sK8) <- {0} [avatar component clause 236]
239. 1 <=> ~main_query_1(sK8) [avatar definition]
240. ~main_query_1(sK8) <- {1} [avatar component clause 239]
242. 2 <=> main_query_2(sK8) [avatar definition]
243. main_query_2(sK8) <- {2} [avatar component clause 242]
245. 3 <=> ~main_query_2(sK8) [avatar definition]
246. ~main_query_2(sK8) <- {3} [avatar component clause 245]
247. ~1 | ~3 [avatar split clause 211,245,239]
248. 0 | 2 [avatar split clause 210,242,236]
249. select_3(sK8) <- {2} [resolution 177,243]
251. from_7(sK8) <- {2} [resolution 249,175]
254. query_4(sK8) <- {2} [resolution 251,173]
257. select_5(sK8) <- {2} [resolution 254,171]
263. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 163,161]
265. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_5(X0) [resolution 169,167]
267. from_6(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 263,168]
270. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_5(X1) [resolution 265,162]
272. ~select_1(X1) | select_5(X1) [resolution 267,170]
274. ~select_5(X1) | select_1(X1) [resolution 270,164]
275. select_1(sK8) <- {2} [resolution 274,257]
277. main_query_1(sK8) <- {2} [resolution 275,166]
278. $false <- {1, 2} [subsumption resolution 277,240]
279. 1 | ~2 [avatar contradiction clause 278]
280. select_1(sK8) <- {0} [resolution 237,165]
281. select_5(sK8) <- {0} [resolution 280,272]
284. query_4(sK8) <- {0} [resolution 281,172]
285. from_7(sK8) <- {0} [resolution 284,174]
287. select_3(sK8) <- {0} [resolution 285,176]
289. main_query_2(sK8) <- {0} [resolution 287,178]
291. $false <- {0, 3} [subsumption resolution 289,246]
292. ~0 | 3 [avatar contradiction clause 291]
293. $false [avatar sat refutation 247,248,279,292]
% SZS output end Proof for query00_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.018 s
% ------------------------------
% ------------------------------
