% Refutation found. Thanks to Tanya!
% SZS status Theorem for query03_query09_tptp4X
% SZS output start Proof for query03_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [input]
6. ! [X0] : (subquery_from_5(X0) <=> select_6(X0)) [input]
7. ! [X0] : (from_4(X0) <=> subquery_from_5(X0)) [input]
8. ! [X0] : (select_3(X0) <=> ($true & from_4(X0))) [input]
9. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
43. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
44. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 43]
45. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
46. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 5]
47. ! [X0] : (select_3(X0) <=> from_4(X0)) [true and false elimination 8]
132. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 44]
133. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 45]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 134]
136. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 135,136]
138. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
139. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 4]
140. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 46]
141. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 140]
142. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
143. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 141,142]
144. ! [X0] : ((subquery_from_5(X0) | ~select_6(X0)) & (select_6(X0) | ~subquery_from_5(X0))) [nnf transformation 6]
145. ! [X0] : ((from_4(X0) | ~subquery_from_5(X0)) & (subquery_from_5(X0) | ~from_4(X0))) [nnf transformation 7]
146. ! [X0] : ((select_3(X0) | ~from_4(X0)) & (from_4(X0) | ~select_3(X0))) [nnf transformation 47]
147. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 9]
150. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 132]
151. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
152. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 150,151]
153. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 133]
154. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 133]
155. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 137]
156. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 137]
157. ~main_query_1(X0) | select_1(X0) [cnf transformation 138]
158. ~select_1(X0) | main_query_1(X0) [cnf transformation 138]
159. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 139]
160. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 139]
161. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 143]
162. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 143]
163. ~subquery_from_5(X0) | select_6(X0) [cnf transformation 144]
164. ~select_6(X0) | subquery_from_5(X0) [cnf transformation 144]
165. ~from_4(X0) | subquery_from_5(X0) [cnf transformation 145]
166. ~subquery_from_5(X0) | from_4(X0) [cnf transformation 145]
167. ~select_3(X0) | from_4(X0) [cnf transformation 146]
168. ~from_4(X0) | select_3(X0) [cnf transformation 146]
169. ~main_query_2(X0) | select_3(X0) [cnf transformation 147]
170. ~select_3(X0) | main_query_2(X0) [cnf transformation 147]
202. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 152]
203. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 152]
228. 0 <=> main_query_1(sK8) [avatar definition]
229. main_query_1(sK8) <- {0} [avatar component clause 228]
231. 1 <=> ~main_query_1(sK8) [avatar definition]
232. ~main_query_1(sK8) <- {1} [avatar component clause 231]
234. 2 <=> main_query_2(sK8) [avatar definition]
235. main_query_2(sK8) <- {2} [avatar component clause 234]
237. 3 <=> ~main_query_2(sK8) [avatar definition]
238. ~main_query_2(sK8) <- {3} [avatar component clause 237]
239. ~1 | ~3 [avatar split clause 203,237,231]
240. 0 | 2 [avatar split clause 202,234,228]
241. select_3(sK8) <- {2} [resolution 169,235]
243. from_4(sK8) <- {2} [resolution 241,167]
246. subquery_from_5(sK8) <- {2} [resolution 243,165]
249. select_6(sK8) <- {2} [resolution 246,163]
255. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 155,153]
257. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 161,159]
259. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 255,160]
262. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 257,154]
264. ~select_1(X1) | select_6(X1) [resolution 259,162]
266. ~select_6(X1) | select_1(X1) [resolution 262,156]
267. select_1(sK8) <- {2} [resolution 266,249]
269. main_query_1(sK8) <- {2} [resolution 267,158]
270. $false <- {1, 2} [subsumption resolution 269,232]
271. 1 | ~2 [avatar contradiction clause 270]
272. select_1(sK8) <- {0} [resolution 229,157]
273. select_6(sK8) <- {0} [resolution 272,264]
276. subquery_from_5(sK8) <- {0} [resolution 273,164]
277. from_4(sK8) <- {0} [resolution 276,166]
279. select_3(sK8) <- {0} [resolution 277,168]
281. main_query_2(sK8) <- {0} [resolution 279,170]
283. $false <- {0, 3} [subsumption resolution 281,238]
284. ~0 | 3 [avatar contradiction clause 283]
285. $false [avatar sat refutation 239,240,271,284]
% SZS output end Proof for query03_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.010 s
% ------------------------------
% ------------------------------
