% Refutation found. Thanks to Tanya!
% SZS status Theorem for query02_query03_tptp4X
% SZS output start Proof for query02_query03_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [input]
6. ! [X6] : (main_query_2(X6) <=> select_3(X6)) [input]
40. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
41. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 40]
42. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
43. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
44. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 5]
45. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 6]
79. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 41]
131. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 79]
132. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
133. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 42]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 133]
135. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
136. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 134,135]
137. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 43]
138. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
139. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 44]
140. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 139]
141. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
142. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0))) [skolemisation 140,141]
143. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 45]
146. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 131]
147. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
148. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 146,147]
149. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 132]
150. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 132]
151. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 136]
152. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 136]
153. ~main_query_1(X0) | select_1(X0) [cnf transformation 137]
154. ~select_1(X0) | main_query_1(X0) [cnf transformation 137]
155. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 138]
156. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 138]
157. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 142]
158. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 142]
159. ~main_query_2(X0) | select_3(X0) [cnf transformation 143]
160. ~select_3(X0) | main_query_2(X0) [cnf transformation 143]
192. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 148]
193. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 148]
218. 0 <=> main_query_1(sK8) [avatar definition]
219. main_query_1(sK8) <- {0} [avatar component clause 218]
221. 1 <=> ~main_query_1(sK8) [avatar definition]
222. ~main_query_1(sK8) <- {1} [avatar component clause 221]
224. 2 <=> main_query_2(sK8) [avatar definition]
225. main_query_2(sK8) <- {2} [avatar component clause 224]
227. 3 <=> ~main_query_2(sK8) [avatar definition]
228. ~main_query_2(sK8) <- {3} [avatar component clause 227]
229. ~1 | ~3 [avatar split clause 193,227,221]
230. 0 | 2 [avatar split clause 192,224,218]
231. select_3(sK8) <- {2} [resolution 159,225]
239. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 151,149]
241. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 157,155]
243. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 239,156]
246. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 241,150]
248. ~select_1(X1) | select_3(X1) [resolution 243,158]
250. ~select_3(X1) | select_1(X1) [resolution 246,152]
251. select_1(sK8) <- {2} [resolution 250,231]
253. main_query_1(sK8) <- {2} [resolution 251,154]
254. $false <- {1, 2} [subsumption resolution 253,222]
255. 1 | ~2 [avatar contradiction clause 254]
256. select_1(sK8) <- {0} [resolution 219,153]
257. select_3(sK8) <- {0} [resolution 256,248]
260. main_query_2(sK8) <- {0} [resolution 257,160]
261. $false <- {0, 3} [subsumption resolution 260,228]
262. ~0 | 3 [avatar contradiction clause 261]
263. $false [avatar sat refutation 229,230,255,262]
% SZS output end Proof for query02_query03_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.015 s
% ------------------------------
% ------------------------------
