% Refutation found. Thanks to Tanya!
% SZS status Theorem for query04_query08_tptp4X
% SZS output start Proof for query04_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0] : (where_6(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
7. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
12. ! [X7] : lte(X7,X7) [input]
18. ! [X7,X8] : (gte(X7,X8) <=> lte(X8,X7)) [input]
42. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
43. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 42]
44. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
45. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
46. ! [X0] : eq(X0,X0) [rectify 9]
49. ! [X0] : lte(X0,X0) [rectify 12]
55. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 18]
79. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 43]
131. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 79]
132. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
133. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 134]
136. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 135,136]
138. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 44]
139. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
140. ! [X0] : ((where_6(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [nnf transformation 6]
141. ! [X0] : ((where_6(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [flattening 140]
142. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 7]
143. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 142]
144. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
145. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 143,144]
146. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 45]
149. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 55]
150. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 131]
151. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
152. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 150,151]
153. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 132]
154. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 132]
156. ~eq(X0,X0) | where_3(X0) [cnf transformation 133]
157. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 137]
159. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 137]
160. ~main_query_1(X0) | select_1(X0) [cnf transformation 138]
161. ~select_1(X0) | main_query_1(X0) [cnf transformation 138]
162. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 139]
163. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 139]
166. ~gte(X0,X0) | where_6(X0) | ~lte(X0,X0) [cnf transformation 141]
167. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 145]
169. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 145]
170. ~main_query_2(X0) | select_4(X0) [cnf transformation 146]
171. ~select_4(X0) | main_query_2(X0) [cnf transformation 146]
172. eq(X0,X0) [cnf transformation 46]
175. lte(X0,X0) [cnf transformation 49]
181. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 149]
205. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 152]
206. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 152]
231. 0 <=> main_query_1(sK8) [avatar definition]
232. main_query_1(sK8) <- {0} [avatar component clause 231]
234. 1 <=> ~main_query_1(sK8) [avatar definition]
235. ~main_query_1(sK8) <- {1} [avatar component clause 234]
237. 2 <=> main_query_2(sK8) [avatar definition]
238. main_query_2(sK8) <- {2} [avatar component clause 237]
240. 3 <=> ~main_query_2(sK8) [avatar definition]
241. ~main_query_2(sK8) <- {3} [avatar component clause 240]
242. ~1 | ~3 [avatar split clause 206,240,234]
243. 0 | 2 [avatar split clause 205,237,231]
244. select_4(sK8) <- {2} [resolution 170,238]
248. where_3(X1) [resolution 156,172]
257. gte(X1,X1) [resolution 181,175]
259. where_6(X0) | ~lte(X0,X0) [resolution 166,257]
261. where_6(X0) [subsumption resolution 259,175]
268. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 157,153]
270. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 167,162]
272. from_5(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 268,163]
275. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_4(X1) [resolution 270,154]
277. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 272,169]
278. ~select_1(X1) | select_4(X1) [subsumption resolution 277,261]
280. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 275,159]
281. ~select_4(X1) | select_1(X1) [subsumption resolution 280,248]
282. select_1(sK8) <- {2} [resolution 281,244]
284. main_query_1(sK8) <- {2} [resolution 282,161]
286. $false <- {1, 2} [subsumption resolution 284,235]
287. 1 | ~2 [avatar contradiction clause 286]
288. select_1(sK8) <- {0} [resolution 232,160]
289. select_4(sK8) <- {0} [resolution 288,278]
293. main_query_2(sK8) <- {0} [resolution 289,171]
295. $false <- {0, 3} [subsumption resolution 293,241]
296. ~0 | 3 [avatar contradiction clause 295]
297. $false [avatar sat refutation 242,243,287,296]
% SZS output end Proof for query04_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.021 s
% ------------------------------
% ------------------------------
