% Refutation found. Thanks to Tanya!
% SZS status Theorem for query07_query08_tptp4X
% SZS output start Proof for query07_query08_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X5,X1,X2,X3,X4] : (from_5(X5,X1,X2,X3,X4) <=> tab1(X5,X1,X2,X3,X4)) [input]
6. ! [X5] : (where_6(X5) <=> (gte(X5,X5) & lte(X5,X5))) [input]
7. ! [X5] : (select_4(X5) <=> ? [X1,X2,X3,X4] : (where_6(X5) & from_5(X5,X1,X2,X3,X4))) [input]
8. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
9. ! [X7] : eq(X7,X7) [input]
12. ! [X7] : lte(X7,X7) [input]
16. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
17. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
19. ! [X7,X8] : (gte(X7,X8) <=> lte(X8,X7)) [input]
43. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
44. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 43]
45. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
46. ! [X0] : (where_6(X0) <=> (gte(X0,X0) & lte(X0,X0))) [rectify 6]
47. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
48. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
49. ! [X0] : eq(X0,X0) [rectify 9]
52. ! [X0] : lte(X0,X0) [rectify 12]
56. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 16]
57. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 17]
59. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 19]
134. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 44]
135. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
136. ! [X0] : ((where_3(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
137. ! [X0] : ((where_3(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_3(X0))) [flattening 136]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
139. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 138]
140. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
141. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 139,140]
142. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
143. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 45]
144. ! [X0] : ((where_6(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [nnf transformation 46]
145. ! [X0] : ((where_6(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_6(X0))) [flattening 144]
146. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0))) [nnf transformation 47]
147. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) | ~select_4(X0))) [rectify 146]
148. ! [X0] : (? [X5,X6,X7,X8] : (where_6(X0) & from_5(X0,X5,X6,X7,X8)) => (where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
149. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_4(X0))) [skolemisation 147,148]
150. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 48]
153. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 56]
154. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 153]
155. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 57]
156. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 59]
157. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 134]
158. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
159. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 157,158]
160. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 135]
161. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 135]
164. gt(X0,X0) | where_3(X0) | lt(X0,X0) [cnf transformation 137]
165. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 141]
167. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 141]
168. ~main_query_1(X0) | select_1(X0) [cnf transformation 142]
169. ~select_1(X0) | main_query_1(X0) [cnf transformation 142]
170. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 143]
171. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 143]
174. ~gte(X0,X0) | where_6(X0) | ~lte(X0,X0) [cnf transformation 145]
175. from_5(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [cnf transformation 149]
177. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0) [cnf transformation 149]
178. ~main_query_2(X0) | select_4(X0) [cnf transformation 150]
179. ~select_4(X0) | main_query_2(X0) [cnf transformation 150]
180. eq(X0,X0) [cnf transformation 49]
183. lte(X0,X0) [cnf transformation 52]
190. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 154]
192. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 155]
195. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 156]
219. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 159]
220. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 159]
222. 0 <=> main_query_1(sK8) [avatar definition]
223. main_query_1(sK8) <- {0} [avatar component clause 222]
225. 1 <=> ~main_query_1(sK8) [avatar definition]
226. ~main_query_1(sK8) <- {1} [avatar component clause 225]
228. 2 <=> main_query_2(sK8) [avatar definition]
229. main_query_2(sK8) <- {2} [avatar component clause 228]
231. 3 <=> ~main_query_2(sK8) [avatar definition]
232. ~main_query_2(sK8) <- {3} [avatar component clause 231]
233. ~1 | ~3 [avatar split clause 220,231,225]
234. 0 | 2 [avatar split clause 219,228,222]
235. select_4(sK8) <- {2} [resolution 178,229]
242. ~lt(X0,X0) [resolution 190,180]
245. gte(X1,X1) [resolution 195,183]
249. where_3(X1) | lt(X1,X1) | lt(X1,X1) [resolution 164,192]
250. where_3(X1) | lt(X1,X1) [duplicate literal removal 249]
251. where_3(X1) [subsumption resolution 250,242]
254. where_6(X1) | ~lte(X1,X1) [resolution 245,174]
255. where_6(X1) [subsumption resolution 254,183]
334. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 165,160]
375. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 175,170]
612. from_5(X1,sK0(X1),sK1(X1),sK2(X1),sK3(X1)) | ~select_1(X1) [resolution 334,171]
623. from_2(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_4(X0) [resolution 375,161]
754. ~select_1(X1) | ~where_6(X1) | select_4(X1) [resolution 612,177]
760. ~select_1(X1) | select_4(X1) [subsumption resolution 754,255]
762. ~select_4(X1) | ~where_3(X1) | select_1(X1) [resolution 623,167]
768. ~select_4(X1) | select_1(X1) [subsumption resolution 762,251]
769. select_1(sK8) <- {2} [resolution 768,235]
778. main_query_1(sK8) <- {2} [resolution 769,169]
780. $false <- {1, 2} [subsumption resolution 778,226]
781. 1 | ~2 [avatar contradiction clause 780]
782. select_1(sK8) <- {0} [resolution 223,168]
783. select_4(sK8) <- {0} [resolution 782,760]
787. main_query_2(sK8) <- {0} [resolution 783,179]
789. $false <- {0, 3} [subsumption resolution 787,232]
790. ~0 | 3 [avatar contradiction clause 789]
791. $false [avatar sat refutation 233,234,781,790]
% SZS output end Proof for query07_query08_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.026 s
% ------------------------------
% ------------------------------
