% Refutation found. Thanks to Tanya!
% SZS status Theorem for query06_query07_tptp4X
% SZS output start Proof for query06_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (query_2(X5) <=> select_3(X5)) [input]
4. ! [X6] : (from_5(X6) <=> query_2(X6)) [input]
5. ! [X6] : (select_1(X6) <=> ($true & from_5(X6))) [input]
6. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
7. ! [X7,X8,X9,X10,X11] : (from_7(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
8. ! [X7] : (where_8(X7) <=> (~gt(X7,X7) & ~lt(X7,X7))) [input]
9. ! [X7] : (select_6(X7) <=> ? [X8,X9,X10,X11] : (where_8(X7) & from_7(X7,X8,X9,X10,X11))) [input]
10. ! [X12] : (main_query_2(X12) <=> select_6(X12)) [input]
11. ! [X13] : eq(X13,X13) [input]
17. ! [X13,X14] : (lt(X13,X14) <=> (~eq(X13,X14) & lte(X13,X14))) [input]
18. ! [X13,X14] : (gt(X13,X14) <=> lt(X14,X13)) [input]
44. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
45. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 44]
46. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 2]
47. ! [X0] : (query_2(X0) <=> select_3(X0)) [rectify 3]
48. ! [X0] : (from_5(X0) <=> query_2(X0)) [rectify 4]
49. ! [X0] : (select_1(X0) <=> ($true & from_5(X0))) [rectify 5]
50. ! [X0] : (select_1(X0) <=> from_5(X0)) [true and false elimination 49]
51. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 6]
52. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 7]
53. ! [X0] : (where_8(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [rectify 8]
54. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4))) [rectify 9]
55. ! [X0] : (main_query_2(X0) <=> select_6(X0)) [rectify 10]
56. ! [X0] : eq(X0,X0) [rectify 11]
62. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 17]
63. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 18]
89. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 45]
141. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 89]
142. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 1]
143. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 46]
144. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 143]
145. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
146. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0))) [skolemisation 144,145]
147. ! [X0] : ((query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~query_2(X0))) [nnf transformation 47]
148. ! [X0] : ((from_5(X0) | ~query_2(X0)) & (query_2(X0) | ~from_5(X0))) [nnf transformation 48]
149. ! [X0] : ((select_1(X0) | ~from_5(X0)) & (from_5(X0) | ~select_1(X0))) [nnf transformation 50]
150. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 51]
151. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 52]
152. ! [X0] : ((where_8(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_8(X0))) [nnf transformation 53]
153. ! [X0] : ((where_8(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_8(X0))) [flattening 152]
154. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4)) | ~select_6(X0))) [nnf transformation 54]
155. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) | ~select_6(X0))) [rectify 154]
156. ! [X0] : (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) => (where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
157. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & ((where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_6(X0))) [skolemisation 155,156]
158. ! [X0] : ((main_query_2(X0) | ~select_6(X0)) & (select_6(X0) | ~main_query_2(X0))) [nnf transformation 55]
161. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 62]
162. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 161]
163. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 63]
164. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 141]
165. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
166. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 164,165]
167. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 142]
168. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 142]
169. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [cnf transformation 146]
170. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 146]
171. ~query_2(X0) | select_3(X0) [cnf transformation 147]
172. ~select_3(X0) | query_2(X0) [cnf transformation 147]
173. ~from_5(X0) | query_2(X0) [cnf transformation 148]
174. ~query_2(X0) | from_5(X0) [cnf transformation 148]
175. ~select_1(X0) | from_5(X0) [cnf transformation 149]
176. ~from_5(X0) | select_1(X0) [cnf transformation 149]
177. ~main_query_1(X0) | select_1(X0) [cnf transformation 150]
178. ~select_1(X0) | main_query_1(X0) [cnf transformation 150]
179. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 151]
180. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 151]
183. gt(X0,X0) | where_8(X0) | lt(X0,X0) [cnf transformation 153]
184. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 157]
186. ~from_7(X0,X1,X2,X3,X4) | ~where_8(X0) | select_6(X0) [cnf transformation 157]
187. ~main_query_2(X0) | select_6(X0) [cnf transformation 158]
188. ~select_6(X0) | main_query_2(X0) [cnf transformation 158]
189. eq(X0,X0) [cnf transformation 56]
198. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 162]
200. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 163]
225. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 166]
226. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 166]
251. 0 <=> main_query_1(sK8) [avatar definition]
252. main_query_1(sK8) <- {0} [avatar component clause 251]
254. 1 <=> ~main_query_1(sK8) [avatar definition]
255. ~main_query_1(sK8) <- {1} [avatar component clause 254]
257. 2 <=> main_query_2(sK8) [avatar definition]
258. main_query_2(sK8) <- {2} [avatar component clause 257]
260. 3 <=> ~main_query_2(sK8) [avatar definition]
261. ~main_query_2(sK8) <- {3} [avatar component clause 260]
262. ~1 | ~3 [avatar split clause 226,260,254]
263. 0 | 2 [avatar split clause 225,257,251]
264. select_6(sK8) <- {2} [resolution 187,258]
270. ~lt(X0,X0) [resolution 198,189]
272. where_8(X1) | lt(X1,X1) | lt(X1,X1) [resolution 183,200]
273. where_8(X1) | lt(X1,X1) [duplicate literal removal 272]
274. where_8(X1) [subsumption resolution 273,270]
279. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [resolution 169,167]
281. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 184,179]
283. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [resolution 279,180]
286. from_4(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 281,168]
288. ~select_3(X1) | ~where_8(X1) | select_6(X1) [resolution 283,186]
289. ~select_3(X1) | select_6(X1) [subsumption resolution 288,274]
291. ~select_6(X1) | select_3(X1) [resolution 286,170]
292. select_3(sK8) <- {2} [resolution 291,264]
294. query_2(sK8) <- {2} [resolution 292,172]
295. from_5(sK8) <- {2} [resolution 294,174]
297. select_1(sK8) <- {2} [resolution 295,176]
299. main_query_1(sK8) <- {2} [resolution 297,178]
301. $false <- {1, 2} [subsumption resolution 299,255]
302. 1 | ~2 [avatar contradiction clause 301]
303. select_1(sK8) <- {0} [resolution 252,177]
305. from_5(sK8) <- {0} [resolution 303,175]
307. query_2(sK8) <- {0} [resolution 305,173]
309. select_3(sK8) <- {0} [resolution 307,171]
310. select_6(sK8) <- {0} [resolution 309,289]
313. main_query_2(sK8) <- {0} [resolution 310,188]
315. $false <- {0, 3} [subsumption resolution 313,261]
316. ~0 | 3 [avatar contradiction clause 315]
317. $false [avatar sat refutation 262,263,302,316]
% SZS output end Proof for query06_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.029 s
% ------------------------------
% ------------------------------
