% Refutation found. Thanks to Tanya!
% SZS status Theorem for query06_query07_tptp4X
% SZS output start Proof for query06_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : ($true & from_4(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (query_2(X5) <=> select_3(X5)) [input]
4. ! [X6] : (from_5(X6) <=> query_2(X6)) [input]
5. ! [X6] : (select_1(X6) <=> ($true & from_5(X6))) [input]
6. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
7. ! [X7,X8,X9,X10,X11] : (from_7(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
8. ! [X7] : (where_8(X7) <=> (~gt(X7,X7) & ~lt(X7,X7))) [input]
9. ! [X7] : (select_6(X7) <=> ? [X8,X9,X10,X11] : (where_8(X7) & from_7(X7,X8,X9,X10,X11))) [input]
10. ! [X12] : (main_query_2(X12) <=> select_6(X12)) [input]
11. ! [X13] : eq(X13,X13) [input]
17. ! [X13,X14] : (lt(X13,X14) <=> (~eq(X13,X14) & lte(X13,X14))) [input]
18. ! [X13,X14] : (gt(X13,X14) <=> lt(X14,X13)) [input]
52. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
53. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 52]
54. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4)) [true and false elimination 2]
55. ! [X0] : (query_2(X0) <=> select_3(X0)) [rectify 3]
56. ! [X0] : (from_5(X0) <=> query_2(X0)) [rectify 4]
57. ! [X0] : (select_1(X0) <=> ($true & from_5(X0))) [rectify 5]
58. ! [X0] : (select_1(X0) <=> from_5(X0)) [true and false elimination 57]
59. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 6]
60. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 7]
61. ! [X0] : (where_8(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [rectify 8]
62. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4))) [rectify 9]
63. ! [X0] : (main_query_2(X0) <=> select_6(X0)) [rectify 10]
64. ! [X0] : eq(X0,X0) [rectify 11]
70. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 17]
71. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 18]
105. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 53]
173. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 105]
174. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 1]
175. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_4(X0,X1,X2,X3,X4) | ~select_3(X0))) [nnf transformation 54]
176. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) | ~select_3(X0))) [rectify 175]
177. ! [X0] : (? [X5,X6,X7,X8] : from_4(X0,X5,X6,X7,X8) => from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
178. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : ~from_4(X0,X1,X2,X3,X4)) & (from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0))) [skolemisation 176,177]
179. ! [X0] : ((query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~query_2(X0))) [nnf transformation 55]
180. ! [X0] : ((from_5(X0) | ~query_2(X0)) & (query_2(X0) | ~from_5(X0))) [nnf transformation 56]
181. ! [X0] : ((select_1(X0) | ~from_5(X0)) & (from_5(X0) | ~select_1(X0))) [nnf transformation 58]
182. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 59]
183. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 60]
184. ! [X0] : ((where_8(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_8(X0))) [nnf transformation 61]
185. ! [X0] : ((where_8(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_8(X0))) [flattening 184]
186. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_8(X0) & from_7(X0,X1,X2,X3,X4)) | ~select_6(X0))) [nnf transformation 62]
187. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) | ~select_6(X0))) [rectify 186]
188. ! [X0] : (? [X5,X6,X7,X8] : (where_8(X0) & from_7(X0,X5,X6,X7,X8)) => (where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
189. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : (~where_8(X0) | ~from_7(X0,X1,X2,X3,X4))) & ((where_8(X0) & from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_6(X0))) [skolemisation 187,188]
190. ! [X0] : ((main_query_2(X0) | ~select_6(X0)) & (select_6(X0) | ~main_query_2(X0))) [nnf transformation 63]
193. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 70]
194. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 193]
195. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 71]
196. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 173]
197. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
198. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 196,197]
199. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 174]
200. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 174]
201. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [cnf transformation 178]
202. ~from_4(X0,X1,X2,X3,X4) | select_3(X0) [cnf transformation 178]
203. ~query_2(X0) | select_3(X0) [cnf transformation 179]
204. ~select_3(X0) | query_2(X0) [cnf transformation 179]
205. ~from_5(X0) | query_2(X0) [cnf transformation 180]
206. ~query_2(X0) | from_5(X0) [cnf transformation 180]
207. ~select_1(X0) | from_5(X0) [cnf transformation 181]
208. ~from_5(X0) | select_1(X0) [cnf transformation 181]
209. ~main_query_1(X0) | select_1(X0) [cnf transformation 182]
210. ~select_1(X0) | main_query_1(X0) [cnf transformation 182]
211. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 183]
212. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 183]
215. gt(X0,X0) | where_8(X0) | lt(X0,X0) [cnf transformation 185]
216. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 189]
218. ~from_7(X0,X1,X2,X3,X4) | ~where_8(X0) | select_6(X0) [cnf transformation 189]
219. ~main_query_2(X0) | select_6(X0) [cnf transformation 190]
220. ~select_6(X0) | main_query_2(X0) [cnf transformation 190]
221. eq(X0,X0) [cnf transformation 64]
230. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 194]
232. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 195]
265. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 198]
266. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 198]
299. 0 <=> main_query_1(sK8) [avatar definition]
300. main_query_1(sK8) <- {0} [avatar component clause 299]
302. 1 <=> ~main_query_1(sK8) [avatar definition]
303. ~main_query_1(sK8) <- {1} [avatar component clause 302]
305. 2 <=> main_query_2(sK8) [avatar definition]
306. main_query_2(sK8) <- {2} [avatar component clause 305]
308. 3 <=> ~main_query_2(sK8) [avatar definition]
309. ~main_query_2(sK8) <- {3} [avatar component clause 308]
310. ~1 | ~3 [avatar split clause 266,308,302]
311. 0 | 2 [avatar split clause 265,305,299]
312. select_6(sK8) <- {2} [resolution 219,306]
318. ~lt(X0,X0) [resolution 230,221]
320. where_8(X1) | lt(X1,X1) | lt(X1,X1) [resolution 215,232]
321. where_8(X1) | lt(X1,X1) [duplicate literal removal 320]
322. where_8(X1) [subsumption resolution 321,318]
327. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [resolution 201,199]
329. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 216,211]
331. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_3(X0) [resolution 327,212]
334. from_4(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 329,200]
336. ~select_3(X1) | ~where_8(X1) | select_6(X1) [resolution 331,218]
337. ~select_3(X1) | select_6(X1) [subsumption resolution 336,322]
339. ~select_6(X1) | select_3(X1) [resolution 334,202]
340. select_3(sK8) <- {2} [resolution 339,312]
342. query_2(sK8) <- {2} [resolution 340,204]
343. from_5(sK8) <- {2} [resolution 342,206]
345. select_1(sK8) <- {2} [resolution 343,208]
347. main_query_1(sK8) <- {2} [resolution 345,210]
349. $false <- {1, 2} [subsumption resolution 347,303]
350. 1 | ~2 [avatar contradiction clause 349]
351. select_1(sK8) <- {0} [resolution 300,209]
353. from_5(sK8) <- {0} [resolution 351,207]
355. query_2(sK8) <- {0} [resolution 353,205]
357. select_3(sK8) <- {0} [resolution 355,203]
358. select_6(sK8) <- {0} [resolution 357,337]
361. main_query_2(sK8) <- {0} [resolution 358,220]
363. $false <- {0, 3} [subsumption resolution 361,309]
364. ~0 | 3 [avatar contradiction clause 363]
365. $false [avatar sat refutation 310,311,350,364]
% SZS output end Proof for query06_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
