% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_06_query_09_tptp4X
% SZS output start Proof for query_06_query_09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> (gte(X0,X0) & lte(X0,X0))) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (where_6_subquery_7(X5) <=> select_8(X5)) [input]
9. ! [X0] : (where_6(X0) <=> where_6_subquery_7(X0)) [input]
10. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
15. ! [X7] : lte(X7,X7) [input]
21. ! [X7,X8] : (gte(X7,X8) <=> lte(X8,X7)) [input]
55. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
56. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 55]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
58. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 7]
59. ! [X0] : (where_6_subquery_7(X0) <=> select_8(X0)) [rectify 8]
60. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
64. ! [X0] : lte(X0,X0) [rectify 15]
70. ! [X0,X1] : (gte(X0,X1) <=> lte(X1,X0)) [rectify 21]
104. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 56]
176. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 104]
177. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
178. ! [X0] : ((where_3(X0) | (~gte(X0,X0) | ~lte(X0,X0))) & ((gte(X0,X0) & lte(X0,X0)) | ~where_3(X0))) [nnf transformation 2]
179. ! [X0] : ((where_3(X0) | ~gte(X0,X0) | ~lte(X0,X0)) & ((gte(X0,X0) & lte(X0,X0)) | ~where_3(X0))) [flattening 178]
180. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
181. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 180]
182. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
183. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 181,182]
184. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
185. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
186. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
187. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 58]
188. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 187]
189. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
190. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 188,189]
191. ! [X0] : ((where_6_subquery_7(X0) | ~select_8(X0)) & (select_8(X0) | ~where_6_subquery_7(X0))) [nnf transformation 59]
192. ! [X0] : ((where_6(X0) | ~where_6_subquery_7(X0)) & (where_6_subquery_7(X0) | ~where_6(X0))) [nnf transformation 9]
193. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 10]
194. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 193]
195. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
196. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 194,195]
197. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 60]
200. ! [X0,X1] : ((gte(X0,X1) | ~lte(X1,X0)) & (lte(X1,X0) | ~gte(X0,X1))) [nnf transformation 70]
201. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 176]
202. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
203. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 201,202]
204. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 177]
205. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 177]
208. ~gte(X0,X0) | where_3(X0) | ~lte(X0,X0) [cnf transformation 179]
209. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 183]
211. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0,X1) [cnf transformation 183]
212. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 184]
213. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 184]
214. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 185]
215. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 185]
217. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 186]
219. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 190]
221. ~select_8(X0) | where_6_subquery_7(X0) [cnf transformation 191]
223. ~where_6_subquery_7(X0) | where_6(X0) [cnf transformation 192]
224. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 196]
226. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0) | select_4(X0,X1) [cnf transformation 196]
227. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 197]
228. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 197]
232. lte(X0,X0) [cnf transformation 64]
238. ~lte(X1,X0) | gte(X0,X1) [cnf transformation 200]
272. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 203]
273. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 203]
275. 0 <=> main_query_1(sK10,sK11) [avatar definition]
276. main_query_1(sK10,sK11) <- {0} [avatar component clause 275]
278. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
279. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 278]
281. 2 <=> main_query_2(sK10,sK11) [avatar definition]
282. main_query_2(sK10,sK11) <- {2} [avatar component clause 281]
284. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
285. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 284]
286. ~1 | ~3 [avatar split clause 273,284,278]
287. 0 | 2 [avatar split clause 272,281,275]
288. select_4(sK10,sK11) <- {2} [resolution 227,282]
300. gte(X1,X1) [resolution 238,232]
303. where_3(X1) | ~lte(X1,X1) [resolution 300,208]
304. where_3(X1) [subsumption resolution 303,232]
320. tab1(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 209,204]
325. tab1(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 224,214]
369. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 320,215]
370. from_9(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 320,217]
376. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 325,205]
503. ~select_1(X2,X3) | ~where_6(X2) | select_4(X2,X3) [resolution 369,226]
525. ~select_1(X2,X3) | select_8(X2) [resolution 370,219]
547. ~select_4(X2,X3) | ~where_3(X2) | select_1(X2,X3) [resolution 376,211]
553. ~select_4(X2,X3) | select_1(X2,X3) [subsumption resolution 547,304]
571. select_1(sK10,sK11) <- {2} [resolution 553,288]
583. main_query_1(sK10,sK11) <- {2} [resolution 571,213]
585. $false <- {1, 2} [subsumption resolution 583,279]
586. 1 | ~2 [avatar contradiction clause 585]
587. select_1(sK10,sK11) <- {0} [resolution 276,212]
588. select_8(sK10) <- {0} [resolution 587,525]
589. ~where_6(sK10) | select_4(sK10,sK11) <- {0} [resolution 587,503]
596. 4 <=> select_4(sK10,sK11) [avatar definition]
597. select_4(sK10,sK11) <- {4} [avatar component clause 596]
599. 6 <=> where_6(sK10) [avatar definition]
602. 7 <=> ~where_6(sK10) [avatar definition]
604. 4 | ~7 | ~0 [avatar split clause 589,275,602,596]
605. where_6_subquery_7(sK10) <- {0} [resolution 588,221]
622. where_6(sK10) <- {0} [resolution 605,223]
626. 6 | ~0 [avatar split clause 622,275,599]
646. main_query_2(sK10,sK11) <- {4} [resolution 597,228]
648. $false <- {3, 4} [subsumption resolution 646,285]
649. 3 | ~4 [avatar contradiction clause 648]
650. $false [avatar sat refutation 286,287,586,604,626,649]
% SZS output end Proof for query_06_query_09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.012 s
% ------------------------------
% ------------------------------
