% Refutation found. Thanks to Tanya!
% SZS status Theorem for query10_query11_tptp4X
% SZS output start Proof for query10_query11_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X8,X7] : (where_6(X8,X7) <=> eq(X7,X8)) [input]
7. ! [X8,X7] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X8,X7) & from_5(X7,X8,X9,X10,X11))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
54. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
55. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 54]
56. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
57. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
58. ! [X0,X1] : (where_6(X0,X1) <=> eq(X1,X0)) [rectify 6]
59. ! [X0,X1] : (select_4(X1,X0) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X1,X0,X2,X3,X4))) [rectify 7]
60. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
106. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 55]
182. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 106]
183. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
184. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
185. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
186. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 185]
187. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
188. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 186,187]
189. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 56]
190. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 57]
191. ! [X0,X1] : ((where_6(X0,X1) | ~eq(X1,X0)) & (eq(X1,X0) | ~where_6(X0,X1))) [nnf transformation 58]
192. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X1,X0,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X1,X0,X2,X3,X4)) | ~select_4(X1,X0))) [nnf transformation 59]
193. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X1,X0,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X1,X0,X5,X6,X7)) | ~select_4(X1,X0))) [rectify 192]
194. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X1,X0,X5,X6,X7)) => (where_6(X0,X1) & from_5(X1,X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
195. ! [X0,X1] : ((select_4(X1,X0) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X1,X0,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X1,X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X1,X0))) [skolemisation 193,194]
196. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 60]
199. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 182]
200. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
201. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 199,200]
202. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 183]
203. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 183]
204. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 184]
205. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 184]
206. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 188]
207. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 188]
208. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 188]
209. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 189]
210. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 189]
211. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 190]
212. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 190]
213. ~where_6(X0,X1) | eq(X1,X0) [cnf transformation 191]
214. ~eq(X1,X0) | where_6(X0,X1) [cnf transformation 191]
215. from_5(X1,X0,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X1,X0) [cnf transformation 195]
216. ~select_4(X1,X0) | where_6(X0,X1) [cnf transformation 195]
217. ~from_5(X1,X0,X2,X3,X4) | ~where_6(X0,X1) | select_4(X1,X0) [cnf transformation 195]
218. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 196]
219. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 196]
263. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 201]
264. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 201]
301. 0 <=> main_query_1(sK6,sK7) [avatar definition]
302. main_query_1(sK6,sK7) <- {0} [avatar component clause 301]
304. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
305. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 304]
307. 2 <=> main_query_2(sK6,sK7) [avatar definition]
308. main_query_2(sK6,sK7) <- {2} [avatar component clause 307]
310. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
311. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 310]
312. ~1 | ~3 [avatar split clause 264,310,304]
313. 0 | 2 [avatar split clause 263,307,301]
318. select_4(sK6,sK7) <- {2} [resolution 218,308]
320. where_6(sK7,sK6) <- {2} [resolution 318,216]
322. eq(sK6,sK7) <- {2} [resolution 320,213]
327. where_3(sK6,sK7) <- {2} [resolution 322,205]
356. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 206,202]
358. tab1(X2,X3,sK3(X3,X2),sK4(X3,X2),sK5(X3,X2)) | ~select_4(X2,X3) [resolution 215,211]
359. from_5(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 356,212]
362. from_2(X2,X3,sK3(X3,X2),sK4(X3,X2),sK5(X3,X2)) | ~select_4(X2,X3) [resolution 358,203]
363. ~where_6(X1,X0) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 359,217]
381. ~select_4(X0,X1) | ~where_3(X0,X1) | select_1(X0,X1) [resolution 362,208]
383. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 381,318]
384. select_1(sK6,sK7) <- {2} [subsumption resolution 383,327]
385. main_query_1(sK6,sK7) <- {2} [resolution 384,210]
387. $false <- {1, 2} [subsumption resolution 385,305]
388. 1 | ~2 [avatar contradiction clause 387]
389. select_1(sK6,sK7) <- {0} [resolution 302,209]
391. where_3(sK6,sK7) <- {0} [resolution 389,207]
392. eq(sK6,sK7) <- {0} [resolution 391,204]
397. where_6(sK7,sK6) <- {0} [resolution 392,214]
409. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 397,363]
411. select_4(sK6,sK7) <- {0} [subsumption resolution 409,389]
418. main_query_2(sK6,sK7) <- {0} [resolution 411,219]
420. $false <- {0, 3} [subsumption resolution 418,311]
421. ~0 | 3 [avatar contradiction clause 420]
422. $false [avatar sat refutation 312,313,388,421]
% SZS output end Proof for query10_query11_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.020 s
% ------------------------------
% ------------------------------
