% Refutation found. Thanks to Tanya!
% SZS status Theorem for query21_query22_tptp4X
% SZS output start Proof for query21_query22_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X5,X6] : (where_5(X5,X6) <=> eq(X5,X6)) [input]
3. ! [X5,X6] : (select_3(X5,X6) <=> ? [X0,X1,X2,X3,X4] : (where_5(X5,X6) & from_4(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (query_2(X5,X6) <=> select_3(X5,X6)) [input]
5. ! [X7,X8] : (from_6(X7,X8) <=> query_2(X7,X8)) [input]
6. ! [X7] : (select_1(X7) <=> ? [X8] : ($true & from_6(X7,X8))) [input]
7. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
8. ! [X0,X1,X2,X3,X4] : (from_10(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
9. ! [X5,X6] : (where_11(X5,X6) <=> eq(X5,X6)) [input]
10. ! [X5,X6] : (select_9(X5,X6) <=> ? [X0,X1,X2,X3,X4] : (where_11(X5,X6) & from_10(X0,X1,X2,X3,X4))) [input]
11. ! [X9,X6] : (query_8(X9,X6) <=> select_9(X9,X6)) [input]
12. ! [X10,X8] : (from_12(X10,X8) <=> query_8(X10,X8)) [input]
13. ! [X10] : (select_7(X10) <=> ? [X8] : ($true & from_12(X10,X8))) [input]
14. ! [X9] : (main_query_2(X9) <=> select_7(X9)) [input]
15. ! [X11] : eq(X11,X11) [input]
65. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
66. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 65]
67. ! [X0,X1] : (where_5(X0,X1) <=> eq(X0,X1)) [rectify 2]
68. ! [X0,X1] : (select_3(X0,X1) <=> ? [X2,X3,X4,X5,X6] : (where_5(X0,X1) & from_4(X2,X3,X4,X5,X6))) [rectify 3]
69. ! [X0,X1] : (query_2(X0,X1) <=> select_3(X0,X1)) [rectify 4]
70. ! [X0,X1] : (from_6(X0,X1) <=> query_2(X0,X1)) [rectify 5]
71. ! [X0] : (select_1(X0) <=> ? [X1] : ($true & from_6(X0,X1))) [rectify 6]
72. ! [X0] : (select_1(X0) <=> ? [X1] : from_6(X0,X1)) [true and false elimination 71]
73. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 7]
74. ! [X0,X1] : (where_11(X0,X1) <=> eq(X0,X1)) [rectify 9]
75. ! [X0,X1] : (select_9(X0,X1) <=> ? [X2,X3,X4,X5,X6] : (where_11(X0,X1) & from_10(X2,X3,X4,X5,X6))) [rectify 10]
76. ! [X0,X1] : (query_8(X0,X1) <=> select_9(X0,X1)) [rectify 11]
77. ! [X0,X1] : (from_12(X0,X1) <=> query_8(X0,X1)) [rectify 12]
78. ! [X0] : (select_7(X0) <=> ? [X1] : ($true & from_12(X0,X1))) [rectify 13]
79. ! [X0] : (select_7(X0) <=> ? [X1] : from_12(X0,X1)) [true and false elimination 78]
80. ! [X0] : (main_query_2(X0) <=> select_7(X0)) [rectify 14]
81. ! [X0] : eq(X0,X0) [rectify 15]
131. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 66]
217. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 131]
218. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 1]
219. ! [X0,X1] : ((where_5(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_5(X0,X1))) [nnf transformation 67]
220. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_5(X0,X1) | ~from_4(X2,X3,X4,X5,X6))) & (? [X2,X3,X4,X5,X6] : (where_5(X0,X1) & from_4(X2,X3,X4,X5,X6)) | ~select_3(X0,X1))) [nnf transformation 68]
221. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_5(X0,X1) | ~from_4(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_5(X0,X1) & from_4(X7,X8,X9,X10,X11)) | ~select_3(X0,X1))) [rectify 220]
222. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_5(X0,X1) & from_4(X7,X8,X9,X10,X11)) => (where_5(X0,X1) & from_4(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)))) [choice axiom]
223. ! [X0,X1] : ((select_3(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_5(X0,X1) | ~from_4(X2,X3,X4,X5,X6))) & ((where_5(X0,X1) & from_4(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1))) | ~select_3(X0,X1))) [skolemisation 221,222]
224. ! [X0,X1] : ((query_2(X0,X1) | ~select_3(X0,X1)) & (select_3(X0,X1) | ~query_2(X0,X1))) [nnf transformation 69]
225. ! [X0,X1] : ((from_6(X0,X1) | ~query_2(X0,X1)) & (query_2(X0,X1) | ~from_6(X0,X1))) [nnf transformation 70]
226. ! [X0] : ((select_1(X0) | ! [X1] : ~from_6(X0,X1)) & (? [X1] : from_6(X0,X1) | ~select_1(X0))) [nnf transformation 72]
227. ! [X0] : ((select_1(X0) | ! [X1] : ~from_6(X0,X1)) & (? [X2] : from_6(X0,X2) | ~select_1(X0))) [rectify 226]
228. ! [X0] : (? [X2] : from_6(X0,X2) => from_6(X0,sK5(X0))) [choice axiom]
229. ! [X0] : ((select_1(X0) | ! [X1] : ~from_6(X0,X1)) & (from_6(X0,sK5(X0)) | ~select_1(X0))) [skolemisation 227,228]
230. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 73]
231. ! [X0,X1,X2,X3,X4] : ((from_10(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_10(X0,X1,X2,X3,X4))) [nnf transformation 8]
232. ! [X0,X1] : ((where_11(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_11(X0,X1))) [nnf transformation 74]
233. ! [X0,X1] : ((select_9(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_11(X0,X1) | ~from_10(X2,X3,X4,X5,X6))) & (? [X2,X3,X4,X5,X6] : (where_11(X0,X1) & from_10(X2,X3,X4,X5,X6)) | ~select_9(X0,X1))) [nnf transformation 75]
234. ! [X0,X1] : ((select_9(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_11(X0,X1) | ~from_10(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_11(X0,X1) & from_10(X7,X8,X9,X10,X11)) | ~select_9(X0,X1))) [rectify 233]
235. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_11(X0,X1) & from_10(X7,X8,X9,X10,X11)) => (where_11(X0,X1) & from_10(sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)))) [choice axiom]
236. ! [X0,X1] : ((select_9(X0,X1) | ! [X2,X3,X4,X5,X6] : (~where_11(X0,X1) | ~from_10(X2,X3,X4,X5,X6))) & ((where_11(X0,X1) & from_10(sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1))) | ~select_9(X0,X1))) [skolemisation 234,235]
237. ! [X0,X1] : ((query_8(X0,X1) | ~select_9(X0,X1)) & (select_9(X0,X1) | ~query_8(X0,X1))) [nnf transformation 76]
238. ! [X0,X1] : ((from_12(X0,X1) | ~query_8(X0,X1)) & (query_8(X0,X1) | ~from_12(X0,X1))) [nnf transformation 77]
239. ! [X0] : ((select_7(X0) | ! [X1] : ~from_12(X0,X1)) & (? [X1] : from_12(X0,X1) | ~select_7(X0))) [nnf transformation 79]
240. ! [X0] : ((select_7(X0) | ! [X1] : ~from_12(X0,X1)) & (? [X2] : from_12(X0,X2) | ~select_7(X0))) [rectify 239]
241. ! [X0] : (? [X2] : from_12(X0,X2) => from_12(X0,sK11(X0))) [choice axiom]
242. ! [X0] : ((select_7(X0) | ! [X1] : ~from_12(X0,X1)) & (from_12(X0,sK11(X0)) | ~select_7(X0))) [skolemisation 240,241]
243. ! [X0] : ((main_query_2(X0) | ~select_7(X0)) & (select_7(X0) | ~main_query_2(X0))) [nnf transformation 80]
246. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 217]
247. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK12) | ~main_query_1(sK12)) & (main_query_2(sK12) | main_query_1(sK12))) [choice axiom]
248. (~main_query_2(sK12) | ~main_query_1(sK12)) & (main_query_2(sK12) | main_query_1(sK12)) [skolemisation 246,247]
249. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 218]
250. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 218]
252. ~eq(X0,X1) | where_5(X0,X1) [cnf transformation 219]
253. from_4(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_3(X0,X1) [cnf transformation 223]
255. select_3(X0,X1) | ~where_5(X0,X1) | ~from_4(X2,X3,X4,X5,X6) [cnf transformation 223]
256. ~query_2(X0,X1) | select_3(X0,X1) [cnf transformation 224]
257. ~select_3(X0,X1) | query_2(X0,X1) [cnf transformation 224]
258. ~from_6(X0,X1) | query_2(X0,X1) [cnf transformation 225]
259. ~query_2(X0,X1) | from_6(X0,X1) [cnf transformation 225]
260. from_6(X0,sK5(X0)) | ~select_1(X0) [cnf transformation 229]
261. ~from_6(X0,X1) | select_1(X0) [cnf transformation 229]
262. ~main_query_1(X0) | select_1(X0) [cnf transformation 230]
263. ~select_1(X0) | main_query_1(X0) [cnf transformation 230]
264. ~from_10(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 231]
265. ~tab1(X0,X1,X2,X3,X4) | from_10(X0,X1,X2,X3,X4) [cnf transformation 231]
267. ~eq(X0,X1) | where_11(X0,X1) [cnf transformation 232]
268. from_10(sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) | ~select_9(X0,X1) [cnf transformation 236]
270. select_9(X0,X1) | ~where_11(X0,X1) | ~from_10(X2,X3,X4,X5,X6) [cnf transformation 236]
271. ~query_8(X0,X1) | select_9(X0,X1) [cnf transformation 237]
272. ~select_9(X0,X1) | query_8(X0,X1) [cnf transformation 237]
273. ~from_12(X0,X1) | query_8(X0,X1) [cnf transformation 238]
274. ~query_8(X0,X1) | from_12(X0,X1) [cnf transformation 238]
275. from_12(X0,sK11(X0)) | ~select_7(X0) [cnf transformation 242]
276. ~from_12(X0,X1) | select_7(X0) [cnf transformation 242]
277. ~main_query_2(X0) | select_7(X0) [cnf transformation 243]
278. ~select_7(X0) | main_query_2(X0) [cnf transformation 243]
279. eq(X0,X0) [cnf transformation 81]
327. main_query_2(sK12) | main_query_1(sK12) [cnf transformation 248]
328. ~main_query_2(sK12) | ~main_query_1(sK12) [cnf transformation 248]
370. 0 <=> ! [X3,X5,X2,X4,X6] : ~from_4(X2,X3,X4,X5,X6) [avatar definition]
371. ~from_4(X2,X3,X4,X5,X6) <- {0} [avatar component clause 370]
373. 2 <=> ! [X1,X0] : (select_3(X0,X1) | ~where_5(X0,X1)) [avatar definition]
374. ~where_5(X0,X1) | select_3(X0,X1) <- {2} [avatar component clause 373]
375. 0 | 2 [avatar split clause 255,373,370]
377. 4 <=> ! [X3,X5,X2,X4,X6] : ~from_10(X2,X3,X4,X5,X6) [avatar definition]
378. ~from_10(X2,X3,X4,X5,X6) <- {4} [avatar component clause 377]
380. 6 <=> ! [X1,X0] : (select_9(X0,X1) | ~where_11(X0,X1)) [avatar definition]
381. ~where_11(X0,X1) | select_9(X0,X1) <- {6} [avatar component clause 380]
382. 4 | 6 [avatar split clause 270,380,377]
384. 8 <=> main_query_1(sK12) [avatar definition]
385. main_query_1(sK12) <- {8} [avatar component clause 384]
387. 9 <=> ~main_query_1(sK12) [avatar definition]
388. ~main_query_1(sK12) <- {9} [avatar component clause 387]
390. 10 <=> main_query_2(sK12) [avatar definition]
391. main_query_2(sK12) <- {10} [avatar component clause 390]
393. 11 <=> ~main_query_2(sK12) [avatar definition]
394. ~main_query_2(sK12) <- {11} [avatar component clause 393]
395. ~9 | ~11 [avatar split clause 328,393,387]
396. 8 | 10 [avatar split clause 327,390,384]
397. select_1(sK12) <- {8} [resolution 262,385]
399. where_5(X0,X0) [resolution 252,279]
401. query_2(X0,sK5(X0)) | ~select_1(X0) [resolution 260,258]
403. where_11(X0,X0) [resolution 267,279]
405. query_8(X0,sK11(X0)) | ~select_7(X0) [resolution 275,273]
423. select_3(X1,sK5(X1)) | ~select_1(X1) [resolution 401,256]
429. tab1(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_3(X0,X1) [resolution 253,249]
432. tab1(sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) | ~select_9(X0,X1) [resolution 268,264]
433. ~select_9(X2,X3) <- {4} [resolution 268,378]
435. select_9(X1,sK11(X1)) | ~select_7(X1) [resolution 405,271]
436. ~select_7(X1) <- {4} [subsumption resolution 435,433]
437. from_10(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_3(X0,X1) [resolution 429,265]
439. ~select_3(X0,X1) <- {4} [subsumption resolution 437,378]
444. ~select_1(X1) <- {4} [subsumption resolution 423,439]
446. $false <- {4, 8} [resolution 444,397]
447. ~4 | ~8 [avatar contradiction clause 446]
448. select_7(sK12) <- {10} [resolution 391,277]
449. $false <- {4, 10} [subsumption resolution 448,436]
450. ~4 | ~10 [avatar contradiction clause 449]
452. select_9(X0,X0) <- {6} [resolution 381,403]
455. from_4(sK6(X2,X3),sK7(X2,X3),sK8(X2,X3),sK9(X2,X3),sK10(X2,X3)) | ~select_9(X2,X3) [resolution 432,250]
456. ~select_9(X2,X3) <- {0} [subsumption resolution 455,371]
457. query_8(X0,X0) <- {6} [resolution 452,272]
459. $false <- {0, 6} [resolution 456,452]
460. ~0 | ~6 [avatar contradiction clause 459]
461. from_12(X0,X0) <- {6} [resolution 457,274]
463. select_3(X0,X0) <- {2} [resolution 374,399]
464. query_2(X0,X0) <- {2} [resolution 463,257]
468. from_6(X0,X0) <- {2} [resolution 464,259]
473. select_7(X1) <- {6} [resolution 461,276]
475. main_query_2(X0) <- {6} [resolution 473,278]
481. select_1(X1) <- {2} [resolution 468,261]
485. main_query_1(X0) <- {2} [resolution 481,263]
493. $false <- {2, 9} [resolution 485,388]
494. ~2 | 9 [avatar contradiction clause 493]
495. $false <- {6, 11} [subsumption resolution 394,475]
496. ~6 | 11 [avatar contradiction clause 495]
497. $false [avatar sat refutation 375,382,395,396,447,450,460,494,496]
% SZS output end Proof for query21_query22_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.013 s
% ------------------------------
% ------------------------------
