% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_12_query_13_tptp4X
% SZS output start Proof for query_12_query_13_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8] : (where_6(X7,X8) <=> (lte(X8,X7) & lte(X7,X8))) [input]
7. ! [X7,X8] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X7,X8) & from_5(X7,X8,X9,X10,X11))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
12. ! [X12] : lte(X12,X12) [input]
13. ! [X12,X13] : ((lte(X13,X12) & lte(X12,X13)) <=> eq(X12,X13)) [input]
28. ! [X12,X7,X8] : ((where_6(X7,X8) & eq(X8,X12)) => where_6(X7,X12)) [input]
52. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
53. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 52]
54. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
55. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
56. ! [X0,X1] : (where_6(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [rectify 6]
57. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
58. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
62. ! [X0] : lte(X0,X0) [rectify 12]
63. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
78. ! [X0,X1,X2] : ((where_6(X1,X2) & eq(X2,X0)) => where_6(X1,X0)) [rectify 28]
102. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 53]
124. ! [X0,X1,X2] : (where_6(X1,X0) | (~where_6(X1,X2) | ~eq(X2,X0))) [ennf transformation 78]
125. ! [X0,X1,X2] : (where_6(X1,X0) | ~where_6(X1,X2) | ~eq(X2,X0)) [flattening 124]
172. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 102]
173. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
174. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
175. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
176. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 175]
177. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
178. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 176,177]
179. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 54]
180. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 55]
181. ! [X0,X1] : ((where_6(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 56]
182. ! [X0,X1] : ((where_6(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [flattening 181]
183. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 57]
184. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 183]
185. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
186. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 184,185]
187. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 58]
188. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 63]
189. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 188]
190. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 172]
191. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
192. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 190,191]
193. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 173]
194. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 173]
195. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 174]
196. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 174]
197. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 178]
198. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 178]
199. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 178]
200. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 179]
201. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 179]
202. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 180]
203. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 180]
204. ~where_6(X0,X1) | lte(X0,X1) [cnf transformation 182]
205. ~where_6(X0,X1) | lte(X1,X0) [cnf transformation 182]
206. ~lte(X1,X0) | where_6(X0,X1) | ~lte(X0,X1) [cnf transformation 182]
207. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 186]
208. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 186]
209. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 186]
210. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 187]
211. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 187]
215. lte(X0,X0) [cnf transformation 62]
216. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 189]
229. ~where_6(X1,X2) | where_6(X1,X0) | ~eq(X2,X0) [cnf transformation 125]
253. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 192]
254. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 192]
256. 0 <=> main_query_1(sK6,sK7) [avatar definition]
257. main_query_1(sK6,sK7) <- {0} [avatar component clause 256]
259. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
260. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 259]
262. 2 <=> main_query_2(sK6,sK7) [avatar definition]
263. main_query_2(sK6,sK7) <- {2} [avatar component clause 262]
265. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
266. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 265]
267. ~1 | ~3 [avatar split clause 254,265,259]
268. 0 | 2 [avatar split clause 253,262,256]
271. select_4(sK6,sK7) <- {2} [resolution 210,263]
272. where_6(sK6,sK7) <- {2} [resolution 271,208]
274. lte(sK7,sK6) <- {2} [resolution 272,205]
275. lte(sK6,sK7) <- {2} [resolution 272,204]
280. where_6(X0,X0) | ~lte(X0,X0) [resolution 206,215]
285. where_6(X0,X0) [subsumption resolution 280,215]
295. eq(sK6,sK7) | ~lte(sK6,sK7) <- {2} [resolution 216,274]
297. eq(sK6,sK7) <- {2} [subsumption resolution 295,275]
303. where_3(sK6,sK7) <- {2} [resolution 297,196]
323. ~eq(X1,X2) | where_6(X1,X2) [resolution 229,285]
377. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 197,193]
381. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 207,202]
487. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 377,203]
506. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 381,194]
829. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 487,209]
839. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 506,199]
852. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 839,271]
856. select_1(sK6,sK7) <- {2} [subsumption resolution 852,303]
869. main_query_1(sK6,sK7) <- {2} [resolution 856,201]
871. $false <- {1, 2} [subsumption resolution 869,260]
872. 1 | ~2 [avatar contradiction clause 871]
875. select_1(sK6,sK7) <- {0} [resolution 257,200]
879. where_3(sK6,sK7) <- {0} [resolution 875,198]
880. eq(sK6,sK7) <- {0} [resolution 879,195]
890. where_6(sK6,sK7) <- {0} [resolution 880,323]
949. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 890,829]
954. select_4(sK6,sK7) <- {0} [subsumption resolution 949,875]
1018. main_query_2(sK6,sK7) <- {0} [resolution 954,211]
1020. $false <- {0, 3} [subsumption resolution 1018,266]
1021. ~0 | 3 [avatar contradiction clause 1020]
1022. $false [avatar sat refutation 267,268,872,1021]
% SZS output end Proof for query_12_query_13_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.022 s
% ------------------------------
% ------------------------------
