% Refutation found. Thanks to Tanya!
% SZS status Theorem for query02_query04_tptp4X
% SZS output start Proof for query02_query04_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
4. ! [X6,X7,X8,X9,X10] : (from_4(X6,X7,X8,X9,X10) <=> tab1(X6,X7,X8,X9,X10)) [input]
5. ! [X6] : (where_5(X6) <=> eq(X6,X6)) [input]
6. ! [X6] : (select_3(X6) <=> ? [X7,X8,X9,X10] : (where_5(X6) & from_4(X6,X7,X8,X9,X10))) [input]
7. ! [X6] : (main_query_2(X6) <=> select_3(X6)) [input]
8. ! [X11] : eq(X11,X11) [input]
42. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
43. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 42]
44. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
45. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
46. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
47. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [rectify 5]
48. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
49. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 7]
50. ! [X0] : eq(X0,X0) [rectify 8]
84. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 43]
136. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 84]
137. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
138. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 44]
139. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 138]
140. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
141. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 139,140]
142. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 45]
143. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 46]
144. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 47]
145. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 48]
146. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 145]
147. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
148. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 146,147]
149. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 49]
152. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 136]
153. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
154. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 152,153]
155. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 137]
156. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 137]
157. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 141]
158. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 141]
159. ~main_query_1(X0) | select_1(X0) [cnf transformation 142]
160. ~select_1(X0) | main_query_1(X0) [cnf transformation 142]
161. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 143]
162. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 143]
164. ~eq(X0,X0) | where_5(X0) [cnf transformation 144]
165. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 148]
167. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 148]
168. ~main_query_2(X0) | select_3(X0) [cnf transformation 149]
169. ~select_3(X0) | main_query_2(X0) [cnf transformation 149]
170. eq(X0,X0) [cnf transformation 50]
202. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 154]
203. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 154]
205. 0 <=> main_query_1(sK8) [avatar definition]
206. main_query_1(sK8) <- {0} [avatar component clause 205]
208. 1 <=> ~main_query_1(sK8) [avatar definition]
209. ~main_query_1(sK8) <- {1} [avatar component clause 208]
211. 2 <=> main_query_2(sK8) [avatar definition]
212. main_query_2(sK8) <- {2} [avatar component clause 211]
214. 3 <=> ~main_query_2(sK8) [avatar definition]
215. ~main_query_2(sK8) <- {3} [avatar component clause 214]
216. ~1 | ~3 [avatar split clause 203,214,208]
217. 0 | 2 [avatar split clause 202,211,205]
218. select_3(sK8) <- {2} [resolution 168,212]
222. where_5(X1) [resolution 164,170]
256. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 157,155]
258. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 165,161]
260. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 256,162]
264. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 258,156]
267. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 260,167]
268. ~select_1(X1) | select_3(X1) [subsumption resolution 267,222]
276. ~select_3(X9) | select_1(X9) [resolution 264,158]
279. select_1(sK8) <- {2} [resolution 276,218]
281. main_query_1(sK8) <- {2} [resolution 279,160]
282. $false <- {1, 2} [subsumption resolution 281,209]
283. 1 | ~2 [avatar contradiction clause 282]
284. select_1(sK8) <- {0} [resolution 206,159]
285. select_3(sK8) <- {0} [resolution 284,268]
288. main_query_2(sK8) <- {0} [resolution 285,169]
290. $false <- {0, 3} [subsumption resolution 288,215]
291. ~0 | 3 [avatar contradiction clause 290]
292. $false [avatar sat refutation 216,217,283,291]
% SZS output end Proof for query02_query04_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
