% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_20_query_30_tptp4X
% SZS output start Proof for query_20_query_30_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X5,X6,X7,X8,X9] : (from_6(X5,X6,X7,X8,X9) <=> tab2(X5,X6,X7,X8,X9)) [input]
3. ! [X5] : (select_5(X5) <=> ? [X6,X7,X8,X9] : ($true & from_6(X5,X6,X7,X8,X9))) [input]
4. ! [X10] : (where_3_subquery_4(X10) <=> select_5(X10)) [input]
5. ! [X0] : (where_3(X0) <=> where_3_subquery_4(X0)) [input]
6. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
7. ! [X11,X12] : (main_query_1(X11,X12) <=> select_1(X11,X12)) [input]
8. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) <=> (tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4))) [input]
10. ! [X5,X0] : (where_13(X5,X0) <=> neq(X5,X0)) [input]
11. ! [X5,X0] : (select_11(X5,X0) <=> ? [X6,X7,X8,X9] : (where_13(X5,X0) & from_12(X5,X6,X7,X8,X9))) [input]
12. ! [X10,X0] : (where_9_subquery_10(X10,X0) <=> select_11(X10,X0)) [input]
13. ! [X5,X6,X7,X8,X9] : (from_16(X5,X6,X7,X8,X9) <=> tab2(X5,X6,X7,X8,X9)) [input]
14. ! [X5,X0] : (where_17(X5,X0) <=> eq(X5,X0)) [input]
15. ! [X5,X0] : (select_15(X5,X0) <=> ? [X6,X7,X8,X9] : (where_17(X5,X0) & from_16(X5,X6,X7,X8,X9))) [input]
16. ! [X10,X0] : (where_9_subquery_14(X10,X0) <=> select_15(X10,X0)) [input]
17. ! [X0] : (where_9(X0) <=> (where_9_subquery_14(X0,X0) & ~where_9_subquery_10(X0,X0))) [input]
18. ! [X0,X1] : (select_7(X0,X1) <=> ? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_9(X0) & from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [input]
19. ! [X11,X12] : (main_query_2(X11,X12) <=> select_7(X11,X12)) [input]
20. ! [X13] : eq(X13,X13) [input]
21. ! [X13,X14] : (eq(X13,X14) => eq(X14,X13)) [input]
28. ! [X13,X14] : (neq(X13,X14) <=> ~eq(X13,X14)) [input]
37. ! [X13,X5,X0] : ((where_13(X5,X0) & eq(X5,X13)) => where_13(X13,X0)) [input]
90. ! [X11,X12] : (main_query_1(X11,X12) <=> main_query_2(X11,X12)) [input]
91. ~! [X11,X12] : (main_query_1(X11,X12) <=> main_query_2(X11,X12)) [negated conjecture 90]
92. ! [X0,X1,X2,X3,X4] : (from_6(X0,X1,X2,X3,X4) <=> tab2(X0,X1,X2,X3,X4)) [rectify 2]
93. ! [X0] : (select_5(X0) <=> ? [X1,X2,X3,X4] : ($true & from_6(X0,X1,X2,X3,X4))) [rectify 3]
94. ! [X0] : (select_5(X0) <=> ? [X1,X2,X3,X4] : from_6(X0,X1,X2,X3,X4)) [true and false elimination 93]
95. ! [X0] : (where_3_subquery_4(X0) <=> select_5(X0)) [rectify 4]
96. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 7]
98. ! [X0,X1] : (where_13(X0,X1) <=> neq(X0,X1)) [rectify 10]
99. ! [X0,X1] : (select_11(X0,X1) <=> ? [X2,X3,X4,X5] : (where_13(X0,X1) & from_12(X0,X2,X3,X4,X5))) [rectify 11]
100. ! [X0,X1] : (where_9_subquery_10(X0,X1) <=> select_11(X0,X1)) [rectify 12]
101. ! [X0,X1,X2,X3,X4] : (from_16(X0,X1,X2,X3,X4) <=> tab2(X0,X1,X2,X3,X4)) [rectify 13]
102. ! [X0,X1] : (where_17(X0,X1) <=> eq(X0,X1)) [rectify 14]
103. ! [X0,X1] : (select_15(X0,X1) <=> ? [X2,X3,X4,X5] : (where_17(X0,X1) & from_16(X0,X2,X3,X4,X5))) [rectify 15]
104. ! [X0,X1] : (where_9_subquery_14(X0,X1) <=> select_15(X0,X1)) [rectify 16]
105. ! [X0,X1] : (main_query_2(X0,X1) <=> select_7(X0,X1)) [rectify 19]
106. ! [X0] : eq(X0,X0) [rectify 20]
107. ! [X0,X1] : (eq(X0,X1) => eq(X1,X0)) [rectify 21]
114. ! [X0,X1] : (neq(X0,X1) <=> ~eq(X0,X1)) [rectify 28]
123. ! [X0,X1,X2] : ((where_13(X1,X2) & eq(X1,X0)) => where_13(X0,X2)) [rectify 37]
176. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 91]
177. ! [X0,X1] : (eq(X1,X0) | ~eq(X0,X1)) [ennf transformation 107]
196. ! [X0,X1,X2] : (where_13(X0,X2) | (~where_13(X1,X2) | ~eq(X1,X0))) [ennf transformation 123]
197. ! [X0,X1,X2] : (where_13(X0,X2) | ~where_13(X1,X2) | ~eq(X1,X0)) [flattening 196]
302. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 176]
303. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
304. ! [X0,X1,X2,X3,X4] : ((from_6(X0,X1,X2,X3,X4) | ~tab2(X0,X1,X2,X3,X4)) & (tab2(X0,X1,X2,X3,X4) | ~from_6(X0,X1,X2,X3,X4))) [nnf transformation 92]
305. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_6(X0,X1,X2,X3,X4) | ~select_5(X0))) [nnf transformation 94]
306. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_6(X0,X5,X6,X7,X8) | ~select_5(X0))) [rectify 305]
307. ! [X0] : (? [X5,X6,X7,X8] : from_6(X0,X5,X6,X7,X8) => from_6(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
308. ! [X0] : ((select_5(X0) | ! [X1,X2,X3,X4] : ~from_6(X0,X1,X2,X3,X4)) & (from_6(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_5(X0))) [skolemisation 306,307]
309. ! [X0] : ((where_3_subquery_4(X0) | ~select_5(X0)) & (select_5(X0) | ~where_3_subquery_4(X0))) [nnf transformation 95]
310. ! [X0] : ((where_3(X0) | ~where_3_subquery_4(X0)) & (where_3_subquery_4(X0) | ~where_3(X0))) [nnf transformation 5]
311. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 6]
312. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 311]
313. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0) & from_2(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1)))) [choice axiom]
314. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1))) | ~select_1(X0,X1))) [skolemisation 312,313]
315. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 96]
316. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | (~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4))) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [nnf transformation 8]
317. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4)) & ((tab2(X5,X6,X7,X8,X9) & tab1(X0,X1,X2,X3,X4)) | ~from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) [flattening 316]
319. ! [X0,X1] : ((where_13(X0,X1) | ~neq(X0,X1)) & (neq(X0,X1) | ~where_13(X0,X1))) [nnf transformation 98]
320. ! [X0,X1] : ((select_11(X0,X1) | ! [X2,X3,X4,X5] : (~where_13(X0,X1) | ~from_12(X0,X2,X3,X4,X5))) & (? [X2,X3,X4,X5] : (where_13(X0,X1) & from_12(X0,X2,X3,X4,X5)) | ~select_11(X0,X1))) [nnf transformation 99]
321. ! [X0,X1] : ((select_11(X0,X1) | ! [X2,X3,X4,X5] : (~where_13(X0,X1) | ~from_12(X0,X2,X3,X4,X5))) & (? [X6,X7,X8,X9] : (where_13(X0,X1) & from_12(X0,X6,X7,X8,X9)) | ~select_11(X0,X1))) [rectify 320]
322. ! [X1,X0] : (? [X6,X7,X8,X9] : (where_13(X0,X1) & from_12(X0,X6,X7,X8,X9)) => (where_13(X0,X1) & from_12(X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)))) [choice axiom]
323. ! [X0,X1] : ((select_11(X0,X1) | ! [X2,X3,X4,X5] : (~where_13(X0,X1) | ~from_12(X0,X2,X3,X4,X5))) & ((where_13(X0,X1) & from_12(X0,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1),sK10(X0,X1))) | ~select_11(X0,X1))) [skolemisation 321,322]
324. ! [X0,X1] : ((where_9_subquery_10(X0,X1) | ~select_11(X0,X1)) & (select_11(X0,X1) | ~where_9_subquery_10(X0,X1))) [nnf transformation 100]
325. ! [X0,X1,X2,X3,X4] : ((from_16(X0,X1,X2,X3,X4) | ~tab2(X0,X1,X2,X3,X4)) & (tab2(X0,X1,X2,X3,X4) | ~from_16(X0,X1,X2,X3,X4))) [nnf transformation 101]
326. ! [X0,X1] : ((where_17(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_17(X0,X1))) [nnf transformation 102]
327. ! [X0,X1] : ((select_15(X0,X1) | ! [X2,X3,X4,X5] : (~where_17(X0,X1) | ~from_16(X0,X2,X3,X4,X5))) & (? [X2,X3,X4,X5] : (where_17(X0,X1) & from_16(X0,X2,X3,X4,X5)) | ~select_15(X0,X1))) [nnf transformation 103]
328. ! [X0,X1] : ((select_15(X0,X1) | ! [X2,X3,X4,X5] : (~where_17(X0,X1) | ~from_16(X0,X2,X3,X4,X5))) & (? [X6,X7,X8,X9] : (where_17(X0,X1) & from_16(X0,X6,X7,X8,X9)) | ~select_15(X0,X1))) [rectify 327]
329. ! [X1,X0] : (? [X6,X7,X8,X9] : (where_17(X0,X1) & from_16(X0,X6,X7,X8,X9)) => (where_17(X0,X1) & from_16(X0,sK11(X0,X1),sK12(X0,X1),sK13(X0,X1),sK14(X0,X1)))) [choice axiom]
330. ! [X0,X1] : ((select_15(X0,X1) | ! [X2,X3,X4,X5] : (~where_17(X0,X1) | ~from_16(X0,X2,X3,X4,X5))) & ((where_17(X0,X1) & from_16(X0,sK11(X0,X1),sK12(X0,X1),sK13(X0,X1),sK14(X0,X1))) | ~select_15(X0,X1))) [skolemisation 328,329]
331. ! [X0,X1] : ((where_9_subquery_14(X0,X1) | ~select_15(X0,X1)) & (select_15(X0,X1) | ~where_9_subquery_14(X0,X1))) [nnf transformation 104]
332. ! [X0] : ((where_9(X0) | (~where_9_subquery_14(X0,X0) | where_9_subquery_10(X0,X0))) & ((where_9_subquery_14(X0,X0) & ~where_9_subquery_10(X0,X0)) | ~where_9(X0))) [nnf transformation 17]
333. ! [X0] : ((where_9(X0) | ~where_9_subquery_14(X0,X0) | where_9_subquery_10(X0,X0)) & ((where_9_subquery_14(X0,X0) & ~where_9_subquery_10(X0,X0)) | ~where_9(X0))) [flattening 332]
334. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_9(X0) | ~from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & (? [X2,X3,X4,X5,X6,X7,X8,X9] : (where_9(X0) & from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9)) | ~select_7(X0,X1))) [nnf transformation 18]
335. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_9(X0) | ~from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_9(X0) & from_8(X0,X1,X10,X11,X12,X13,X14,X15,X16,X17)) | ~select_7(X0,X1))) [rectify 334]
336. ! [X1,X0] : (? [X10,X11,X12,X13,X14,X15,X16,X17] : (where_9(X0) & from_8(X0,X1,X10,X11,X12,X13,X14,X15,X16,X17)) => (where_9(X0) & from_8(X0,X1,sK15(X0,X1),sK16(X0,X1),sK17(X0,X1),sK18(X0,X1),sK19(X0,X1),sK20(X0,X1),sK21(X0,X1),sK22(X0,X1)))) [choice axiom]
337. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8,X9] : (~where_9(X0) | ~from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9))) & ((where_9(X0) & from_8(X0,X1,sK15(X0,X1),sK16(X0,X1),sK17(X0,X1),sK18(X0,X1),sK19(X0,X1),sK20(X0,X1),sK21(X0,X1),sK22(X0,X1))) | ~select_7(X0,X1))) [skolemisation 335,336]
338. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_7(X0,X1)) & (select_7(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 105]
341. ! [X0,X1] : ((neq(X0,X1) | eq(X0,X1)) & (~eq(X0,X1) | ~neq(X0,X1))) [nnf transformation 114]
342. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 302]
343. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK23,sK24) | ~main_query_1(sK23,sK24)) & (main_query_2(sK23,sK24) | main_query_1(sK23,sK24))) [choice axiom]
344. (~main_query_2(sK23,sK24) | ~main_query_1(sK23,sK24)) & (main_query_2(sK23,sK24) | main_query_1(sK23,sK24)) [skolemisation 342,343]
345. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 303]
346. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 303]
347. ~from_6(X0,X1,X2,X3,X4) | tab2(X0,X1,X2,X3,X4) [cnf transformation 304]
348. ~tab2(X0,X1,X2,X3,X4) | from_6(X0,X1,X2,X3,X4) [cnf transformation 304]
349. from_6(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_5(X0) [cnf transformation 308]
350. ~from_6(X0,X1,X2,X3,X4) | select_5(X0) [cnf transformation 308]
351. ~where_3_subquery_4(X0) | select_5(X0) [cnf transformation 309]
352. ~select_5(X0) | where_3_subquery_4(X0) [cnf transformation 309]
353. ~where_3(X0) | where_3_subquery_4(X0) [cnf transformation 310]
354. ~where_3_subquery_4(X0) | where_3(X0) [cnf transformation 310]
355. from_2(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1)) | ~select_1(X0,X1) [cnf transformation 314]
356. ~select_1(X0,X1) | where_3(X0) [cnf transformation 314]
357. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0,X1) [cnf transformation 314]
358. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 315]
359. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 315]
360. ~from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | tab1(X0,X1,X2,X3,X4) [cnf transformation 317]
362. from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~tab2(X5,X6,X7,X8,X9) | ~tab1(X0,X1,X2,X3,X4) [cnf transformation 317]
365. ~where_13(X0,X1) | neq(X0,X1) [cnf transformation 319]
368. ~select_11(X0,X1) | where_13(X0,X1) [cnf transformation 323]
370. ~where_9_subquery_10(X0,X1) | select_11(X0,X1) [cnf transformation 324]
372. ~from_16(X0,X1,X2,X3,X4) | tab2(X0,X1,X2,X3,X4) [cnf transformation 325]
373. ~tab2(X0,X1,X2,X3,X4) | from_16(X0,X1,X2,X3,X4) [cnf transformation 325]
375. ~eq(X0,X1) | where_17(X0,X1) [cnf transformation 326]
376. from_16(X0,sK11(X0,X1),sK12(X0,X1),sK13(X0,X1),sK14(X0,X1)) | ~select_15(X0,X1) [cnf transformation 330]
378. ~from_16(X0,X2,X3,X4,X5) | ~where_17(X0,X1) | select_15(X0,X1) [cnf transformation 330]
379. ~where_9_subquery_14(X0,X1) | select_15(X0,X1) [cnf transformation 331]
380. ~select_15(X0,X1) | where_9_subquery_14(X0,X1) [cnf transformation 331]
382. where_9_subquery_14(X0,X0) | ~where_9(X0) [cnf transformation 333]
383. ~where_9_subquery_14(X0,X0) | where_9(X0) | where_9_subquery_10(X0,X0) [cnf transformation 333]
384. from_8(X0,X1,sK15(X0,X1),sK16(X0,X1),sK17(X0,X1),sK18(X0,X1),sK19(X0,X1),sK20(X0,X1),sK21(X0,X1),sK22(X0,X1)) | ~select_7(X0,X1) [cnf transformation 337]
385. ~select_7(X0,X1) | where_9(X0) [cnf transformation 337]
386. ~from_8(X0,X1,X2,X3,X4,X5,X6,X7,X8,X9) | ~where_9(X0) | select_7(X0,X1) [cnf transformation 337]
387. ~main_query_2(X0,X1) | select_7(X0,X1) [cnf transformation 338]
388. ~select_7(X0,X1) | main_query_2(X0,X1) [cnf transformation 338]
389. eq(X0,X0) [cnf transformation 106]
390. ~eq(X0,X1) | eq(X1,X0) [cnf transformation 177]
397. ~eq(X0,X1) | ~neq(X0,X1) [cnf transformation 341]
398. eq(X0,X1) | neq(X0,X1) [cnf transformation 341]
406. ~eq(X1,X0) | ~where_13(X1,X2) | where_13(X0,X2) [cnf transformation 197]
459. main_query_2(sK23,sK24) | main_query_1(sK23,sK24) [cnf transformation 344]
460. ~main_query_2(sK23,sK24) | ~main_query_1(sK23,sK24) [cnf transformation 344]
462. 0 <=> main_query_1(sK23,sK24) [avatar definition]
463. main_query_1(sK23,sK24) <- {0} [avatar component clause 462]
465. 1 <=> ~main_query_1(sK23,sK24) [avatar definition]
466. ~main_query_1(sK23,sK24) <- {1} [avatar component clause 465]
468. 2 <=> main_query_2(sK23,sK24) [avatar definition]
469. main_query_2(sK23,sK24) <- {2} [avatar component clause 468]
471. 3 <=> ~main_query_2(sK23,sK24) [avatar definition]
472. ~main_query_2(sK23,sK24) <- {3} [avatar component clause 471]
473. ~1 | ~3 [avatar split clause 460,471,465]
474. 0 | 2 [avatar split clause 459,468,462]
475. where_17(X0,X0) [resolution 375,389]
477. select_15(X0,X0) | ~where_9(X0) [resolution 379,382]
480. select_7(sK23,sK24) <- {2} [resolution 387,469]
482. where_9(sK23) <- {2} [resolution 480,385]
486. ~neq(X0,X0) [resolution 397,389]
490. eq(X7,X6) | neq(X6,X7) [resolution 398,390]
516. ~where_13(X2,X3) | where_13(X4,X3) | neq(X4,X2) [resolution 406,490]
611. tab2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_5(X0) [resolution 349,347]
639. tab1(X0,X1,sK4(X0,X1),sK5(X0,X1),sK6(X0,X1)) | ~select_1(X0,X1) [resolution 355,345]
780. tab2(X0,sK11(X0,X1),sK12(X0,X1),sK13(X0,X1),sK14(X0,X1)) | ~select_15(X0,X1) [resolution 376,372]
794. ~tab2(X20,X21,X22,X23,X24) | ~tab1(X25,X26,X27,X28,X29) | ~where_9(X25) | select_7(X25,X26) [resolution 362,386]
796. 4 <=> ! [X25,X27,X29,X26,X28] : (~tab1(X25,X26,X27,X28,X29) | select_7(X25,X26) | ~where_9(X25)) [avatar definition]
797. ~tab1(X25,X26,X27,X28,X29) | select_7(X25,X26) | ~where_9(X25) <- {4} [avatar component clause 796]
799. 6 <=> ! [X20,X22,X21,X23,X24] : ~tab2(X20,X21,X22,X23,X24) [avatar definition]
800. ~tab2(X20,X21,X22,X23,X24) <- {6} [avatar component clause 799]
801. 4 | 6 [avatar split clause 794,799,796]
849. tab1(X0,X1,sK15(X0,X1),sK16(X0,X1),sK17(X0,X1)) | ~select_7(X0,X1) [resolution 384,360]
864. from_16(X2,sK0(X2),sK1(X2),sK2(X2),sK3(X2)) | ~select_5(X2) [resolution 611,373]
911. ~select_1(X17,X18) | select_7(X17,X18) | ~where_9(X17) <- {4} [resolution 639,797]
967. from_2(X0,X1,sK15(X0,X1),sK16(X0,X1),sK17(X0,X1)) | ~select_7(X0,X1) [resolution 849,346]
1031. from_6(X0,sK11(X0,X1),sK12(X0,X1),sK13(X0,X1),sK14(X0,X1)) | ~select_15(X0,X1) [resolution 780,348]
1840. ~where_17(X1,X2) | ~select_5(X1) | select_15(X1,X2) [resolution 864,378]
1993. 10 <=> select_5(sK23) [avatar definition]
1994. select_5(sK23) <- {10} [avatar component clause 1993]
1996. 11 <=> ~select_5(sK23) [avatar definition]
1997. ~select_5(sK23) <- {11} [avatar component clause 1996]
2076. ~select_7(X2,X3) | ~where_3(X2) | select_1(X2,X3) [resolution 967,357]
2111. ~select_15(X2,X3) | select_5(X2) [resolution 1031,350]
2117. ~where_9(X0) | select_5(X0) [resolution 2111,477]
2129. select_5(sK23) <- {2} [resolution 2117,482]
2130. $false <- {2, 11} [subsumption resolution 2129,1997]
2131. ~2 | 11 [avatar contradiction clause 2130]
2136. select_1(sK23,sK24) <- {0} [resolution 463,358]
2137. where_3(sK23) <- {0} [resolution 2136,356]
2143. select_7(sK23,sK24) | ~where_9(sK23) <- {0, 4} [resolution 2136,911]
2145. 16 <=> where_9(sK23) [avatar definition]
2148. 17 <=> ~where_9(sK23) [avatar definition]
2154. 18 <=> select_7(sK23,sK24) [avatar definition]
2155. select_7(sK23,sK24) <- {18} [avatar component clause 2154]
2156. ~17 | 18 | ~0 | ~4 [avatar split clause 2143,796,462,2154,2148]
2157. where_3_subquery_4(sK23) <- {0} [resolution 2137,353]
2169. select_5(sK23) <- {0} [resolution 2157,351]
2172. 10 | ~0 [avatar split clause 2169,462,1993]
2178. where_3_subquery_4(sK23) <- {10} [resolution 1994,352]
2321. where_3(sK23) <- {10} [resolution 2178,354]
2325. select_7(sK23,sK24) <- {2} [resolution 469,387]
2326. 18 | ~2 [avatar split clause 2325,468,2154]
2341. where_9(sK23) <- {18} [resolution 2155,385]
2344. 16 | ~18 [avatar split clause 2341,2154,2145]
2690. select_15(X0,X0) | ~select_5(X0) [resolution 1840,475]
3030. ~where_3(sK23) | select_1(sK23,sK24) <- {18} [resolution 2076,2155]
3039. select_1(sK23,sK24) <- {10, 18} [subsumption resolution 3030,2321]
3056. main_query_1(sK23,sK24) <- {10, 18} [resolution 3039,359]
3062. $false <- {1, 10, 18} [subsumption resolution 3056,466]
3063. 1 | ~10 | ~18 [avatar contradiction clause 3062]
3181. where_9_subquery_14(X1,X1) | ~select_5(X1) [resolution 2690,380]
3389. where_9_subquery_10(X5,X5) | where_9(X5) | ~select_5(X5) [resolution 3181,383]
4182. select_11(X9,X9) | ~select_5(X9) | where_9(X9) [resolution 3389,370]
4214. where_13(X0,X0) | where_9(X0) | ~select_5(X0) [resolution 4182,368]
4237. where_9(X3) | ~select_5(X3) | where_13(X4,X3) | neq(X4,X3) [resolution 4214,516]
4244. neq(X4,X3) | ~select_5(X3) | where_9(X3) [subsumption resolution 4237,365]
4284. ~select_5(X9) | where_9(X9) [resolution 4244,486]
4298. where_9(sK23) <- {10} [resolution 4284,1994]
4339. 16 | ~10 [avatar split clause 4298,1993,2145]
4347. ~select_5(X8) <- {6} [resolution 800,611]
4368. $false <- {6, 10} [resolution 4347,1994]
4372. ~6 | ~10 [avatar contradiction clause 4368]
4412. main_query_2(sK23,sK24) <- {18} [resolution 2155,388]
4414. $false <- {3, 18} [subsumption resolution 4412,472]
4415. 3 | ~18 [avatar contradiction clause 4414]
4416. $false [avatar sat refutation 473,474,801,2131,2156,2172,2326,2344,3063,4339,4372,4415]
% SZS output end Proof for query_20_query_30_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 6780
% Time elapsed: 0.414 s
% ------------------------------
% ------------------------------
