% Refutation found. Thanks to Tanya!
% SZS status Theorem for query13_query14_tptp4X
% SZS output start Proof for query13_query14_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> (~gt(X0,X1) & ~lt(X0,X1))) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [input]
8. ! [X5] : (where_6_subquery_7(X5) <=> select_8(X5)) [input]
9. ! [X0,X1] : (where_6(X0,X1) <=> (where_6_subquery_7(X0) & eq(X0,X1))) [input]
10. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
13. ! [X7,X8] : (eq(X7,X8) => eq(X8,X7)) [input]
16. ! [X7,X8] : ((lte(X8,X7) & lte(X7,X8)) <=> eq(X7,X8)) [input]
17. ! [X7,X8,X9] : ((lte(X8,X9) & lte(X7,X8)) => lte(X7,X9)) [input]
18. ! [X7,X8] : (lte(X8,X7) | lte(X7,X8)) [input]
19. ! [X7,X8] : (lt(X7,X8) <=> (~eq(X7,X8) & lte(X7,X8))) [input]
20. ! [X7,X8] : (gt(X7,X8) <=> lt(X8,X7)) [input]
60. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
61. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 60]
62. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
63. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 7]
64. ! [X0] : (where_6_subquery_7(X0) <=> select_8(X0)) [rectify 8]
65. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
67. ! [X0,X1] : (eq(X0,X1) => eq(X1,X0)) [rectify 13]
70. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 16]
71. ! [X0,X1,X2] : ((lte(X1,X2) & lte(X0,X1)) => lte(X0,X2)) [rectify 17]
72. ! [X0,X1] : (lte(X1,X0) | lte(X0,X1)) [rectify 18]
73. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 19]
74. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 20]
114. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 61]
115. ! [X0,X1] : (eq(X1,X0) | ~eq(X0,X1)) [ennf transformation 67]
118. ! [X0,X1,X2] : (lte(X0,X2) | (~lte(X1,X2) | ~lte(X0,X1))) [ennf transformation 71]
119. ! [X0,X1,X2] : (lte(X0,X2) | ~lte(X1,X2) | ~lte(X0,X1)) [flattening 118]
194. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 114]
195. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
196. ! [X0,X1] : ((where_3(X0,X1) | (gt(X0,X1) | lt(X0,X1))) & ((~gt(X0,X1) & ~lt(X0,X1)) | ~where_3(X0,X1))) [nnf transformation 2]
197. ! [X0,X1] : ((where_3(X0,X1) | gt(X0,X1) | lt(X0,X1)) & ((~gt(X0,X1) & ~lt(X0,X1)) | ~where_3(X0,X1))) [flattening 196]
198. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
199. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 198]
200. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
201. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 199,200]
202. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 62]
203. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
204. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 6]
205. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 63]
206. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 205]
207. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
208. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 206,207]
209. ! [X0] : ((where_6_subquery_7(X0) | ~select_8(X0)) & (select_8(X0) | ~where_6_subquery_7(X0))) [nnf transformation 64]
210. ! [X0,X1] : ((where_6(X0,X1) | (~where_6_subquery_7(X0) | ~eq(X0,X1))) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 9]
211. ! [X0,X1] : ((where_6(X0,X1) | ~where_6_subquery_7(X0) | ~eq(X0,X1)) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 210]
212. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 10]
213. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 212]
214. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
215. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 213,214]
216. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 65]
217. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 70]
218. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 217]
219. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 73]
220. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 219]
221. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 74]
222. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 194]
223. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
224. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 222,223]
225. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 195]
226. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 195]
227. ~lt(X0,X1) | ~where_3(X0,X1) [cnf transformation 197]
228. ~gt(X0,X1) | ~where_3(X0,X1) [cnf transformation 197]
229. gt(X0,X1) | where_3(X0,X1) | lt(X0,X1) [cnf transformation 197]
230. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 201]
231. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 201]
232. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 201]
233. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 202]
234. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 202]
235. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 203]
236. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 203]
238. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 204]
240. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 208]
242. ~select_8(X0) | where_6_subquery_7(X0) [cnf transformation 209]
243. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 211]
245. ~eq(X0,X1) | ~where_6_subquery_7(X0) | where_6(X0,X1) [cnf transformation 211]
246. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 215]
247. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 215]
248. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 215]
249. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 216]
250. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 216]
252. ~eq(X0,X1) | eq(X1,X0) [cnf transformation 115]
257. ~eq(X0,X1) | lte(X1,X0) [cnf transformation 218]
258. ~lte(X1,X2) | lte(X0,X2) | ~lte(X0,X1) [cnf transformation 119]
259. lte(X1,X0) | lte(X0,X1) [cnf transformation 72]
261. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 220]
262. ~lte(X0,X1) | eq(X0,X1) | lt(X0,X1) [cnf transformation 220]
263. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 221]
264. ~lt(X1,X0) | gt(X0,X1) [cnf transformation 221]
302. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 224]
303. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 224]
305. 0 <=> main_query_1(sK10,sK11) [avatar definition]
306. main_query_1(sK10,sK11) <- {0} [avatar component clause 305]
308. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
309. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 308]
311. 2 <=> main_query_2(sK10,sK11) [avatar definition]
312. main_query_2(sK10,sK11) <- {2} [avatar component clause 311]
314. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
315. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 314]
316. ~1 | ~3 [avatar split clause 303,314,308]
317. 0 | 2 [avatar split clause 302,311,305]
354. lt(X1,X0) | lt(X0,X1) | where_3(X0,X1) [resolution 229,263]
369. ~lte(X2,X4) | lte(X2,X3) | lte(X3,X4) [resolution 258,259]
376. eq(X1,X2) | lt(X1,X2) | lte(X2,X1) [resolution 262,259]
380. lte(X2,X1) | lt(X1,X2) [subsumption resolution 376,257]
382. ~lte(X4,X3) | lte(X4,X2) | lt(X2,X3) [resolution 380,258]
413. gt(X0,X1) | where_3(X1,X0) | lt(X0,X1) [resolution 354,264]
420. lte(X2,X3) | lte(X3,X4) | lt(X4,X2) [resolution 369,380]
444. 9 <=> ~lt(sK11,sK10) [avatar definition]
445. ~lt(sK11,sK10) <- {9} [avatar component clause 444]
447. 8 <=> lt(sK11,sK10) [avatar definition]
448. lt(sK11,sK10) <- {8} [avatar component clause 447]
453. 10 <=> where_3(sK11,sK10) [avatar definition]
454. where_3(sK11,sK10) <- {10} [avatar component clause 453]
458. select_4(sK10,sK11) <- {2} [resolution 312,249]
462. where_6(sK10,sK11) <- {2} [resolution 458,247]
463. eq(sK10,sK11) <- {2} [resolution 462,243]
470. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 230,225]
472. eq(sK11,sK10) <- {2} [resolution 463,252]
486. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 246,235]
492. ~lt(sK11,sK10) <- {2} [resolution 472,261]
501. ~9 | ~2 [avatar split clause 492,311,444]
504. lt(sK10,sK11) | where_3(sK10,sK11) <- {9} [resolution 445,354]
533. 13 <=> ~lt(sK10,sK11) [avatar definition]
534. ~lt(sK10,sK11) <- {13} [avatar component clause 533]
536. 12 <=> lt(sK10,sK11) [avatar definition]
537. lt(sK10,sK11) <- {12} [avatar component clause 536]
542. 14 <=> where_3(sK10,sK11) [avatar definition]
543. where_3(sK10,sK11) <- {14} [avatar component clause 542]
545. 14 | 12 | 9 [avatar split clause 504,444,536,542]
546. select_1(sK10,sK11) <- {0} [resolution 306,233]
549. where_3(sK10,sK11) <- {0} [resolution 546,231]
550. 14 | ~0 [avatar split clause 549,305,542]
554. select_4(sK10,sK11) <- {2} [resolution 312,249]
560. where_6(sK10,sK11) <- {2} [resolution 554,247]
562. eq(sK10,sK11) <- {2} [resolution 560,243]
574. ~lt(sK10,sK11) <- {2} [resolution 562,261]
583. $false <- {2, 12} [subsumption resolution 574,537]
584. ~2 | ~12 [avatar contradiction clause 583]
585. select_1(sK10,sK11) <- {0} [resolution 306,233]
592. ~where_3(sK10,sK11) <- {12} [resolution 537,227]
593. $false <- {12, 14} [subsumption resolution 592,543]
594. ~12 | ~14 [avatar contradiction clause 593]
632. 16 <=> where_6(sK10,sK11) [avatar definition]
633. where_6(sK10,sK11) <- {16} [avatar component clause 632]
635. 18 <=> where_6_subquery_7(sK10) [avatar definition]
636. where_6_subquery_7(sK10) <- {18} [avatar component clause 635]
664. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 470,236]
665. from_9(X4,X5,sK0(X4,X5),sK1(X4,X5),sK2(X4,X5)) | ~select_1(X4,X5) [resolution 470,238]
672. from_2(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [resolution 486,226]
732. where_3(X2,X3) | lt(X3,X2) | ~where_3(X3,X2) [resolution 413,228]
733. ~where_3(X3,X2) | where_3(X2,X3) [subsumption resolution 732,227]
743. where_3(sK11,sK10) <- {14} [resolution 733,543]
765. lte(sK10,X11) | lte(X11,sK11) <- {9} [resolution 420,445]
775. lte(X0,sK11) | lte(sK10,X1) | lt(X1,X0) <- {9} [resolution 765,382]
995. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 664,248]
1003. ~select_1(X2,X3) | select_8(X2) [resolution 665,240]
1009. select_8(sK10) <- {0} [resolution 1003,585]
1011. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 672,232]
1017. where_6_subquery_7(sK10) <- {0} [resolution 1009,242]
1020. 18 | ~0 [avatar split clause 1017,305,635]
1155. lte(sK10,sK11) | lte(sK10,sK11) <- {9} [resolution 775,445]
1158. lte(sK10,sK11) <- {9} [duplicate literal removal 1155]
1163. eq(sK10,sK11) | lt(sK10,sK11) <- {9} [resolution 1158,262]
1166. eq(sK10,sK11) <- {9, 13} [subsumption resolution 1163,534]
1187. ~where_6_subquery_7(sK10) | where_6(sK10,sK11) <- {9, 13} [resolution 1166,245]
1203. where_6(sK10,sK11) <- {9, 13, 18} [subsumption resolution 1187,636]
1204. 16 | 9 | 13 | ~18 [avatar split clause 1203,635,533,444,632]
1217. ~where_3(sK11,sK10) <- {8} [resolution 448,227]
1220. $false <- {8, 10} [subsumption resolution 1217,454]
1221. ~8 | ~10 [avatar contradiction clause 1220]
1302. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {16} [resolution 633,995]
1305. select_4(sK10,sK11) <- {0, 16} [subsumption resolution 1302,585]
1388. main_query_2(sK10,sK11) <- {0, 16} [resolution 1305,250]
1390. $false <- {0, 3, 16} [subsumption resolution 1388,315]
1391. ~0 | 3 | ~16 [avatar contradiction clause 1390]
1392. 10 | ~14 [avatar split clause 743,542,453]
1404. select_4(sK10,sK11) <- {2} [resolution 312,249]
1406. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 1404,1011]
1412. select_1(sK10,sK11) <- {2, 14} [subsumption resolution 1406,543]
1432. main_query_1(sK10,sK11) <- {2, 14} [resolution 1412,234]
1434. $false <- {1, 2, 14} [subsumption resolution 1432,309]
1435. 1 | ~2 | ~14 [avatar contradiction clause 1434]
1436. $false [avatar sat refutation 316,317,501,545,550,584,594,1020,1204,1221,1391,1392,1435]
% SZS output end Proof for query13_query14_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.013 s
% ------------------------------
% ------------------------------
