% Refutation found. Thanks to Tanya!
% SZS status Theorem for query11_query12_tptp4X
% SZS output start Proof for query11_query12_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8] : (where_6(X7,X8) <=> (lte(X8,X7) & lte(X7,X8))) [input]
7. ! [X7,X8] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X7,X8) & from_5(X7,X8,X9,X10,X11))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
12. ! [X12] : lte(X12,X12) [input]
13. ! [X12,X13] : ((lte(X13,X12) & lte(X12,X13)) <=> eq(X12,X13)) [input]
28. ! [X12,X7,X8] : ((where_6(X7,X8) & eq(X8,X12)) => where_6(X7,X12)) [input]
55. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
56. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 55]
57. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
58. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
59. ! [X0,X1] : (where_6(X0,X1) <=> (lte(X1,X0) & lte(X0,X1))) [rectify 6]
60. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
61. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
65. ! [X0] : lte(X0,X0) [rectify 12]
66. ! [X0,X1] : ((lte(X1,X0) & lte(X0,X1)) <=> eq(X0,X1)) [rectify 13]
81. ! [X0,X1,X2] : ((where_6(X1,X2) & eq(X2,X0)) => where_6(X1,X0)) [rectify 28]
108. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 56]
130. ! [X0,X1,X2] : (where_6(X1,X0) | (~where_6(X1,X2) | ~eq(X2,X0))) [ennf transformation 81]
131. ! [X0,X1,X2] : (where_6(X1,X0) | ~where_6(X1,X2) | ~eq(X2,X0)) [flattening 130]
184. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 108]
185. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
186. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
187. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
188. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 187]
189. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
190. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 188,189]
191. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 57]
192. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 58]
193. ! [X0,X1] : ((where_6(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1))) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 59]
194. ! [X0,X1] : ((where_6(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1)) & ((lte(X1,X0) & lte(X0,X1)) | ~where_6(X0,X1))) [flattening 193]
195. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 60]
196. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 195]
197. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
198. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 196,197]
199. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 61]
200. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | (~lte(X1,X0) | ~lte(X0,X1)))) [nnf transformation 66]
201. ! [X0,X1] : (((lte(X1,X0) & lte(X0,X1)) | ~eq(X0,X1)) & (eq(X0,X1) | ~lte(X1,X0) | ~lte(X0,X1))) [flattening 200]
202. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 184]
203. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
204. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 202,203]
205. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 185]
206. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 185]
207. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 186]
208. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 186]
209. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 190]
210. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 190]
211. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 190]
212. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 191]
213. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 191]
214. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 192]
215. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 192]
216. ~where_6(X0,X1) | lte(X0,X1) [cnf transformation 194]
217. ~where_6(X0,X1) | lte(X1,X0) [cnf transformation 194]
218. ~lte(X1,X0) | where_6(X0,X1) | ~lte(X0,X1) [cnf transformation 194]
219. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 198]
220. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 198]
221. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 198]
222. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 199]
223. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 199]
227. lte(X0,X0) [cnf transformation 65]
228. ~lte(X1,X0) | eq(X0,X1) | ~lte(X0,X1) [cnf transformation 201]
241. ~where_6(X1,X2) | where_6(X1,X0) | ~eq(X2,X0) [cnf transformation 131]
268. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 204]
269. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 204]
271. 0 <=> main_query_1(sK6,sK7) [avatar definition]
272. main_query_1(sK6,sK7) <- {0} [avatar component clause 271]
274. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
275. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 274]
277. 2 <=> main_query_2(sK6,sK7) [avatar definition]
278. main_query_2(sK6,sK7) <- {2} [avatar component clause 277]
280. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
281. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 280]
282. ~1 | ~3 [avatar split clause 269,280,274]
283. 0 | 2 [avatar split clause 268,277,271]
286. select_4(sK6,sK7) <- {2} [resolution 222,278]
287. where_6(sK6,sK7) <- {2} [resolution 286,220]
289. lte(sK7,sK6) <- {2} [resolution 287,217]
290. lte(sK6,sK7) <- {2} [resolution 287,216]
295. where_6(X0,X0) | ~lte(X0,X0) [resolution 218,227]
300. where_6(X0,X0) [subsumption resolution 295,227]
310. eq(sK6,sK7) | ~lte(sK6,sK7) <- {2} [resolution 228,289]
312. eq(sK6,sK7) <- {2} [subsumption resolution 310,290]
318. where_3(sK6,sK7) <- {2} [resolution 312,208]
350. ~eq(X1,X2) | where_6(X1,X2) [resolution 241,300]
398. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 209,205]
402. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 219,214]
515. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 398,215]
532. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 402,206]
906. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 515,221]
916. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 532,211]
929. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 916,286]
933. select_1(sK6,sK7) <- {2} [subsumption resolution 929,318]
946. main_query_1(sK6,sK7) <- {2} [resolution 933,213]
948. $false <- {1, 2} [subsumption resolution 946,275]
949. 1 | ~2 [avatar contradiction clause 948]
952. select_1(sK6,sK7) <- {0} [resolution 272,212]
956. where_3(sK6,sK7) <- {0} [resolution 952,210]
957. eq(sK6,sK7) <- {0} [resolution 956,207]
967. where_6(sK6,sK7) <- {0} [resolution 957,350]
1026. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 967,906]
1031. select_4(sK6,sK7) <- {0} [subsumption resolution 1026,952]
1095. main_query_2(sK6,sK7) <- {0} [resolution 1031,223]
1097. $false <- {0, 3} [subsumption resolution 1095,281]
1098. ~0 | 3 [avatar contradiction clause 1097]
1099. $false [avatar sat refutation 282,283,949,1098]
% SZS output end Proof for query11_query12_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.137 s
% ------------------------------
% ------------------------------
