% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_11_query_12_tptp4X
% SZS output start Proof for query_11_query_12_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8] : (where_6(X7,X8) <=> eq(X7,X8)) [input]
7. ! [X7,X8] : (select_4(X7,X8) <=> ? [X9,X10,X11] : (where_6(X7,X8) & from_5(X7,X8,X9,X10,X11))) [input]
8. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
51. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
52. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 51]
53. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
54. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
55. ! [X0,X1] : (where_6(X0,X1) <=> eq(X0,X1)) [rectify 6]
56. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 7]
57. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 8]
100. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 52]
170. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 100]
171. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
172. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
173. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
174. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 173]
175. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
176. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 174,175]
177. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 53]
178. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 54]
179. ! [X0,X1] : ((where_6(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_6(X0,X1))) [nnf transformation 55]
180. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 56]
181. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 180]
182. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
183. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_4(X0,X1))) [skolemisation 181,182]
184. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 57]
187. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 170]
188. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
189. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 187,188]
190. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 171]
191. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 171]
192. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 172]
193. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 172]
194. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 176]
195. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 176]
196. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 176]
197. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 177]
198. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 177]
199. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 178]
200. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 178]
201. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 179]
202. ~eq(X0,X1) | where_6(X0,X1) [cnf transformation 179]
203. from_5(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [cnf transformation 183]
204. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 183]
205. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 183]
206. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 184]
207. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 184]
248. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 189]
249. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 189]
251. 0 <=> main_query_1(sK6,sK7) [avatar definition]
252. main_query_1(sK6,sK7) <- {0} [avatar component clause 251]
254. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
255. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 254]
257. 2 <=> main_query_2(sK6,sK7) [avatar definition]
258. main_query_2(sK6,sK7) <- {2} [avatar component clause 257]
260. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
261. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 260]
262. ~1 | ~3 [avatar split clause 249,260,254]
263. 0 | 2 [avatar split clause 248,257,251]
268. select_4(sK6,sK7) <- {2} [resolution 206,258]
270. where_6(sK6,sK7) <- {2} [resolution 268,204]
272. eq(sK6,sK7) <- {2} [resolution 270,201]
277. where_3(sK6,sK7) <- {2} [resolution 272,193]
353. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 194,190]
357. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_4(X2,X3) [resolution 203,199]
408. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 353,200]
414. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_4(X0,X1) [resolution 357,191]
422. ~where_6(X2,X3) | ~select_1(X2,X3) | select_4(X2,X3) [resolution 408,205]
453. ~select_4(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 414,196]
466. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 453,268]
470. select_1(sK6,sK7) <- {2} [subsumption resolution 466,277]
483. main_query_1(sK6,sK7) <- {2} [resolution 470,198]
485. $false <- {1, 2} [subsumption resolution 483,255]
486. 1 | ~2 [avatar contradiction clause 485]
489. select_1(sK6,sK7) <- {0} [resolution 252,197]
493. where_3(sK6,sK7) <- {0} [resolution 489,195]
494. eq(sK6,sK7) <- {0} [resolution 493,192]
497. where_6(sK6,sK7) <- {0} [resolution 494,202]
507. ~select_1(sK6,sK7) | select_4(sK6,sK7) <- {0} [resolution 497,422]
511. select_4(sK6,sK7) <- {0} [subsumption resolution 507,489]
586. main_query_2(sK6,sK7) <- {0} [resolution 511,207]
588. $false <- {0, 3} [subsumption resolution 586,261]
589. ~0 | 3 [avatar contradiction clause 588]
590. $false [avatar sat refutation 262,263,486,589]
% SZS output end Proof for query_11_query_12_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
