% Refutation found. Thanks to Tanya!
% SZS status Theorem for query11_query14_tptp4X
% SZS output start Proof for query11_query14_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [input]
5. ! [X0,X1,X5,X6,X7] : (from_5(X0,X1,X5,X6,X7) <=> tab1(X0,X1,X5,X6,X7)) [input]
6. ! [X0,X8,X5,X6,X7] : (from_9(X0,X8,X5,X6,X7) <=> tab1(X0,X8,X5,X6,X7)) [input]
7. ! [X0] : (select_8(X0) <=> ? [X8,X5,X6,X7] : ($true & from_9(X0,X8,X5,X6,X7))) [input]
8. ! [X0] : (where_6_subquery_7(X0) <=> select_8(X0)) [input]
9. ! [X0,X1] : (where_6(X0,X1) <=> (where_6_subquery_7(X0) & eq(X0,X1))) [input]
10. ! [X0,X1] : (select_4(X0,X1) <=> ? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7))) [input]
11. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [input]
62. ! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [input]
63. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [negated conjecture 62]
64. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
65. ! [X0,X1,X2,X3,X4] : (from_9(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 6]
66. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : ($true & from_9(X0,X1,X2,X3,X4))) [rectify 7]
67. ! [X0] : (select_8(X0) <=> ? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4)) [true and false elimination 66]
68. ! [X0,X1] : (select_4(X0,X1) <=> ? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4))) [rectify 10]
204. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 63]
205. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
206. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
207. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
208. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 207]
209. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
210. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 208,209]
211. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 4]
212. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 64]
213. ! [X0,X1,X2,X3,X4] : ((from_9(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_9(X0,X1,X2,X3,X4))) [nnf transformation 65]
214. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_9(X0,X1,X2,X3,X4) | ~select_8(X0))) [nnf transformation 67]
215. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) | ~select_8(X0))) [rectify 214]
216. ! [X0] : (? [X5,X6,X7,X8] : from_9(X0,X5,X6,X7,X8) => from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0))) [choice axiom]
217. ! [X0] : ((select_8(X0) | ! [X1,X2,X3,X4] : ~from_9(X0,X1,X2,X3,X4)) & (from_9(X0,sK3(X0),sK4(X0),sK5(X0),sK6(X0)) | ~select_8(X0))) [skolemisation 215,216]
218. ! [X0] : ((where_6_subquery_7(X0) | ~select_8(X0)) & (select_8(X0) | ~where_6_subquery_7(X0))) [nnf transformation 8]
219. ! [X0,X1] : ((where_6(X0,X1) | (~where_6_subquery_7(X0) | ~eq(X0,X1))) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [nnf transformation 9]
220. ! [X0,X1] : ((where_6(X0,X1) | ~where_6_subquery_7(X0) | ~eq(X0,X1)) & ((where_6_subquery_7(X0) & eq(X0,X1)) | ~where_6(X0,X1))) [flattening 219]
221. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_6(X0,X1) & from_5(X0,X1,X2,X3,X4)) | ~select_4(X0,X1))) [nnf transformation 68]
222. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) | ~select_4(X0,X1))) [rectify 221]
223. ! [X1,X0] : (? [X5,X6,X7] : (where_6(X0,X1) & from_5(X0,X1,X5,X6,X7)) => (where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
224. ! [X0,X1] : ((select_4(X0,X1) | ! [X2,X3,X4] : (~where_6(X0,X1) | ~from_5(X0,X1,X2,X3,X4))) & ((where_6(X0,X1) & from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X0,X1))) [skolemisation 222,223]
225. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 11]
228. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 204]
229. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11))) [choice axiom]
230. (~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11)) & (main_query_2(sK10,sK11) | main_query_1(sK10,sK11)) [skolemisation 228,229]
231. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 205]
232. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 205]
233. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 206]
234. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 206]
235. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 210]
236. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 210]
237. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 210]
238. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 211]
239. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 211]
240. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 212]
241. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 212]
243. ~tab1(X0,X1,X2,X3,X4) | from_9(X0,X1,X2,X3,X4) [cnf transformation 213]
245. ~from_9(X0,X1,X2,X3,X4) | select_8(X0) [cnf transformation 217]
247. ~select_8(X0) | where_6_subquery_7(X0) [cnf transformation 218]
248. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 220]
250. ~eq(X0,X1) | ~where_6_subquery_7(X0) | where_6(X0,X1) [cnf transformation 220]
251. from_5(X0,X1,sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X0,X1) [cnf transformation 224]
252. ~select_4(X0,X1) | where_6(X0,X1) [cnf transformation 224]
253. ~from_5(X0,X1,X2,X3,X4) | ~where_6(X0,X1) | select_4(X0,X1) [cnf transformation 224]
254. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 225]
255. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 225]
304. main_query_2(sK10,sK11) | main_query_1(sK10,sK11) [cnf transformation 230]
305. ~main_query_2(sK10,sK11) | ~main_query_1(sK10,sK11) [cnf transformation 230]
347. 0 <=> main_query_1(sK10,sK11) [avatar definition]
348. main_query_1(sK10,sK11) <- {0} [avatar component clause 347]
350. 1 <=> ~main_query_1(sK10,sK11) [avatar definition]
351. ~main_query_1(sK10,sK11) <- {1} [avatar component clause 350]
353. 2 <=> main_query_2(sK10,sK11) [avatar definition]
354. main_query_2(sK10,sK11) <- {2} [avatar component clause 353]
356. 3 <=> ~main_query_2(sK10,sK11) [avatar definition]
357. ~main_query_2(sK10,sK11) <- {3} [avatar component clause 356]
358. ~1 | ~3 [avatar split clause 305,356,350]
359. 0 | 2 [avatar split clause 304,353,347]
362. select_4(sK10,sK11) <- {2} [resolution 254,354]
363. where_6(sK10,sK11) <- {2} [resolution 362,252]
365. eq(sK10,sK11) <- {2} [resolution 363,248]
375. where_3(sK10,sK11) <- {2} [resolution 365,234]
423. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 235,231]
427. tab1(X2,X3,sK7(X2,X3),sK8(X2,X3),sK9(X2,X3)) | ~select_4(X2,X3) [resolution 251,240]
428. from_9(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [resolution 423,243]
429. from_5(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 423,241]
435. from_2(X4,X5,sK7(X4,X5),sK8(X4,X5),sK9(X4,X5)) | ~select_4(X4,X5) [resolution 427,232]
437. ~select_1(X2,X3) | select_8(X2) [resolution 428,245]
438. ~where_6(X0,X1) | ~select_1(X0,X1) | select_4(X0,X1) [resolution 429,253]
445. ~select_4(X0,X1) | ~where_3(X0,X1) | select_1(X0,X1) [resolution 435,237]
447. ~where_3(sK10,sK11) | select_1(sK10,sK11) <- {2} [resolution 445,362]
448. select_1(sK10,sK11) <- {2} [subsumption resolution 447,375]
450. main_query_1(sK10,sK11) <- {2} [resolution 448,239]
452. $false <- {1, 2} [subsumption resolution 450,351]
453. 1 | ~2 [avatar contradiction clause 452]
454. select_1(sK10,sK11) <- {0} [resolution 348,238]
455. select_8(sK10) <- {0} [resolution 454,437]
457. where_3(sK10,sK11) <- {0} [resolution 454,236]
458. where_6_subquery_7(sK10) <- {0} [resolution 455,247]
459. eq(sK10,sK11) <- {0} [resolution 457,233]
462. ~where_6_subquery_7(sK10) | where_6(sK10,sK11) <- {0} [resolution 459,250]
467. where_6(sK10,sK11) <- {0} [subsumption resolution 462,458]
478. ~select_1(sK10,sK11) | select_4(sK10,sK11) <- {0} [resolution 467,438]
481. select_4(sK10,sK11) <- {0} [subsumption resolution 478,454]
487. main_query_2(sK10,sK11) <- {0} [resolution 481,255]
489. $false <- {0, 3} [subsumption resolution 487,357]
490. ~0 | 3 [avatar contradiction clause 489]
491. $false [avatar sat refutation 358,359,453,490]
% SZS output end Proof for query11_query14_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.011 s
% ------------------------------
% ------------------------------
