% Refutation found. Thanks to Tanya!
% SZS status Theorem for query04_query09_tptp4X
% SZS output start Proof for query04_query09_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [input]
3. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
5. ! [X0,X5,X6,X7,X8] : (from_8(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
6. ! [X0] : (select_7(X0) <=> ? [X5,X6,X7,X8] : ($true & from_8(X0,X5,X6,X7,X8))) [input]
7. ! [X0] : (subquery_from_6(X0) <=> select_7(X0)) [input]
8. ! [X9] : (from_5(X9) <=> subquery_from_6(X9)) [input]
9. ! [X9] : (select_4(X9) <=> ($true & from_5(X9))) [input]
10. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [input]
11. ! [X10] : eq(X10,X10) [input]
45. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
46. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 45]
47. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
48. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : ($true & from_8(X0,X1,X2,X3,X4))) [rectify 6]
49. ! [X0] : (select_7(X0) <=> ? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4)) [true and false elimination 48]
50. ! [X0] : (from_5(X0) <=> subquery_from_6(X0)) [rectify 8]
51. ! [X0] : (select_4(X0) <=> ($true & from_5(X0))) [rectify 9]
52. ! [X0] : (select_4(X0) <=> from_5(X0)) [true and false elimination 51]
53. ! [X0] : eq(X0,X0) [rectify 11]
138. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 46]
139. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
140. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 2]
141. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_3(X0) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0))) [nnf transformation 3]
142. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) | ~select_1(X0))) [rectify 141]
143. ! [X0] : (? [X5,X6,X7,X8] : (where_3(X0) & from_2(X0,X5,X6,X7,X8)) => (where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
144. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : (~where_3(X0) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0) & from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) | ~select_1(X0))) [skolemisation 142,143]
145. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 4]
146. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 47]
147. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_8(X0,X1,X2,X3,X4) | ~select_7(X0))) [nnf transformation 49]
148. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) | ~select_7(X0))) [rectify 147]
149. ! [X0] : (? [X5,X6,X7,X8] : from_8(X0,X5,X6,X7,X8) => from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
150. ! [X0] : ((select_7(X0) | ! [X1,X2,X3,X4] : ~from_8(X0,X1,X2,X3,X4)) & (from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0))) [skolemisation 148,149]
151. ! [X0] : ((subquery_from_6(X0) | ~select_7(X0)) & (select_7(X0) | ~subquery_from_6(X0))) [nnf transformation 7]
152. ! [X0] : ((from_5(X0) | ~subquery_from_6(X0)) & (subquery_from_6(X0) | ~from_5(X0))) [nnf transformation 50]
153. ! [X0] : ((select_4(X0) | ~from_5(X0)) & (from_5(X0) | ~select_4(X0))) [nnf transformation 52]
154. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 10]
157. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 138]
158. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
159. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 157,158]
160. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 139]
161. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 139]
163. ~eq(X0,X0) | where_3(X0) [cnf transformation 140]
164. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 144]
166. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0) | select_1(X0) [cnf transformation 144]
167. ~main_query_1(X0) | select_1(X0) [cnf transformation 145]
168. ~select_1(X0) | main_query_1(X0) [cnf transformation 145]
169. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 146]
170. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 146]
171. from_8(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [cnf transformation 150]
172. ~from_8(X0,X1,X2,X3,X4) | select_7(X0) [cnf transformation 150]
173. ~subquery_from_6(X0) | select_7(X0) [cnf transformation 151]
174. ~select_7(X0) | subquery_from_6(X0) [cnf transformation 151]
175. ~from_5(X0) | subquery_from_6(X0) [cnf transformation 152]
176. ~subquery_from_6(X0) | from_5(X0) [cnf transformation 152]
177. ~select_4(X0) | from_5(X0) [cnf transformation 153]
178. ~from_5(X0) | select_4(X0) [cnf transformation 153]
179. ~main_query_2(X0) | select_4(X0) [cnf transformation 154]
180. ~select_4(X0) | main_query_2(X0) [cnf transformation 154]
181. eq(X0,X0) [cnf transformation 53]
213. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 159]
214. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 159]
216. 0 <=> main_query_1(sK8) [avatar definition]
217. main_query_1(sK8) <- {0} [avatar component clause 216]
219. 1 <=> ~main_query_1(sK8) [avatar definition]
220. ~main_query_1(sK8) <- {1} [avatar component clause 219]
222. 2 <=> main_query_2(sK8) [avatar definition]
223. main_query_2(sK8) <- {2} [avatar component clause 222]
225. 3 <=> ~main_query_2(sK8) [avatar definition]
226. ~main_query_2(sK8) <- {3} [avatar component clause 225]
227. ~1 | ~3 [avatar split clause 214,225,219]
228. 0 | 2 [avatar split clause 213,222,216]
229. select_4(sK8) <- {2} [resolution 179,223]
230. from_5(sK8) <- {2} [resolution 229,177]
233. subquery_from_6(sK8) <- {2} [resolution 230,175]
235. select_7(sK8) <- {2} [resolution 233,173]
237. where_3(X1) [resolution 163,181]
272. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 164,160]
274. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_7(X0) [resolution 171,169]
276. from_8(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 272,170]
279. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_7(X1) [resolution 274,161]
281. ~select_1(X1) | select_7(X1) [resolution 276,172]
283. ~select_7(X1) | ~where_3(X1) | select_1(X1) [resolution 279,166]
284. ~select_7(X1) | select_1(X1) [subsumption resolution 283,237]
285. select_1(sK8) <- {2} [resolution 284,235]
289. main_query_1(sK8) <- {2} [resolution 285,168]
291. $false <- {1, 2} [subsumption resolution 289,220]
292. 1 | ~2 [avatar contradiction clause 291]
293. select_1(sK8) <- {0} [resolution 217,167]
296. select_7(sK8) <- {0} [resolution 293,281]
302. subquery_from_6(sK8) <- {0} [resolution 296,174]
305. from_5(sK8) <- {0} [resolution 302,176]
309. select_4(sK8) <- {0} [resolution 305,178]
313. main_query_2(sK8) <- {0} [resolution 309,180]
315. $false <- {0, 3} [subsumption resolution 313,226]
316. ~0 | 3 [avatar contradiction clause 315]
317. $false [avatar sat refutation 227,228,292,316]
% SZS output end Proof for query04_query09_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
