% Refutation found. Thanks to Tanya!
% SZS status Theorem for query_12_query_16_tptp4X
% SZS output start Proof for query_12_query_16_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [input]
3. ! [X0,X1] : (select_1(X0,X1) <=> ? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5,X6] : (main_query_1(X5,X6) <=> select_1(X5,X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_8(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X7,X8] : (where_9(X7,X8) <=> eq(X7,X8)) [input]
7. ! [X7,X8] : (select_7(X7,X8) <=> ? [X9,X10,X11] : (where_9(X7,X8) & from_8(X7,X8,X9,X10,X11))) [input]
8. ! [X5,X6] : (subquery_from_6(X5,X6) <=> select_7(X5,X6)) [input]
9. ! [X12,X13] : (from_5(X12,X13) <=> subquery_from_6(X12,X13)) [input]
10. ! [X12,X13] : (select_4(X12,X13) <=> ($true & from_5(X12,X13))) [input]
11. ! [X5,X6] : (main_query_2(X5,X6) <=> select_4(X5,X6)) [input]
60. ! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [input]
61. ~! [X5,X6] : (main_query_1(X5,X6) <=> main_query_2(X5,X6)) [negated conjecture 60]
62. ! [X0,X1] : (main_query_1(X0,X1) <=> select_1(X0,X1)) [rectify 4]
63. ! [X0,X1,X2,X3,X4] : (from_8(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
64. ! [X0,X1] : (where_9(X0,X1) <=> eq(X0,X1)) [rectify 6]
65. ! [X0,X1] : (select_7(X0,X1) <=> ? [X2,X3,X4] : (where_9(X0,X1) & from_8(X0,X1,X2,X3,X4))) [rectify 7]
66. ! [X0,X1] : (subquery_from_6(X0,X1) <=> select_7(X0,X1)) [rectify 8]
67. ! [X0,X1] : (from_5(X0,X1) <=> subquery_from_6(X0,X1)) [rectify 9]
68. ! [X0,X1] : (select_4(X0,X1) <=> ($true & from_5(X0,X1))) [rectify 10]
69. ! [X0,X1] : (select_4(X0,X1) <=> from_5(X0,X1)) [true and false elimination 68]
70. ! [X0,X1] : (main_query_2(X0,X1) <=> select_4(X0,X1)) [rectify 11]
119. ~! [X0,X1] : (main_query_1(X0,X1) <=> main_query_2(X0,X1)) [rectify 61]
201. ? [X0,X1] : (main_query_1(X0,X1) <~> main_query_2(X0,X1)) [ennf transformation 119]
202. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
203. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 2]
204. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_3(X0,X1) & from_2(X0,X1,X2,X3,X4)) | ~select_1(X0,X1))) [nnf transformation 3]
205. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) | ~select_1(X0,X1))) [rectify 204]
206. ! [X1,X0] : (? [X5,X6,X7] : (where_3(X0,X1) & from_2(X0,X1,X5,X6,X7)) => (where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)))) [choice axiom]
207. ! [X0,X1] : ((select_1(X0,X1) | ! [X2,X3,X4] : (~where_3(X0,X1) | ~from_2(X0,X1,X2,X3,X4))) & ((where_3(X0,X1) & from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1))) | ~select_1(X0,X1))) [skolemisation 205,206]
208. ! [X0,X1] : ((main_query_1(X0,X1) | ~select_1(X0,X1)) & (select_1(X0,X1) | ~main_query_1(X0,X1))) [nnf transformation 62]
209. ! [X0,X1,X2,X3,X4] : ((from_8(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_8(X0,X1,X2,X3,X4))) [nnf transformation 63]
210. ! [X0,X1] : ((where_9(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_9(X0,X1))) [nnf transformation 64]
211. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : (~where_9(X0,X1) | ~from_8(X0,X1,X2,X3,X4))) & (? [X2,X3,X4] : (where_9(X0,X1) & from_8(X0,X1,X2,X3,X4)) | ~select_7(X0,X1))) [nnf transformation 65]
212. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : (~where_9(X0,X1) | ~from_8(X0,X1,X2,X3,X4))) & (? [X5,X6,X7] : (where_9(X0,X1) & from_8(X0,X1,X5,X6,X7)) | ~select_7(X0,X1))) [rectify 211]
213. ! [X1,X0] : (? [X5,X6,X7] : (where_9(X0,X1) & from_8(X0,X1,X5,X6,X7)) => (where_9(X0,X1) & from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
214. ! [X0,X1] : ((select_7(X0,X1) | ! [X2,X3,X4] : (~where_9(X0,X1) | ~from_8(X0,X1,X2,X3,X4))) & ((where_9(X0,X1) & from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1))) | ~select_7(X0,X1))) [skolemisation 212,213]
215. ! [X0,X1] : ((subquery_from_6(X0,X1) | ~select_7(X0,X1)) & (select_7(X0,X1) | ~subquery_from_6(X0,X1))) [nnf transformation 66]
216. ! [X0,X1] : ((from_5(X0,X1) | ~subquery_from_6(X0,X1)) & (subquery_from_6(X0,X1) | ~from_5(X0,X1))) [nnf transformation 67]
217. ! [X0,X1] : ((select_4(X0,X1) | ~from_5(X0,X1)) & (from_5(X0,X1) | ~select_4(X0,X1))) [nnf transformation 69]
218. ! [X0,X1] : ((main_query_2(X0,X1) | ~select_4(X0,X1)) & (select_4(X0,X1) | ~main_query_2(X0,X1))) [nnf transformation 70]
221. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) [nnf transformation 201]
222. ? [X0,X1] : ((~main_query_2(X0,X1) | ~main_query_1(X0,X1)) & (main_query_2(X0,X1) | main_query_1(X0,X1))) => ((~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7))) [choice axiom]
223. (~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7)) & (main_query_2(sK6,sK7) | main_query_1(sK6,sK7)) [skolemisation 221,222]
224. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 202]
225. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 202]
226. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 203]
227. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 203]
228. from_2(X0,X1,sK0(X0,X1),sK1(X0,X1),sK2(X0,X1)) | ~select_1(X0,X1) [cnf transformation 207]
229. ~select_1(X0,X1) | where_3(X0,X1) [cnf transformation 207]
230. ~from_2(X0,X1,X2,X3,X4) | ~where_3(X0,X1) | select_1(X0,X1) [cnf transformation 207]
231. ~main_query_1(X0,X1) | select_1(X0,X1) [cnf transformation 208]
232. ~select_1(X0,X1) | main_query_1(X0,X1) [cnf transformation 208]
233. ~from_8(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 209]
234. ~tab1(X0,X1,X2,X3,X4) | from_8(X0,X1,X2,X3,X4) [cnf transformation 209]
235. ~where_9(X0,X1) | eq(X0,X1) [cnf transformation 210]
236. ~eq(X0,X1) | where_9(X0,X1) [cnf transformation 210]
237. from_8(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1) [cnf transformation 214]
238. ~select_7(X0,X1) | where_9(X0,X1) [cnf transformation 214]
239. ~from_8(X0,X1,X2,X3,X4) | ~where_9(X0,X1) | select_7(X0,X1) [cnf transformation 214]
240. ~subquery_from_6(X0,X1) | select_7(X0,X1) [cnf transformation 215]
241. ~select_7(X0,X1) | subquery_from_6(X0,X1) [cnf transformation 215]
242. ~from_5(X0,X1) | subquery_from_6(X0,X1) [cnf transformation 216]
243. ~subquery_from_6(X0,X1) | from_5(X0,X1) [cnf transformation 216]
244. ~select_4(X0,X1) | from_5(X0,X1) [cnf transformation 217]
245. ~from_5(X0,X1) | select_4(X0,X1) [cnf transformation 217]
246. ~main_query_2(X0,X1) | select_4(X0,X1) [cnf transformation 218]
247. ~select_4(X0,X1) | main_query_2(X0,X1) [cnf transformation 218]
294. main_query_2(sK6,sK7) | main_query_1(sK6,sK7) [cnf transformation 223]
295. ~main_query_2(sK6,sK7) | ~main_query_1(sK6,sK7) [cnf transformation 223]
297. 0 <=> main_query_1(sK6,sK7) [avatar definition]
298. main_query_1(sK6,sK7) <- {0} [avatar component clause 297]
300. 1 <=> ~main_query_1(sK6,sK7) [avatar definition]
301. ~main_query_1(sK6,sK7) <- {1} [avatar component clause 300]
303. 2 <=> main_query_2(sK6,sK7) [avatar definition]
304. main_query_2(sK6,sK7) <- {2} [avatar component clause 303]
306. 3 <=> ~main_query_2(sK6,sK7) [avatar definition]
307. ~main_query_2(sK6,sK7) <- {3} [avatar component clause 306]
308. ~1 | ~3 [avatar split clause 295,306,300]
309. 0 | 2 [avatar split clause 294,303,297]
314. select_4(sK6,sK7) <- {2} [resolution 246,304]
316. from_5(sK6,sK7) <- {2} [resolution 314,244]
319. subquery_from_6(sK6,sK7) <- {2} [resolution 316,242]
322. select_7(sK6,sK7) <- {2} [resolution 319,240]
325. where_9(sK6,sK7) <- {2} [resolution 322,238]
327. eq(sK6,sK7) <- {2} [resolution 325,235]
334. where_3(sK6,sK7) <- {2} [resolution 327,227]
429. tab1(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 228,224]
433. tab1(X2,X3,sK3(X2,X3),sK4(X2,X3),sK5(X2,X3)) | ~select_7(X2,X3) [resolution 237,233]
496. from_8(X2,X3,sK0(X2,X3),sK1(X2,X3),sK2(X2,X3)) | ~select_1(X2,X3) [resolution 429,234]
504. from_2(X0,X1,sK3(X0,X1),sK4(X0,X1),sK5(X0,X1)) | ~select_7(X0,X1) [resolution 433,225]
604. ~where_9(X2,X3) | ~select_1(X2,X3) | select_7(X2,X3) [resolution 496,239]
613. ~select_7(X2,X3) | ~where_3(X2,X3) | select_1(X2,X3) [resolution 504,230]
756. ~where_3(sK6,sK7) | select_1(sK6,sK7) <- {2} [resolution 613,322]
760. select_1(sK6,sK7) <- {2} [subsumption resolution 756,334]
773. main_query_1(sK6,sK7) <- {2} [resolution 760,232]
775. $false <- {1, 2} [subsumption resolution 773,301]
776. 1 | ~2 [avatar contradiction clause 775]
779. select_1(sK6,sK7) <- {0} [resolution 298,231]
783. where_3(sK6,sK7) <- {0} [resolution 779,229]
784. eq(sK6,sK7) <- {0} [resolution 783,226]
787. where_9(sK6,sK7) <- {0} [resolution 784,236]
797. ~select_1(sK6,sK7) | select_7(sK6,sK7) <- {0} [resolution 787,604]
801. select_7(sK6,sK7) <- {0} [subsumption resolution 797,779]
875. subquery_from_6(sK6,sK7) <- {0} [resolution 801,241]
925. from_5(sK6,sK7) <- {0} [resolution 875,243]
957. select_4(sK6,sK7) <- {0} [resolution 925,245]
990. main_query_2(sK6,sK7) <- {0} [resolution 957,247]
992. $false <- {0, 3} [subsumption resolution 990,307]
993. ~0 | 3 [avatar contradiction clause 992]
994. $false [avatar sat refutation 308,309,776,993]
% SZS output end Proof for query_12_query_16_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.011 s
% ------------------------------
% ------------------------------
