% Refutation found. Thanks to Tanya!
% SZS status Theorem for query03_query06_tptp4X
% SZS output start Proof for query03_query06_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X5,X1,X2,X3,X4] : (from_7(X5,X1,X2,X3,X4) <=> tab1(X5,X1,X2,X3,X4)) [input]
5. ! [X5] : (select_6(X5) <=> ? [X1,X2,X3,X4] : ($true & from_7(X5,X1,X2,X3,X4))) [input]
6. ! [X6] : (subquery_from_5(X6) <=> select_6(X6)) [input]
7. ! [X7] : (from_4(X7) <=> subquery_from_5(X7)) [input]
8. ! [X7] : (select_3(X7) <=> ($true & from_4(X7))) [input]
9. ! [X6] : (main_query_2(X6) <=> select_3(X6)) [input]
44. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
45. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 44]
46. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
47. ! [X0,X1,X2,X3,X4] : (from_7(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
48. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : ($true & from_7(X0,X1,X2,X3,X4))) [rectify 5]
49. ! [X0] : (select_6(X0) <=> ? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4)) [true and false elimination 48]
50. ! [X0] : (subquery_from_5(X0) <=> select_6(X0)) [rectify 6]
51. ! [X0] : (from_4(X0) <=> subquery_from_5(X0)) [rectify 7]
52. ! [X0] : (select_3(X0) <=> ($true & from_4(X0))) [rectify 8]
53. ! [X0] : (select_3(X0) <=> from_4(X0)) [true and false elimination 52]
54. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 9]
140. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 45]
141. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
142. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 46]
143. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 142]
144. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
145. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 143,144]
146. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
147. ! [X0,X1,X2,X3,X4] : ((from_7(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_7(X0,X1,X2,X3,X4))) [nnf transformation 47]
148. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_7(X0,X1,X2,X3,X4) | ~select_6(X0))) [nnf transformation 49]
149. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) | ~select_6(X0))) [rectify 148]
150. ! [X0] : (? [X5,X6,X7,X8] : from_7(X0,X5,X6,X7,X8) => from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) [choice axiom]
151. ! [X0] : ((select_6(X0) | ! [X1,X2,X3,X4] : ~from_7(X0,X1,X2,X3,X4)) & (from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0))) [skolemisation 149,150]
152. ! [X0] : ((subquery_from_5(X0) | ~select_6(X0)) & (select_6(X0) | ~subquery_from_5(X0))) [nnf transformation 50]
153. ! [X0] : ((from_4(X0) | ~subquery_from_5(X0)) & (subquery_from_5(X0) | ~from_4(X0))) [nnf transformation 51]
154. ! [X0] : ((select_3(X0) | ~from_4(X0)) & (from_4(X0) | ~select_3(X0))) [nnf transformation 53]
155. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 54]
158. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 140]
159. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
160. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 158,159]
161. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 141]
162. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 141]
163. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 145]
164. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 145]
165. ~main_query_1(X0) | select_1(X0) [cnf transformation 146]
166. ~select_1(X0) | main_query_1(X0) [cnf transformation 146]
167. ~from_7(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 147]
168. ~tab1(X0,X1,X2,X3,X4) | from_7(X0,X1,X2,X3,X4) [cnf transformation 147]
169. from_7(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [cnf transformation 151]
170. ~from_7(X0,X1,X2,X3,X4) | select_6(X0) [cnf transformation 151]
171. ~subquery_from_5(X0) | select_6(X0) [cnf transformation 152]
172. ~select_6(X0) | subquery_from_5(X0) [cnf transformation 152]
173. ~from_4(X0) | subquery_from_5(X0) [cnf transformation 153]
174. ~subquery_from_5(X0) | from_4(X0) [cnf transformation 153]
175. ~select_3(X0) | from_4(X0) [cnf transformation 154]
176. ~from_4(X0) | select_3(X0) [cnf transformation 154]
177. ~main_query_2(X0) | select_3(X0) [cnf transformation 155]
178. ~select_3(X0) | main_query_2(X0) [cnf transformation 155]
211. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 160]
212. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 160]
214. 0 <=> main_query_1(sK8) [avatar definition]
215. main_query_1(sK8) <- {0} [avatar component clause 214]
217. 1 <=> ~main_query_1(sK8) [avatar definition]
218. ~main_query_1(sK8) <- {1} [avatar component clause 217]
220. 2 <=> main_query_2(sK8) [avatar definition]
221. main_query_2(sK8) <- {2} [avatar component clause 220]
223. 3 <=> ~main_query_2(sK8) [avatar definition]
224. ~main_query_2(sK8) <- {3} [avatar component clause 223]
225. ~1 | ~3 [avatar split clause 212,223,217]
226. 0 | 2 [avatar split clause 211,220,214]
227. select_3(sK8) <- {2} [resolution 177,221]
229. from_4(sK8) <- {2} [resolution 227,175]
232. subquery_from_5(sK8) <- {2} [resolution 229,173]
235. select_6(sK8) <- {2} [resolution 232,171]
264. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 163,161]
266. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_6(X0) [resolution 169,167]
268. from_7(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 264,168]
271. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_6(X1) [resolution 266,162]
273. ~select_1(X1) | select_6(X1) [resolution 268,170]
275. ~select_6(X1) | select_1(X1) [resolution 271,164]
276. select_1(sK8) <- {2} [resolution 275,235]
280. main_query_1(sK8) <- {2} [resolution 276,166]
281. $false <- {1, 2} [subsumption resolution 280,218]
282. 1 | ~2 [avatar contradiction clause 281]
283. select_1(sK8) <- {0} [resolution 215,165]
286. select_6(sK8) <- {0} [resolution 283,273]
291. subquery_from_5(sK8) <- {0} [resolution 286,172]
294. from_4(sK8) <- {0} [resolution 291,174]
298. select_3(sK8) <- {0} [resolution 294,176]
302. main_query_2(sK8) <- {0} [resolution 298,178]
304. $false <- {0, 3} [subsumption resolution 302,224]
305. ~0 | 3 [avatar contradiction clause 304]
306. $false [avatar sat refutation 225,226,282,305]
% SZS output end Proof for query03_query06_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
