% Refutation found. Thanks to Tanya!
% SZS status Theorem for query03_query04_tptp4X
% SZS output start Proof for query03_query04_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X5,X6,X7,X8] : (from_4(X0,X5,X6,X7,X8) <=> tab1(X0,X5,X6,X7,X8)) [input]
5. ! [X0] : (where_5(X0) <=> eq(X0,X0)) [input]
6. ! [X0] : (select_3(X0) <=> ? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8))) [input]
7. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
8. ! [X9] : eq(X9,X9) [input]
42. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
43. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 42]
44. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
45. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
46. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [rectify 6]
47. ! [X0] : eq(X0,X0) [rectify 8]
132. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 43]
133. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
134. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 44]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 134]
136. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
137. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 135,136]
138. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
139. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 45]
140. ! [X0] : ((where_5(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_5(X0))) [nnf transformation 5]
141. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 46]
142. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 141]
143. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
144. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 142,143]
145. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 7]
148. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 132]
149. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
150. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 148,149]
151. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 133]
152. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 133]
153. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 137]
154. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 137]
155. ~main_query_1(X0) | select_1(X0) [cnf transformation 138]
156. ~select_1(X0) | main_query_1(X0) [cnf transformation 138]
157. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 139]
158. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 139]
160. ~eq(X0,X0) | where_5(X0) [cnf transformation 140]
161. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 144]
163. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 144]
164. ~main_query_2(X0) | select_3(X0) [cnf transformation 145]
165. ~select_3(X0) | main_query_2(X0) [cnf transformation 145]
166. eq(X0,X0) [cnf transformation 47]
198. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 150]
199. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 150]
201. 0 <=> main_query_1(sK8) [avatar definition]
202. main_query_1(sK8) <- {0} [avatar component clause 201]
204. 1 <=> ~main_query_1(sK8) [avatar definition]
205. ~main_query_1(sK8) <- {1} [avatar component clause 204]
207. 2 <=> main_query_2(sK8) [avatar definition]
208. main_query_2(sK8) <- {2} [avatar component clause 207]
210. 3 <=> ~main_query_2(sK8) [avatar definition]
211. ~main_query_2(sK8) <- {3} [avatar component clause 210]
212. ~1 | ~3 [avatar split clause 199,210,204]
213. 0 | 2 [avatar split clause 198,207,201]
214. select_3(sK8) <- {2} [resolution 164,208]
218. where_5(X1) [resolution 160,166]
252. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 153,151]
254. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 161,157]
256. from_4(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 252,158]
260. from_2(X1,sK4(X1),sK5(X1),sK6(X1),sK7(X1)) | ~select_3(X1) [resolution 254,152]
265. ~select_1(X5) | ~where_5(X5) | select_3(X5) [resolution 256,163]
266. ~select_1(X5) | select_3(X5) [subsumption resolution 265,218]
272. ~select_3(X1) | select_1(X1) [resolution 260,154]
275. select_1(sK8) <- {2} [resolution 272,214]
277. main_query_1(sK8) <- {2} [resolution 275,156]
278. $false <- {1, 2} [subsumption resolution 277,205]
279. 1 | ~2 [avatar contradiction clause 278]
280. select_1(sK8) <- {0} [resolution 202,155]
281. select_3(sK8) <- {0} [resolution 280,266]
284. main_query_2(sK8) <- {0} [resolution 281,165]
286. $false <- {0, 3} [subsumption resolution 284,211]
287. ~0 | 3 [avatar contradiction clause 286]
288. $false [avatar sat refutation 212,213,279,287]
% SZS output end Proof for query03_query04_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.051 s
% ------------------------------
% ------------------------------
