% Refutation found. Thanks to Tanya!
% SZS status Theorem for query01_query07_tptp4X
% SZS output start Proof for query01_query07_tptp4X
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [input]
4. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
5. ! [X0] : (where_5(X0) <=> (~gt(X0,X0) & ~lt(X0,X0))) [input]
6. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4))) [input]
7. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [input]
8. ! [X5] : eq(X5,X5) [input]
15. ! [X5,X6] : (lt(X5,X6) <=> (~eq(X5,X6) & lte(X5,X6))) [input]
16. ! [X5,X6] : (gt(X5,X6) <=> lt(X6,X5)) [input]
42. ! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [input]
43. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [negated conjecture 42]
44. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4)) [true and false elimination 2]
45. ! [X0] : eq(X0,X0) [rectify 8]
52. ! [X0,X1] : (lt(X0,X1) <=> (~eq(X0,X1) & lte(X0,X1))) [rectify 15]
53. ! [X0,X1] : (gt(X0,X1) <=> lt(X1,X0)) [rectify 16]
130. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 43]
131. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
132. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X1,X2,X3,X4] : from_2(X0,X1,X2,X3,X4) | ~select_1(X0))) [nnf transformation 44]
133. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) | ~select_1(X0))) [rectify 132]
134. ! [X0] : (? [X5,X6,X7,X8] : from_2(X0,X5,X6,X7,X8) => from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0))) [choice axiom]
135. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4] : ~from_2(X0,X1,X2,X3,X4)) & (from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0))) [skolemisation 133,134]
136. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 3]
137. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 4]
138. ! [X0] : ((where_5(X0) | (gt(X0,X0) | lt(X0,X0))) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [nnf transformation 5]
139. ! [X0] : ((where_5(X0) | gt(X0,X0) | lt(X0,X0)) & ((~gt(X0,X0) & ~lt(X0,X0)) | ~where_5(X0))) [flattening 138]
140. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X1,X2,X3,X4] : (where_5(X0) & from_4(X0,X1,X2,X3,X4)) | ~select_3(X0))) [nnf transformation 6]
141. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) | ~select_3(X0))) [rectify 140]
142. ! [X0] : (? [X5,X6,X7,X8] : (where_5(X0) & from_4(X0,X5,X6,X7,X8)) => (where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)))) [choice axiom]
143. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4] : (~where_5(X0) | ~from_4(X0,X1,X2,X3,X4))) & ((where_5(X0) & from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0))) | ~select_3(X0))) [skolemisation 141,142]
144. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 7]
147. ! [X0,X1] : ((lt(X0,X1) | (eq(X0,X1) | ~lte(X0,X1))) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [nnf transformation 52]
148. ! [X0,X1] : ((lt(X0,X1) | eq(X0,X1) | ~lte(X0,X1)) & ((~eq(X0,X1) & lte(X0,X1)) | ~lt(X0,X1))) [flattening 147]
149. ! [X0,X1] : ((gt(X0,X1) | ~lt(X1,X0)) & (lt(X1,X0) | ~gt(X0,X1))) [nnf transformation 53]
150. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 130]
151. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8))) [choice axiom]
152. (~main_query_2(sK8) | ~main_query_1(sK8)) & (main_query_2(sK8) | main_query_1(sK8)) [skolemisation 150,151]
153. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 131]
154. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 131]
155. from_2(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [cnf transformation 135]
156. ~from_2(X0,X1,X2,X3,X4) | select_1(X0) [cnf transformation 135]
157. ~main_query_1(X0) | select_1(X0) [cnf transformation 136]
158. ~select_1(X0) | main_query_1(X0) [cnf transformation 136]
159. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 137]
160. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 137]
163. gt(X0,X0) | where_5(X0) | lt(X0,X0) [cnf transformation 139]
164. from_4(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [cnf transformation 143]
166. ~from_4(X0,X1,X2,X3,X4) | ~where_5(X0) | select_3(X0) [cnf transformation 143]
167. ~main_query_2(X0) | select_3(X0) [cnf transformation 144]
168. ~select_3(X0) | main_query_2(X0) [cnf transformation 144]
169. eq(X0,X0) [cnf transformation 45]
179. ~eq(X0,X1) | ~lt(X0,X1) [cnf transformation 148]
181. ~gt(X0,X1) | lt(X1,X0) [cnf transformation 149]
206. main_query_2(sK8) | main_query_1(sK8) [cnf transformation 152]
207. ~main_query_2(sK8) | ~main_query_1(sK8) [cnf transformation 152]
209. 0 <=> main_query_1(sK8) [avatar definition]
210. main_query_1(sK8) <- {0} [avatar component clause 209]
212. 1 <=> ~main_query_1(sK8) [avatar definition]
213. ~main_query_1(sK8) <- {1} [avatar component clause 212]
215. 2 <=> main_query_2(sK8) [avatar definition]
216. main_query_2(sK8) <- {2} [avatar component clause 215]
218. 3 <=> ~main_query_2(sK8) [avatar definition]
219. ~main_query_2(sK8) <- {3} [avatar component clause 218]
220. ~1 | ~3 [avatar split clause 207,218,212]
221. 0 | 2 [avatar split clause 206,215,209]
222. select_3(sK8) <- {2} [resolution 167,216]
229. ~lt(X0,X0) [resolution 179,169]
231. where_5(X1) | lt(X1,X1) | lt(X1,X1) [resolution 163,181]
232. where_5(X1) | lt(X1,X1) [duplicate literal removal 231]
233. where_5(X1) [subsumption resolution 232,229]
297. tab1(X0,sK0(X0),sK1(X0),sK2(X0),sK3(X0)) | ~select_1(X0) [resolution 155,153]
337. tab1(X0,sK4(X0),sK5(X0),sK6(X0),sK7(X0)) | ~select_3(X0) [resolution 164,159]
420. from_4(X6,sK0(X6),sK1(X6),sK2(X6),sK3(X6)) | ~select_1(X6) [resolution 297,160]
428. from_2(X9,sK4(X9),sK5(X9),sK6(X9),sK7(X9)) | ~select_3(X9) [resolution 337,154]
432. ~select_1(X1) | ~where_5(X1) | select_3(X1) [resolution 420,166]
433. ~select_1(X1) | select_3(X1) [subsumption resolution 432,233]
435. ~select_3(X1) | select_1(X1) [resolution 428,156]
441. select_1(sK8) <- {2} [resolution 435,222]
443. main_query_1(sK8) <- {2} [resolution 441,158]
444. $false <- {1, 2} [subsumption resolution 443,213]
445. 1 | ~2 [avatar contradiction clause 444]
446. select_1(sK8) <- {0} [resolution 210,157]
447. select_3(sK8) <- {0} [resolution 446,433]
452. main_query_2(sK8) <- {0} [resolution 447,168]
454. $false <- {0, 3} [subsumption resolution 452,219]
455. ~0 | 3 [avatar contradiction clause 454]
456. $false [avatar sat refutation 220,221,445,455]
% SZS output end Proof for query01_query07_tptp4X
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.017 s
% ------------------------------
% ------------------------------
