% Refutation found. Thanks to Tanya!
% SZS status Theorem for test9
% SZS output start Proof for test9
1. ! [X0] : (main_query_0(X0) <=> ? [X1,X2,X3,X4,X5] : (X0 = X4 & X0 = X2 & $true & tab(X0,X4,X5) & tab(X1,X2,X3))) [input]
2. ! [X0] : (main_query_1(X0) <=> ? [X4,X5] : (X0 = X4 & $true & tab(X0,X4,X5))) [input]
3. ! [X0] : (main_query_0(X0) <=> main_query_1(X0)) [input]
4. ~! [X0] : (main_query_0(X0) <=> main_query_1(X0)) [negated conjecture 3]
16. ! [X0] : (main_query_0(X0) <=> ? [X1,X2,X3,X4,X5] : (X0 = X4 & X0 = X2 & tab(X0,X4,X5) & tab(X1,X2,X3))) [true and false elimination 1]
17. ! [X0] : (main_query_1(X0) <=> ? [X1,X2] : (X0 = X1 & $true & tab(X0,X1,X2))) [rectify 2]
18. ! [X0] : (main_query_1(X0) <=> ? [X1,X2] : (X0 = X1 & tab(X0,X1,X2))) [true and false elimination 17]
30. ? [X0] : (main_query_0(X0) <~> main_query_1(X0)) [ennf transformation 4]
48. ! [X0] : ((main_query_0(X0) | ! [X1,X2,X3,X4,X5] : (X0 != X4 | X0 != X2 | ~tab(X0,X4,X5) | ~tab(X1,X2,X3))) & (? [X1,X2,X3,X4,X5] : (X0 = X4 & X0 = X2 & tab(X0,X4,X5) & tab(X1,X2,X3)) | ~main_query_0(X0))) [nnf transformation 16]
49. ! [X0] : ((main_query_0(X0) | ! [X1,X2,X3,X4,X5] : (X0 != X4 | X0 != X2 | ~tab(X0,X4,X5) | ~tab(X1,X2,X3))) & (? [X6,X7,X8,X9,X10] : (X0 = X9 & X0 = X7 & tab(X0,X9,X10) & tab(X6,X7,X8)) | ~main_query_0(X0))) [rectify 48]
50. ! [X0] : (? [X6,X7,X8,X9,X10] : (X0 = X9 & X0 = X7 & tab(X0,X9,X10) & tab(X6,X7,X8)) => (sK3(X0) = X0 & sK1(X0) = X0 & tab(X0,sK3(X0),sK4(X0)) & tab(sK0(X0),sK1(X0),sK2(X0)))) [choice axiom]
51. ! [X0] : ((main_query_0(X0) | ! [X1,X2,X3,X4,X5] : (X0 != X4 | X0 != X2 | ~tab(X0,X4,X5) | ~tab(X1,X2,X3))) & ((sK3(X0) = X0 & sK1(X0) = X0 & tab(X0,sK3(X0),sK4(X0)) & tab(sK0(X0),sK1(X0),sK2(X0))) | ~main_query_0(X0))) [skolemisation 49,50]
52. ! [X0] : ((main_query_1(X0) | ! [X1,X2] : (X0 != X1 | ~tab(X0,X1,X2))) & (? [X1,X2] : (X0 = X1 & tab(X0,X1,X2)) | ~main_query_1(X0))) [nnf transformation 18]
53. ! [X0] : ((main_query_1(X0) | ! [X1,X2] : (X0 != X1 | ~tab(X0,X1,X2))) & (? [X3,X4] : (X0 = X3 & tab(X0,X3,X4)) | ~main_query_1(X0))) [rectify 52]
54. ! [X0] : (? [X3,X4] : (X0 = X3 & tab(X0,X3,X4)) => (sK5(X0) = X0 & tab(X0,sK5(X0),sK6(X0)))) [choice axiom]
55. ! [X0] : ((main_query_1(X0) | ! [X1,X2] : (X0 != X1 | ~tab(X0,X1,X2))) & ((sK5(X0) = X0 & tab(X0,sK5(X0),sK6(X0))) | ~main_query_1(X0))) [skolemisation 53,54]
56. ? [X0] : ((~main_query_1(X0) | ~main_query_0(X0)) & (main_query_1(X0) | main_query_0(X0))) [nnf transformation 30]
57. ? [X0] : ((~main_query_1(X0) | ~main_query_0(X0)) & (main_query_1(X0) | main_query_0(X0))) => ((~main_query_1(sK7) | ~main_query_0(sK7)) & (main_query_1(sK7) | main_query_0(sK7))) [choice axiom]
58. (~main_query_1(sK7) | ~main_query_0(sK7)) & (main_query_1(sK7) | main_query_0(sK7)) [skolemisation 56,57]
60. tab(X0,sK3(X0),sK4(X0)) | ~main_query_0(X0) [cnf transformation 51]
62. ~main_query_0(X0) | sK3(X0) = X0 [cnf transformation 51]
63. main_query_0(X0) | X0 != X4 | X0 != X2 | ~tab(X0,X4,X5) | ~tab(X1,X2,X3) [cnf transformation 51]
64. tab(X0,sK5(X0),sK6(X0)) | ~main_query_1(X0) [cnf transformation 55]
65. ~main_query_1(X0) | sK5(X0) = X0 [cnf transformation 55]
66. main_query_1(X0) | X0 != X1 | ~tab(X0,X1,X2) [cnf transformation 55]
67. main_query_1(sK7) | main_query_0(sK7) [cnf transformation 58]
68. ~main_query_1(sK7) | ~main_query_0(sK7) [cnf transformation 58]
80. main_query_0(X4) | X2 != X4 | ~tab(X4,X4,X5) | ~tab(X1,X2,X3) [equality resolution 63]
81. ~tab(X4,X4,X5) | main_query_0(X4) | ~tab(X1,X4,X3) [equality resolution 80]
82. ~tab(X1,X1,X2) | main_query_1(X1) [equality resolution 66]
92. 0 <=> main_query_0(sK7) [avatar definition]
93. main_query_0(sK7) <- {0} [avatar component clause 92]
95. 1 <=> ~main_query_0(sK7) [avatar definition]
96. ~main_query_0(sK7) <- {1} [avatar component clause 95]
98. 2 <=> main_query_1(sK7) [avatar definition]
99. main_query_1(sK7) <- {2} [avatar component clause 98]
101. 3 <=> ~main_query_1(sK7) [avatar definition]
102. ~main_query_1(sK7) <- {3} [avatar component clause 101]
103. ~1 | ~3 [avatar split clause 68,101,95]
104. 0 | 2 [avatar split clause 67,98,92]
105. sK5(sK7) = sK7 <- {2} [resolution 65,99]
106. tab(sK7,sK7,sK6(sK7)) | ~main_query_1(sK7) <- {2} [superposition 64,105]
107. tab(sK7,sK7,sK6(sK7)) <- {2} [subsumption resolution 106,99]
111. main_query_0(sK7) | ~tab(X0,sK7,X1) <- {2} [resolution 81,107]
112. ~tab(X0,sK7,X1) <- {1, 2} [subsumption resolution 111,96]
113. $false <- {1, 2} [resolution 112,107]
114. 1 | ~2 [avatar contradiction clause 113]
115. sK3(sK7) = sK7 <- {0} [resolution 93,62]
117. tab(sK7,sK7,sK4(sK7)) | ~main_query_0(sK7) <- {0} [superposition 60,115]
118. tab(sK7,sK7,sK4(sK7)) <- {0} [subsumption resolution 117,93]
122. main_query_1(sK7) <- {0} [resolution 118,82]
123. $false <- {0, 3} [subsumption resolution 122,102]
124. ~0 | 3 [avatar contradiction clause 123]
125. $false [avatar sat refutation 103,104,114,124]
% SZS output end Proof for test9
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.059 s
% ------------------------------
% ------------------------------
