% Refutation found. Thanks to Tanya!
% SZS status Theorem for test1_symm
% SZS output start Proof for test1_symm
1. ! [X0,X1] : (main_query_0(X0,X1) <=> ? [X2,X3,X4,X5] : (X2 != X4 & lessthanorequal(X2,X4) & X1 = X2 & X0 = X4 & $true & tab(X1,X4,X5) & tab(X0,X2,X3))) [input]
2. ! [X0,X1] : (main_query_1(X0,X1) <=> ? [X2,X3,X4,X5,X6,X7,X8] : (X2 != X4 & lessthanorequal(X2,X4) & X6 = X6 & X1 = X2 & X0 = X4 & $true & tab(X6,X7,X8) & tab(X1,X4,X5) & tab(X0,X2,X3))) [input]
3. ! [X0,X1] : (main_query_0(X0,X1) <=> main_query_1(X0,X1)) [input]
4. ~! [X0,X1] : (main_query_0(X0,X1) <=> main_query_1(X0,X1)) [negated conjecture 3]
13. ! [X13] : lessthanorequal(X13,X13) [input]
14. ! [X13,X14] : ((lessthanorequal(X14,X13) & lessthanorequal(X13,X14)) => X13 = X14) [input]
16. ! [X0,X1] : (main_query_0(X0,X1) <=> ? [X2,X3,X4,X5] : (X2 != X4 & lessthanorequal(X2,X4) & X1 = X2 & X0 = X4 & tab(X1,X4,X5) & tab(X0,X2,X3))) [true and false elimination 1]
17. ! [X0,X1] : (main_query_1(X0,X1) <=> ? [X2,X3,X4,X5,X6,X7,X8] : (X2 != X4 & lessthanorequal(X2,X4) & X6 = X6 & X1 = X2 & X0 = X4 & tab(X6,X7,X8) & tab(X1,X4,X5) & tab(X0,X2,X3))) [true and false elimination 2]
26. ! [X0] : lessthanorequal(X0,X0) [rectify 13]
27. ! [X0,X1] : ((lessthanorequal(X1,X0) & lessthanorequal(X0,X1)) => X0 = X1) [rectify 14]
29. ? [X0,X1] : (main_query_0(X0,X1) <~> main_query_1(X0,X1)) [ennf transformation 4]
43. ! [X0,X1] : (X0 = X1 | (~lessthanorequal(X1,X0) | ~lessthanorequal(X0,X1))) [ennf transformation 27]
44. ! [X0,X1] : (X0 = X1 | ~lessthanorequal(X1,X0) | ~lessthanorequal(X0,X1)) [flattening 43]
47. ! [X0,X1] : ((main_query_0(X0,X1) | ! [X2,X3,X4,X5] : (X2 = X4 | ~lessthanorequal(X2,X4) | X1 != X2 | X0 != X4 | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & (? [X2,X3,X4,X5] : (X2 != X4 & lessthanorequal(X2,X4) & X1 = X2 & X0 = X4 & tab(X1,X4,X5) & tab(X0,X2,X3)) | ~main_query_0(X0,X1))) [nnf transformation 16]
48. ! [X0,X1] : ((main_query_0(X0,X1) | ! [X2,X3,X4,X5] : (X2 = X4 | ~lessthanorequal(X2,X4) | X1 != X2 | X0 != X4 | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & (? [X6,X7,X8,X9] : (X6 != X8 & lessthanorequal(X6,X8) & X1 = X6 & X0 = X8 & tab(X1,X8,X9) & tab(X0,X6,X7)) | ~main_query_0(X0,X1))) [rectify 47]
49. ! [X1,X0] : (? [X6,X7,X8,X9] : (X6 != X8 & lessthanorequal(X6,X8) & X1 = X6 & X0 = X8 & tab(X1,X8,X9) & tab(X0,X6,X7)) => (sK0(X0,X1) != sK2(X0,X1) & lessthanorequal(sK0(X0,X1),sK2(X0,X1)) & sK0(X0,X1) = X1 & sK2(X0,X1) = X0 & tab(X1,sK2(X0,X1),sK3(X0,X1)) & tab(X0,sK0(X0,X1),sK1(X0,X1)))) [choice axiom]
50. ! [X0,X1] : ((main_query_0(X0,X1) | ! [X2,X3,X4,X5] : (X2 = X4 | ~lessthanorequal(X2,X4) | X1 != X2 | X0 != X4 | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & ((sK0(X0,X1) != sK2(X0,X1) & lessthanorequal(sK0(X0,X1),sK2(X0,X1)) & sK0(X0,X1) = X1 & sK2(X0,X1) = X0 & tab(X1,sK2(X0,X1),sK3(X0,X1)) & tab(X0,sK0(X0,X1),sK1(X0,X1))) | ~main_query_0(X0,X1))) [skolemisation 48,49]
51. ! [X0,X1] : ((main_query_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8] : (X2 = X4 | ~lessthanorequal(X2,X4) | X6 != X6 | X1 != X2 | X0 != X4 | ~tab(X6,X7,X8) | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & (? [X2,X3,X4,X5,X6,X7,X8] : (X2 != X4 & lessthanorequal(X2,X4) & X6 = X6 & X1 = X2 & X0 = X4 & tab(X6,X7,X8) & tab(X1,X4,X5) & tab(X0,X2,X3)) | ~main_query_1(X0,X1))) [nnf transformation 17]
52. ! [X0,X1] : ((main_query_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8] : (X2 = X4 | ~lessthanorequal(X2,X4) | X6 != X6 | X1 != X2 | X0 != X4 | ~tab(X6,X7,X8) | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & (? [X9,X10,X11,X12,X13,X14,X15] : (X9 != X11 & lessthanorequal(X9,X11) & X13 = X13 & X1 = X9 & X0 = X11 & tab(X13,X14,X15) & tab(X1,X11,X12) & tab(X0,X9,X10)) | ~main_query_1(X0,X1))) [rectify 51]
53. ! [X1,X0] : (? [X9,X10,X11,X12,X13,X14,X15] : (X9 != X11 & lessthanorequal(X9,X11) & X13 = X13 & X1 = X9 & X0 = X11 & tab(X13,X14,X15) & tab(X1,X11,X12) & tab(X0,X9,X10)) => (sK4(X0,X1) != sK6(X0,X1) & lessthanorequal(sK4(X0,X1),sK6(X0,X1)) & sK8(X0,X1) = sK8(X0,X1) & sK4(X0,X1) = X1 & sK6(X0,X1) = X0 & tab(sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) & tab(X1,sK6(X0,X1),sK7(X0,X1)) & tab(X0,sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
54. ! [X0,X1] : ((main_query_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8] : (X2 = X4 | ~lessthanorequal(X2,X4) | X6 != X6 | X1 != X2 | X0 != X4 | ~tab(X6,X7,X8) | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & ((sK4(X0,X1) != sK6(X0,X1) & lessthanorequal(sK4(X0,X1),sK6(X0,X1)) & sK8(X0,X1) = sK8(X0,X1) & sK4(X0,X1) = X1 & sK6(X0,X1) = X0 & tab(sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) & tab(X1,sK6(X0,X1),sK7(X0,X1)) & tab(X0,sK4(X0,X1),sK5(X0,X1))) | ~main_query_1(X0,X1))) [skolemisation 52,53]
55. ? [X0,X1] : ((~main_query_1(X0,X1) | ~main_query_0(X0,X1)) & (main_query_1(X0,X1) | main_query_0(X0,X1))) [nnf transformation 29]
56. ? [X0,X1] : ((~main_query_1(X0,X1) | ~main_query_0(X0,X1)) & (main_query_1(X0,X1) | main_query_0(X0,X1))) => ((~main_query_1(sK11,sK12) | ~main_query_0(sK11,sK12)) & (main_query_1(sK11,sK12) | main_query_0(sK11,sK12))) [choice axiom]
57. (~main_query_1(sK11,sK12) | ~main_query_0(sK11,sK12)) & (main_query_1(sK11,sK12) | main_query_0(sK11,sK12)) [skolemisation 55,56]
58. tab(X0,sK0(X0,X1),sK1(X0,X1)) | ~main_query_0(X0,X1) [cnf transformation 50]
59. tab(X1,sK2(X0,X1),sK3(X0,X1)) | ~main_query_0(X0,X1) [cnf transformation 50]
60. ~main_query_0(X0,X1) | sK2(X0,X1) = X0 [cnf transformation 50]
61. ~main_query_0(X0,X1) | sK0(X0,X1) = X1 [cnf transformation 50]
62. lessthanorequal(sK0(X0,X1),sK2(X0,X1)) | ~main_query_0(X0,X1) [cnf transformation 50]
63. sK0(X0,X1) != sK2(X0,X1) | ~main_query_0(X0,X1) [cnf transformation 50]
64. main_query_0(X0,X1) | X2 = X4 | ~lessthanorequal(X2,X4) | X1 != X2 | X0 != X4 | ~tab(X1,X4,X5) | ~tab(X0,X2,X3) [cnf transformation 50]
65. tab(X0,sK4(X0,X1),sK5(X0,X1)) | ~main_query_1(X0,X1) [cnf transformation 54]
66. tab(X1,sK6(X0,X1),sK7(X0,X1)) | ~main_query_1(X0,X1) [cnf transformation 54]
68. ~main_query_1(X0,X1) | sK6(X0,X1) = X0 [cnf transformation 54]
69. ~main_query_1(X0,X1) | sK4(X0,X1) = X1 [cnf transformation 54]
71. lessthanorequal(sK4(X0,X1),sK6(X0,X1)) | ~main_query_1(X0,X1) [cnf transformation 54]
72. sK4(X0,X1) != sK6(X0,X1) | ~main_query_1(X0,X1) [cnf transformation 54]
73. main_query_1(X0,X1) | X2 = X4 | ~lessthanorequal(X2,X4) | X6 != X6 | X1 != X2 | X0 != X4 | ~tab(X6,X7,X8) | ~tab(X1,X4,X5) | ~tab(X0,X2,X3) [cnf transformation 54]
74. main_query_1(sK11,sK12) | main_query_0(sK11,sK12) [cnf transformation 57]
75. ~main_query_1(sK11,sK12) | ~main_query_0(sK11,sK12) [cnf transformation 57]
84. lessthanorequal(X0,X0) [cnf transformation 26]
85. ~lessthanorequal(X1,X0) | X0 = X1 | ~lessthanorequal(X0,X1) [cnf transformation 44]
87. main_query_0(X0,X2) | X2 = X4 | ~lessthanorequal(X2,X4) | X0 != X4 | ~tab(X2,X4,X5) | ~tab(X0,X2,X3) [equality resolution 64]
88. ~tab(X4,X2,X3) | X2 = X4 | ~lessthanorequal(X2,X4) | ~tab(X2,X4,X5) | main_query_0(X4,X2) [equality resolution 87]
89. main_query_1(X0,X2) | X2 = X4 | ~lessthanorequal(X2,X4) | X6 != X6 | X0 != X4 | ~tab(X6,X7,X8) | ~tab(X2,X4,X5) | ~tab(X0,X2,X3) [equality resolution 73]
90. main_query_1(X4,X2) | X2 = X4 | ~lessthanorequal(X2,X4) | X6 != X6 | ~tab(X6,X7,X8) | ~tab(X2,X4,X5) | ~tab(X4,X2,X3) [equality resolution 89]
99. main_query_1(X4,X2) | X2 = X4 | ~lessthanorequal(X2,X4) | ~tab(X6,X7,X8) | ~tab(X2,X4,X5) | ~tab(X4,X2,X3) [trivial inequality removal 90]
101. 0 <=> ! [X7,X8,X6] : ~tab(X6,X7,X8) [avatar definition]
102. ~tab(X6,X7,X8) <- {0} [avatar component clause 101]
104. 2 <=> ! [X3,X5,X2,X4] : (main_query_1(X4,X2) | ~tab(X4,X2,X3) | ~tab(X2,X4,X5) | ~lessthanorequal(X2,X4) | X2 = X4) [avatar definition]
105. ~tab(X4,X2,X3) | main_query_1(X4,X2) | ~tab(X2,X4,X5) | ~lessthanorequal(X2,X4) | X2 = X4 <- {2} [avatar component clause 104]
106. 0 | 2 [avatar split clause 99,104,101]
108. 4 <=> main_query_0(sK11,sK12) [avatar definition]
109. main_query_0(sK11,sK12) <- {4} [avatar component clause 108]
111. 5 <=> ~main_query_0(sK11,sK12) [avatar definition]
112. ~main_query_0(sK11,sK12) <- {5} [avatar component clause 111]
114. 6 <=> main_query_1(sK11,sK12) [avatar definition]
115. main_query_1(sK11,sK12) <- {6} [avatar component clause 114]
117. 7 <=> ~main_query_1(sK11,sK12) [avatar definition]
118. ~main_query_1(sK11,sK12) <- {7} [avatar component clause 117]
119. ~5 | ~7 [avatar split clause 75,117,111]
120. 4 | 6 [avatar split clause 74,114,108]
121. sK2(sK11,sK12) = sK11 <- {4} [resolution 60,109]
122. sK0(sK11,sK12) = sK12 <- {4} [resolution 61,109]
127. lessthanorequal(sK12,sK2(sK11,sK12)) | ~main_query_0(sK11,sK12) <- {4} [superposition 62,122]
130. lessthanorequal(sK12,sK2(sK11,sK12)) <- {4} [subsumption resolution 127,109]
131. lessthanorequal(sK12,sK11) <- {4} [forward demodulation 130,121]
143. 11 <=> sK11 != sK12 [avatar definition]
144. sK11 != sK12 <- {11} [avatar component clause 143]
146. 10 <=> sK11 = sK12 [avatar definition]
147. sK11 = sK12 <- {10} [avatar component clause 146]
149. sK0(sK11,sK12) != sK11 | ~main_query_0(sK11,sK12) <- {4} [superposition 63,121]
150. sK0(sK11,sK12) != sK11 <- {4} [subsumption resolution 149,109]
152. ~main_query_1(X3,X4) | sK4(X3,X4) = sK6(X3,X4) | ~lessthanorequal(sK6(X3,X4),sK4(X3,X4)) [resolution 71,85]
153. ~lessthanorequal(sK6(X3,X4),sK4(X3,X4)) | ~main_query_1(X3,X4) [subsumption resolution 152,72]
154. sK11 != sK12 <- {4} [superposition 150,122]
155. ~11 | ~4 [avatar split clause 154,108,143]
156. tab(sK11,sK12,sK1(sK11,sK12)) | ~main_query_0(sK11,sK12) <- {4} [superposition 58,122]
157. tab(sK11,sK12,sK1(sK11,sK12)) <- {4} [subsumption resolution 156,109]
158. tab(sK12,sK11,sK3(sK11,sK12)) | ~main_query_0(sK11,sK12) <- {4} [superposition 59,121]
159. tab(sK12,sK11,sK3(sK11,sK12)) <- {4} [subsumption resolution 158,109]
179. main_query_1(sK11,sK12) | ~tab(sK12,sK11,X3) | ~lessthanorequal(sK12,sK11) | sK11 = sK12 <- {2, 4} [resolution 105,157]
185. ~tab(sK12,sK11,X3) | ~lessthanorequal(sK12,sK11) | sK11 = sK12 <- {2, 4, 7} [subsumption resolution 179,118]
186. ~tab(sK12,sK11,X3) | sK11 = sK12 <- {2, 4, 7} [subsumption resolution 185,131]
187. ~tab(sK12,sK11,X3) <- {2, 4, 7, 11} [subsumption resolution 186,144]
188. $false <- {2, 4, 7, 11} [resolution 187,159]
189. ~2 | ~4 | 7 | 11 [avatar contradiction clause 188]
190. sK4(sK11,sK12) = sK12 <- {6} [resolution 115,69]
191. sK6(sK11,sK12) = sK11 <- {6} [resolution 115,68]
193. tab(sK11,sK12,sK5(sK11,sK12)) | ~main_query_1(sK11,sK12) <- {6} [superposition 65,190]
194. lessthanorequal(sK12,sK6(sK11,sK12)) | ~main_query_1(sK11,sK12) <- {6} [superposition 71,190]
197. tab(sK11,sK12,sK5(sK11,sK12)) <- {6} [subsumption resolution 193,115]
198. lessthanorequal(sK12,sK6(sK11,sK12)) <- {6} [subsumption resolution 194,115]
199. lessthanorequal(sK12,sK11) <- {6} [forward demodulation 198,191]
203. tab(sK12,sK11,sK7(sK11,sK12)) | ~main_query_1(sK11,sK12) <- {6} [superposition 66,191]
208. tab(sK12,sK11,sK7(sK11,sK12)) <- {6} [subsumption resolution 203,115]
218. sK11 = sK12 | ~lessthanorequal(sK12,sK11) | ~tab(sK12,sK11,X1) | main_query_0(sK11,sK12) <- {6} [resolution 197,88]
219. ~lessthanorequal(sK12,sK11) | ~tab(sK12,sK11,X1) | main_query_0(sK11,sK12) <- {6, 11} [subsumption resolution 218,144]
220. ~tab(sK12,sK11,X1) | main_query_0(sK11,sK12) <- {6, 11} [subsumption resolution 219,199]
221. ~tab(sK12,sK11,X1) <- {5, 6, 11} [subsumption resolution 220,112]
224. $false <- {5, 6, 11} [resolution 208,221]
225. 5 | ~6 | 11 [avatar contradiction clause 224]
228. ~main_query_0(X0,X1) <- {0} [resolution 102,58]
233. $false <- {0, 4} [resolution 228,109]
234. ~0 | ~4 [avatar contradiction clause 233]
241. main_query_1(sK11,sK11) <- {6, 10} [forward demodulation 115,147]
243. sK4(sK11,sK11) = sK11 <- {6, 10} [resolution 241,69]
244. sK6(sK11,sK11) = sK11 <- {6, 10} [resolution 241,68]
251. ~lessthanorequal(sK6(sK11,sK11),sK11) | ~main_query_1(sK11,sK11) <- {6, 10} [superposition 153,243]
256. ~lessthanorequal(sK6(sK11,sK11),sK11) <- {6, 10} [subsumption resolution 251,241]
257. ~lessthanorequal(sK11,sK11) <- {6, 10} [forward demodulation 256,244]
258. $false <- {6, 10} [subsumption resolution 257,84]
259. ~6 | ~10 [avatar contradiction clause 258]
263. $false [avatar sat refutation 106,119,120,155,189,225,234,259]
% SZS output end Proof for test1_symm
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.064 s
% ------------------------------
% ------------------------------
