% Refutation found. Thanks to Tanya!
% SZS status Theorem for test9
% SZS output start Proof for test9
1. ! [X0] : (main_query_0(X0) <=> ? [X1,X2,X3,X4,X5] : (b_fst = b_snd & b_fst = a_snd & $true & tab(X0,X4,X5) & tab(X1,X2,X3))) [input]
2. ! [X0] : (main_query_1(X0) <=> ? [X4,X5] : (b_fst = b_snd & $true & tab(X0,X4,X5))) [input]
3. ! [X0] : (main_query_0(X0) <=> main_query_1(X0)) [input]
4. ~! [X0] : (main_query_0(X0) <=> main_query_1(X0)) [negated conjecture 3]
15. ! [X10,X11,X12] : ((X10 = X11 & X10 = X11) => X10 = X12) [input]
16. ! [X0] : (main_query_0(X0) <=> ? [X1,X2,X3,X4,X5] : (b_fst = b_snd & b_fst = a_snd & tab(X0,X4,X5) & tab(X1,X2,X3))) [true and false elimination 1]
17. ! [X0] : (main_query_1(X0) <=> ? [X1,X2] : (b_fst = b_snd & $true & tab(X0,X1,X2))) [rectify 2]
18. ! [X0] : (main_query_1(X0) <=> ? [X1,X2] : (b_fst = b_snd & tab(X0,X1,X2))) [true and false elimination 17]
29. ! [X0,X1,X2] : ((X0 = X1 & X0 = X1) => X0 = X2) [rectify 15]
30. ? [X0] : (main_query_0(X0) <~> main_query_1(X0)) [ennf transformation 4]
48. ! [X0,X1,X2] : (X0 = X2 | (X0 != X1 | X0 != X1)) [ennf transformation 29]
49. ! [X0,X1,X2] : (X0 = X2 | X0 != X1 | X0 != X1) [flattening 48]
50. ! [X0] : ((main_query_0(X0) | ! [X1,X2,X3,X4,X5] : (b_fst != b_snd | b_fst != a_snd | ~tab(X0,X4,X5) | ~tab(X1,X2,X3))) & (? [X1,X2,X3,X4,X5] : (b_fst = b_snd & b_fst = a_snd & tab(X0,X4,X5) & tab(X1,X2,X3)) | ~main_query_0(X0))) [nnf transformation 16]
51. ! [X0] : ((main_query_0(X0) | ! [X1,X2,X3,X4,X5] : (b_fst != b_snd | b_fst != a_snd | ~tab(X0,X4,X5) | ~tab(X1,X2,X3))) & (? [X6,X7,X8,X9,X10] : (b_fst = b_snd & b_fst = a_snd & tab(X0,X9,X10) & tab(X6,X7,X8)) | ~main_query_0(X0))) [rectify 50]
52. ! [X0] : (? [X6,X7,X8,X9,X10] : (b_fst = b_snd & b_fst = a_snd & tab(X0,X9,X10) & tab(X6,X7,X8)) => (b_fst = b_snd & b_fst = a_snd & tab(X0,sK3(X0),sK4(X0)) & tab(sK0(X0),sK1(X0),sK2(X0)))) [choice axiom]
53. ! [X0] : ((main_query_0(X0) | ! [X1,X2,X3,X4,X5] : (b_fst != b_snd | b_fst != a_snd | ~tab(X0,X4,X5) | ~tab(X1,X2,X3))) & ((b_fst = b_snd & b_fst = a_snd & tab(X0,sK3(X0),sK4(X0)) & tab(sK0(X0),sK1(X0),sK2(X0))) | ~main_query_0(X0))) [skolemisation 51,52]
54. ! [X0] : ((main_query_1(X0) | ! [X1,X2] : (b_fst != b_snd | ~tab(X0,X1,X2))) & (? [X1,X2] : (b_fst = b_snd & tab(X0,X1,X2)) | ~main_query_1(X0))) [nnf transformation 18]
55. ! [X0] : ((main_query_1(X0) | ! [X1,X2] : (b_fst != b_snd | ~tab(X0,X1,X2))) & (? [X3,X4] : (b_fst = b_snd & tab(X0,X3,X4)) | ~main_query_1(X0))) [rectify 54]
56. ! [X0] : (? [X3,X4] : (b_fst = b_snd & tab(X0,X3,X4)) => (b_fst = b_snd & tab(X0,sK5(X0),sK6(X0)))) [choice axiom]
57. ! [X0] : ((main_query_1(X0) | ! [X1,X2] : (b_fst != b_snd | ~tab(X0,X1,X2))) & ((b_fst = b_snd & tab(X0,sK5(X0),sK6(X0))) | ~main_query_1(X0))) [skolemisation 55,56]
58. ? [X0] : ((~main_query_1(X0) | ~main_query_0(X0)) & (main_query_1(X0) | main_query_0(X0))) [nnf transformation 30]
59. ? [X0] : ((~main_query_1(X0) | ~main_query_0(X0)) & (main_query_1(X0) | main_query_0(X0))) => ((~main_query_1(sK7) | ~main_query_0(sK7)) & (main_query_1(sK7) | main_query_0(sK7))) [choice axiom]
60. (~main_query_1(sK7) | ~main_query_0(sK7)) & (main_query_1(sK7) | main_query_0(sK7)) [skolemisation 58,59]
62. tab(X0,sK3(X0),sK4(X0)) | ~main_query_0(X0) [cnf transformation 53]
65. main_query_0(X0) | b_fst != b_snd | b_fst != a_snd | ~tab(X0,X4,X5) | ~tab(X1,X2,X3) [cnf transformation 53]
66. tab(X0,sK5(X0),sK6(X0)) | ~main_query_1(X0) [cnf transformation 57]
68. main_query_1(X0) | b_fst != b_snd | ~tab(X0,X1,X2) [cnf transformation 57]
69. main_query_1(sK7) | main_query_0(sK7) [cnf transformation 60]
70. ~main_query_1(sK7) | ~main_query_0(sK7) [cnf transformation 60]
81. X0 = X2 | X0 != X1 | X0 != X1 [cnf transformation 49]
91. X1 = X2 | X1 != X1 [equality resolution 81]
92. X1 = X2 [trivial inequality removal 91]
94. 0 <=> ! [X1,X3,X2] : ~tab(X1,X2,X3) [avatar definition]
95. ~tab(X1,X2,X3) <- {0} [avatar component clause 94]
100. 3 <=> b_fst != a_snd [avatar definition]
101. b_fst != a_snd <- {3} [avatar component clause 100]
106. 5 <=> b_fst != b_snd [avatar definition]
107. b_fst != b_snd <- {5} [avatar component clause 106]
109. 6 <=> ! [X5,X0,X4] : (main_query_0(X0) | ~tab(X0,X4,X5)) [avatar definition]
110. ~tab(X0,X4,X5) | main_query_0(X0) <- {6} [avatar component clause 109]
111. 0 | ~3 | ~5 | 6 [avatar split clause 65,109,106,100,94]
113. 8 <=> ! [X0] : ~main_query_0(X0) [avatar definition]
114. ~main_query_0(X0) <- {8} [avatar component clause 113]
118. 10 <=> ! [X1,X0,X2] : (main_query_1(X0) | ~tab(X0,X1,X2)) [avatar definition]
119. ~tab(X0,X1,X2) | main_query_1(X0) <- {10} [avatar component clause 118]
120. ~5 | 10 [avatar split clause 68,118,106]
122. 12 <=> ! [X0] : ~main_query_1(X0) [avatar definition]
123. ~main_query_1(X0) <- {12} [avatar component clause 122]
126. 14 <=> main_query_0(sK7) [avatar definition]
127. main_query_0(sK7) <- {14} [avatar component clause 126]
129. 15 <=> ~main_query_0(sK7) [avatar definition]
130. ~main_query_0(sK7) <- {15} [avatar component clause 129]
132. 16 <=> main_query_1(sK7) [avatar definition]
133. main_query_1(sK7) <- {16} [avatar component clause 132]
135. 17 <=> ~main_query_1(sK7) [avatar definition]
136. ~main_query_1(sK7) <- {17} [avatar component clause 135]
137. ~15 | ~17 [avatar split clause 70,135,129]
138. 14 | 16 [avatar split clause 69,132,126]
140. main_query_0(X0) <- {14} [superposition 127,92]
141. ~main_query_1(X0) <- {17} [superposition 136,92]
144. 12 | 17 [avatar split clause 141,135,122]
145. ~main_query_0(X0) | main_query_1(X0) <- {10} [resolution 62,119]
151. main_query_1(X0) <- {10, 14} [subsumption resolution 145,140]
152. $false <- {10, 12, 14} [subsumption resolution 151,123]
153. ~10 | ~12 | ~14 [avatar contradiction clause 152]
158. $false <- {5} [subsumption resolution 107,92]
159. 5 [avatar contradiction clause 158]
165. ~main_query_0(X0) <- {15} [superposition 130,92]
167. 8 | 15 [avatar split clause 165,129,113]
168. main_query_1(X0) <- {16} [superposition 133,92]
172. ~main_query_1(X1) | main_query_0(X1) <- {6} [resolution 66,110]
177. main_query_0(X1) <- {6, 16} [subsumption resolution 172,168]
178. $false <- {6, 8, 16} [subsumption resolution 177,114]
179. ~6 | ~8 | ~16 [avatar contradiction clause 178]
182. $false <- {3} [subsumption resolution 101,92]
183. 3 [avatar contradiction clause 182]
185. ~main_query_1(X1) <- {0} [resolution 95,66]
186. $false <- {0, 16} [subsumption resolution 185,168]
187. ~0 | ~16 [avatar contradiction clause 186]
188. $false [avatar sat refutation 111,120,137,138,144,153,159,167,179,183,187]
% SZS output end Proof for test9
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.086 s
% ------------------------------
% ------------------------------
