% Refutation found. Thanks to Tanya!
% SZS status Theorem for test1_symm
% SZS output start Proof for test1_symm
1. ! [X0,X1] : (main_query_0(X0,X1) <=> ? [X2,X3,X4,X5] : (lessthan(X2,X4) & b_fst = a_snd & a_fst = b_snd & $true & tab(X1,X4,X5) & tab(X0,X2,X3))) [input]
2. ! [X0,X1] : (main_query_1(X0,X1) <=> ? [X2,X3,X4,X5,X6,X7,X8] : (lessthan(X2,X4) & c_fst = c_fst & b_fst = a_snd & a_fst = b_snd & $true & tab(X6,X7,X8) & tab(X1,X4,X5) & tab(X0,X2,X3))) [input]
3. ! [X0,X1] : (main_query_0(X0,X1) <=> main_query_1(X0,X1)) [input]
4. ~! [X0,X1] : (main_query_0(X0,X1) <=> main_query_1(X0,X1)) [negated conjecture 3]
15. ! [X13,X14,X15] : ((X13 = X14 & X13 = X14) => X13 = X15) [input]
16. ! [X0,X1] : (main_query_0(X0,X1) <=> ? [X2,X3,X4,X5] : (lessthan(X2,X4) & b_fst = a_snd & a_fst = b_snd & tab(X1,X4,X5) & tab(X0,X2,X3))) [true and false elimination 1]
17. ! [X0,X1] : (main_query_1(X0,X1) <=> ? [X2,X3,X4,X5,X6,X7,X8] : (lessthan(X2,X4) & c_fst = c_fst & b_fst = a_snd & a_fst = b_snd & tab(X6,X7,X8) & tab(X1,X4,X5) & tab(X0,X2,X3))) [true and false elimination 2]
28. ! [X0,X1,X2] : ((X0 = X1 & X0 = X1) => X0 = X2) [rectify 15]
29. ? [X0,X1] : (main_query_0(X0,X1) <~> main_query_1(X0,X1)) [ennf transformation 4]
47. ! [X0,X1,X2] : (X0 = X2 | (X0 != X1 | X0 != X1)) [ennf transformation 28]
48. ! [X0,X1,X2] : (X0 = X2 | X0 != X1 | X0 != X1) [flattening 47]
49. ! [X0,X1] : ((main_query_0(X0,X1) | ! [X2,X3,X4,X5] : (~lessthan(X2,X4) | b_fst != a_snd | a_fst != b_snd | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & (? [X2,X3,X4,X5] : (lessthan(X2,X4) & b_fst = a_snd & a_fst = b_snd & tab(X1,X4,X5) & tab(X0,X2,X3)) | ~main_query_0(X0,X1))) [nnf transformation 16]
50. ! [X0,X1] : ((main_query_0(X0,X1) | ! [X2,X3,X4,X5] : (~lessthan(X2,X4) | b_fst != a_snd | a_fst != b_snd | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & (? [X6,X7,X8,X9] : (lessthan(X6,X8) & b_fst = a_snd & a_fst = b_snd & tab(X1,X8,X9) & tab(X0,X6,X7)) | ~main_query_0(X0,X1))) [rectify 49]
51. ! [X1,X0] : (? [X6,X7,X8,X9] : (lessthan(X6,X8) & b_fst = a_snd & a_fst = b_snd & tab(X1,X8,X9) & tab(X0,X6,X7)) => (lessthan(sK0(X0,X1),sK2(X0,X1)) & b_fst = a_snd & a_fst = b_snd & tab(X1,sK2(X0,X1),sK3(X0,X1)) & tab(X0,sK0(X0,X1),sK1(X0,X1)))) [choice axiom]
52. ! [X0,X1] : ((main_query_0(X0,X1) | ! [X2,X3,X4,X5] : (~lessthan(X2,X4) | b_fst != a_snd | a_fst != b_snd | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & ((lessthan(sK0(X0,X1),sK2(X0,X1)) & b_fst = a_snd & a_fst = b_snd & tab(X1,sK2(X0,X1),sK3(X0,X1)) & tab(X0,sK0(X0,X1),sK1(X0,X1))) | ~main_query_0(X0,X1))) [skolemisation 50,51]
53. ! [X0,X1] : ((main_query_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8] : (~lessthan(X2,X4) | c_fst != c_fst | b_fst != a_snd | a_fst != b_snd | ~tab(X6,X7,X8) | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & (? [X2,X3,X4,X5,X6,X7,X8] : (lessthan(X2,X4) & c_fst = c_fst & b_fst = a_snd & a_fst = b_snd & tab(X6,X7,X8) & tab(X1,X4,X5) & tab(X0,X2,X3)) | ~main_query_1(X0,X1))) [nnf transformation 17]
54. ! [X0,X1] : ((main_query_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8] : (~lessthan(X2,X4) | c_fst != c_fst | b_fst != a_snd | a_fst != b_snd | ~tab(X6,X7,X8) | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & (? [X9,X10,X11,X12,X13,X14,X15] : (lessthan(X9,X11) & c_fst = c_fst & b_fst = a_snd & a_fst = b_snd & tab(X13,X14,X15) & tab(X1,X11,X12) & tab(X0,X9,X10)) | ~main_query_1(X0,X1))) [rectify 53]
55. ! [X1,X0] : (? [X9,X10,X11,X12,X13,X14,X15] : (lessthan(X9,X11) & c_fst = c_fst & b_fst = a_snd & a_fst = b_snd & tab(X13,X14,X15) & tab(X1,X11,X12) & tab(X0,X9,X10)) => (lessthan(sK4(X0,X1),sK6(X0,X1)) & c_fst = c_fst & b_fst = a_snd & a_fst = b_snd & tab(sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) & tab(X1,sK6(X0,X1),sK7(X0,X1)) & tab(X0,sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
56. ! [X0,X1] : ((main_query_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8] : (~lessthan(X2,X4) | c_fst != c_fst | b_fst != a_snd | a_fst != b_snd | ~tab(X6,X7,X8) | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & ((lessthan(sK4(X0,X1),sK6(X0,X1)) & c_fst = c_fst & b_fst = a_snd & a_fst = b_snd & tab(sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) & tab(X1,sK6(X0,X1),sK7(X0,X1)) & tab(X0,sK4(X0,X1),sK5(X0,X1))) | ~main_query_1(X0,X1))) [skolemisation 54,55]
57. ? [X0,X1] : ((~main_query_1(X0,X1) | ~main_query_0(X0,X1)) & (main_query_1(X0,X1) | main_query_0(X0,X1))) [nnf transformation 29]
58. ? [X0,X1] : ((~main_query_1(X0,X1) | ~main_query_0(X0,X1)) & (main_query_1(X0,X1) | main_query_0(X0,X1))) => ((~main_query_1(sK11,sK12) | ~main_query_0(sK11,sK12)) & (main_query_1(sK11,sK12) | main_query_0(sK11,sK12))) [choice axiom]
59. (~main_query_1(sK11,sK12) | ~main_query_0(sK11,sK12)) & (main_query_1(sK11,sK12) | main_query_0(sK11,sK12)) [skolemisation 57,58]
60. tab(X0,sK0(X0,X1),sK1(X0,X1)) | ~main_query_0(X0,X1) [cnf transformation 52]
64. lessthan(sK0(X0,X1),sK2(X0,X1)) | ~main_query_0(X0,X1) [cnf transformation 52]
65. main_query_0(X0,X1) | ~lessthan(X2,X4) | b_fst != a_snd | a_fst != b_snd | ~tab(X1,X4,X5) | ~tab(X0,X2,X3) [cnf transformation 52]
66. tab(X0,sK4(X0,X1),sK5(X0,X1)) | ~main_query_1(X0,X1) [cnf transformation 56]
72. lessthan(sK4(X0,X1),sK6(X0,X1)) | ~main_query_1(X0,X1) [cnf transformation 56]
73. main_query_1(X0,X1) | ~lessthan(X2,X4) | c_fst != c_fst | b_fst != a_snd | a_fst != b_snd | ~tab(X6,X7,X8) | ~tab(X1,X4,X5) | ~tab(X0,X2,X3) [cnf transformation 56]
74. main_query_1(sK11,sK12) | main_query_0(sK11,sK12) [cnf transformation 59]
75. ~main_query_1(sK11,sK12) | ~main_query_0(sK11,sK12) [cnf transformation 59]
86. X0 = X2 | X0 != X1 | X0 != X1 [cnf transformation 48]
96. X1 = X2 | X1 != X1 [equality resolution 86]
97. X1 = X2 [trivial inequality removal 96]
98. main_query_1(X0,X1) | ~lessthan(X2,X4) | b_fst != a_snd | a_fst != b_snd | ~tab(X6,X7,X8) | ~tab(X1,X4,X5) | ~tab(X0,X2,X3) [trivial inequality removal 73]
103. 1 <=> a_fst != b_snd [avatar definition]
104. a_fst != b_snd <- {1} [avatar component clause 103]
109. 3 <=> b_fst != a_snd [avatar definition]
110. b_fst != a_snd <- {3} [avatar component clause 109]
112. 4 <=> ! [X1,X3,X5,X0,X2,X4] : (main_query_0(X0,X1) | ~tab(X0,X2,X3) | ~tab(X1,X4,X5) | ~lessthan(X2,X4)) [avatar definition]
113. ~tab(X1,X4,X5) | ~tab(X0,X2,X3) | main_query_0(X0,X1) | ~lessthan(X2,X4) <- {4} [avatar component clause 112]
114. ~1 | ~3 | 4 [avatar split clause 65,112,109,103]
116. 6 <=> ! [X1,X0] : ~main_query_0(X0,X1) [avatar definition]
117. ~main_query_0(X0,X1) <- {6} [avatar component clause 116]
121. 8 <=> ! [X7,X8,X6] : ~tab(X6,X7,X8) [avatar definition]
122. ~tab(X6,X7,X8) <- {8} [avatar component clause 121]
124. 10 <=> ! [X1,X3,X5,X0,X2,X4] : (main_query_1(X0,X1) | ~tab(X0,X2,X3) | ~tab(X1,X4,X5) | ~lessthan(X2,X4)) [avatar definition]
125. ~tab(X1,X4,X5) | ~tab(X0,X2,X3) | main_query_1(X0,X1) | ~lessthan(X2,X4) <- {10} [avatar component clause 124]
126. 8 | ~1 | ~3 | 10 [avatar split clause 98,124,109,103,121]
128. 12 <=> ! [X1,X0] : ~main_query_1(X0,X1) [avatar definition]
129. ~main_query_1(X0,X1) <- {12} [avatar component clause 128]
133. 14 <=> main_query_0(sK11,sK12) [avatar definition]
134. main_query_0(sK11,sK12) <- {14} [avatar component clause 133]
136. 15 <=> ~main_query_0(sK11,sK12) [avatar definition]
137. ~main_query_0(sK11,sK12) <- {15} [avatar component clause 136]
139. 16 <=> main_query_1(sK11,sK12) [avatar definition]
140. main_query_1(sK11,sK12) <- {16} [avatar component clause 139]
142. 17 <=> ~main_query_1(sK11,sK12) [avatar definition]
143. ~main_query_1(sK11,sK12) <- {17} [avatar component clause 142]
144. ~15 | ~17 [avatar split clause 75,142,136]
145. 14 | 16 [avatar split clause 74,139,133]
146. $false <- {1} [subsumption resolution 104,97]
147. 1 [avatar contradiction clause 146]
148. $false <- {3} [subsumption resolution 110,97]
149. 3 [avatar contradiction clause 148]
157. main_query_0(sK11,X0) <- {14} [superposition 134,97]
164. $false <- {6, 14} [subsumption resolution 157,117]
165. ~6 | ~14 [avatar contradiction clause 164]
166. $false <- {12, 16} [subsumption resolution 140,129]
167. ~12 | ~16 [avatar contradiction clause 166]
168. main_query_1(X0,sK12) <- {16} [superposition 140,97]
172. main_query_1(X1,X0) <- {16} [superposition 168,97]
175. lessthan(X2,sK2(X0,X1)) | ~main_query_0(X0,X1) [superposition 64,97]
181. lessthan(X2,sK6(X0,X1)) | ~main_query_1(X0,X1) [superposition 72,97]
186. ~main_query_0(X0,X1) <- {8} [resolution 60,122]
200. tab(X0,X2,sK5(X0,X1)) | ~main_query_1(X0,X1) [superposition 66,97]
202. tab(X0,sK4(X0,X1),X2) | ~main_query_1(X0,X1) [superposition 66,97]
214. ~main_query_1(X0,X1) <- {8} [subsumption resolution 202,122]
215. $false <- {8, 16} [subsumption resolution 214,172]
216. ~8 | ~16 [avatar contradiction clause 215]
234. 6 | ~8 [avatar split clause 186,121,116]
240. main_query_0(X0,sK12) <- {14} [superposition 134,97]
244. ~main_query_1(X0,sK12) <- {17} [superposition 143,97]
248. main_query_0(X1,X0) <- {14} [superposition 240,97]
263. ~main_query_1(X1,X0) <- {17} [superposition 244,97]
265. 12 | 17 [avatar split clause 263,142,128]
273. ~tab(X0,X1,X2) | main_query_1(X0,X3) | ~lessthan(X1,sK0(X3,X4)) | ~main_query_0(X3,X4) <- {10} [resolution 125,60]
278. ~tab(X0,X1,X2) | ~lessthan(X1,sK0(X3,X4)) | ~main_query_0(X3,X4) <- {10, 12} [subsumption resolution 273,129]
279. ~lessthan(X1,sK0(X3,X4)) | ~tab(X0,X1,X2) <- {10, 12, 14} [subsumption resolution 278,248]
282. lessthan(X3,X2) | ~main_query_0(X0,X1) [superposition 175,97]
284. lessthan(X3,X2) <- {14} [subsumption resolution 282,248]
287. lessthan(X3,X2) | ~main_query_1(X0,X1) [superposition 181,97]
295. ~lessthan(X3,X2) | ~tab(X4,X3,X5) <- {10, 12, 14} [superposition 279,97]
300. ~tab(X4,X3,X5) <- {10, 12, 14} [subsumption resolution 295,284]
301. 8 | ~10 | ~12 | ~14 [avatar split clause 300,133,128,124,121]
305. 18 <=> ! [X3,X2] : lessthan(X2,X3) [avatar definition]
306. lessthan(X2,X3) <- {18} [avatar component clause 305]
310. 12 | 18 [avatar split clause 287,305,128]
314. 6 | 18 [avatar split clause 282,305,116]
317. ~main_query_0(X0,sK12) <- {15} [superposition 137,97]
321. main_query_1(X0,sK12) <- {16} [superposition 140,97]
326. ~main_query_0(X1,X0) <- {15} [superposition 317,97]
328. 6 | 15 [avatar split clause 326,136,116]
331. main_query_1(X1,X0) <- {16} [superposition 321,97]
352. ~main_query_1(X6,X7) | ~tab(X8,X9,X10) | main_query_0(X8,X6) | ~lessthan(X9,X11) <- {4} [resolution 200,113]
355. ~tab(X8,X9,X10) | main_query_0(X8,X6) | ~lessthan(X9,X11) <- {4, 16} [subsumption resolution 352,331]
356. ~tab(X8,X9,X10) | ~lessthan(X9,X11) <- {4, 6, 16} [subsumption resolution 355,117]
357. ~tab(X8,X9,X10) <- {4, 6, 16, 18} [subsumption resolution 356,306]
358. 8 | ~4 | ~6 | ~16 | ~18 [avatar split clause 357,305,139,116,112,121]
360. $false [avatar sat refutation 114,126,144,145,147,149,165,167,216,234,265,301,310,314,328,358]
% SZS output end Proof for test1_symm
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.101 s
% ------------------------------
% ------------------------------
