% Refutation found. Thanks to Tanya!
% SZS status Theorem for test1_symm
% SZS output start Proof for test1_symm
1. ! [X0,X1] : (main_query_0(X0,X1) <=> ? [X2,X3,X4,X5] : (lessthan(X2,X4) & b_fst = a_snd & a_fst = b_snd & $true & tab(X1,X4,X5) & tab(X0,X2,X3))) [input]
2. ! [X0,X1] : (main_query_1(X0,X1) <=> ? [X2,X3,X4,X5,X6,X7,X8] : (lessthan(X2,X4) & c_fst = c_fst & b_fst = a_snd & a_fst = b_snd & $true & tab(X6,X7,X8) & tab(X1,X4,X5) & tab(X0,X2,X3))) [input]
3. ! [X0,X1] : (main_query_0(X0,X1) <=> main_query_1(X0,X1)) [input]
4. ~! [X0,X1] : (main_query_0(X0,X1) <=> main_query_1(X0,X1)) [negated conjecture 3]
5. ! [X0,X1] : (main_query_0(X0,X1) <=> ? [X2,X3,X4,X5] : (lessthan(X2,X4) & b_fst = a_snd & a_fst = b_snd & tab(X1,X4,X5) & tab(X0,X2,X3))) [true and false elimination 1]
6. ! [X0,X1] : (main_query_1(X0,X1) <=> ? [X2,X3,X4,X5,X6,X7,X8] : (lessthan(X2,X4) & c_fst = c_fst & b_fst = a_snd & a_fst = b_snd & tab(X6,X7,X8) & tab(X1,X4,X5) & tab(X0,X2,X3))) [true and false elimination 2]
7. ? [X0,X1] : (main_query_0(X0,X1) <~> main_query_1(X0,X1)) [ennf transformation 4]
8. ! [X0,X1] : ((main_query_0(X0,X1) | ! [X2,X3,X4,X5] : (~lessthan(X2,X4) | b_fst != a_snd | a_fst != b_snd | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & (? [X2,X3,X4,X5] : (lessthan(X2,X4) & b_fst = a_snd & a_fst = b_snd & tab(X1,X4,X5) & tab(X0,X2,X3)) | ~main_query_0(X0,X1))) [nnf transformation 5]
9. ! [X0,X1] : ((main_query_0(X0,X1) | ! [X2,X3,X4,X5] : (~lessthan(X2,X4) | b_fst != a_snd | a_fst != b_snd | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & (? [X6,X7,X8,X9] : (lessthan(X6,X8) & b_fst = a_snd & a_fst = b_snd & tab(X1,X8,X9) & tab(X0,X6,X7)) | ~main_query_0(X0,X1))) [rectify 8]
10. ! [X1,X0] : (? [X6,X7,X8,X9] : (lessthan(X6,X8) & b_fst = a_snd & a_fst = b_snd & tab(X1,X8,X9) & tab(X0,X6,X7)) => (lessthan(sK0(X0,X1),sK2(X0,X1)) & b_fst = a_snd & a_fst = b_snd & tab(X1,sK2(X0,X1),sK3(X0,X1)) & tab(X0,sK0(X0,X1),sK1(X0,X1)))) [choice axiom]
11. ! [X0,X1] : ((main_query_0(X0,X1) | ! [X2,X3,X4,X5] : (~lessthan(X2,X4) | b_fst != a_snd | a_fst != b_snd | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & ((lessthan(sK0(X0,X1),sK2(X0,X1)) & b_fst = a_snd & a_fst = b_snd & tab(X1,sK2(X0,X1),sK3(X0,X1)) & tab(X0,sK0(X0,X1),sK1(X0,X1))) | ~main_query_0(X0,X1))) [skolemisation 9,10]
12. ! [X0,X1] : ((main_query_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8] : (~lessthan(X2,X4) | c_fst != c_fst | b_fst != a_snd | a_fst != b_snd | ~tab(X6,X7,X8) | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & (? [X2,X3,X4,X5,X6,X7,X8] : (lessthan(X2,X4) & c_fst = c_fst & b_fst = a_snd & a_fst = b_snd & tab(X6,X7,X8) & tab(X1,X4,X5) & tab(X0,X2,X3)) | ~main_query_1(X0,X1))) [nnf transformation 6]
13. ! [X0,X1] : ((main_query_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8] : (~lessthan(X2,X4) | c_fst != c_fst | b_fst != a_snd | a_fst != b_snd | ~tab(X6,X7,X8) | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & (? [X9,X10,X11,X12,X13,X14,X15] : (lessthan(X9,X11) & c_fst = c_fst & b_fst = a_snd & a_fst = b_snd & tab(X13,X14,X15) & tab(X1,X11,X12) & tab(X0,X9,X10)) | ~main_query_1(X0,X1))) [rectify 12]
14. ! [X1,X0] : (? [X9,X10,X11,X12,X13,X14,X15] : (lessthan(X9,X11) & c_fst = c_fst & b_fst = a_snd & a_fst = b_snd & tab(X13,X14,X15) & tab(X1,X11,X12) & tab(X0,X9,X10)) => (lessthan(sK4(X0,X1),sK6(X0,X1)) & c_fst = c_fst & b_fst = a_snd & a_fst = b_snd & tab(sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) & tab(X1,sK6(X0,X1),sK7(X0,X1)) & tab(X0,sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
15. ! [X0,X1] : ((main_query_1(X0,X1) | ! [X2,X3,X4,X5,X6,X7,X8] : (~lessthan(X2,X4) | c_fst != c_fst | b_fst != a_snd | a_fst != b_snd | ~tab(X6,X7,X8) | ~tab(X1,X4,X5) | ~tab(X0,X2,X3))) & ((lessthan(sK4(X0,X1),sK6(X0,X1)) & c_fst = c_fst & b_fst = a_snd & a_fst = b_snd & tab(sK8(X0,X1),sK9(X0,X1),sK10(X0,X1)) & tab(X1,sK6(X0,X1),sK7(X0,X1)) & tab(X0,sK4(X0,X1),sK5(X0,X1))) | ~main_query_1(X0,X1))) [skolemisation 13,14]
16. ? [X0,X1] : ((~main_query_1(X0,X1) | ~main_query_0(X0,X1)) & (main_query_1(X0,X1) | main_query_0(X0,X1))) [nnf transformation 7]
17. ? [X0,X1] : ((~main_query_1(X0,X1) | ~main_query_0(X0,X1)) & (main_query_1(X0,X1) | main_query_0(X0,X1))) => ((~main_query_1(sK11,sK12) | ~main_query_0(sK11,sK12)) & (main_query_1(sK11,sK12) | main_query_0(sK11,sK12))) [choice axiom]
18. (~main_query_1(sK11,sK12) | ~main_query_0(sK11,sK12)) & (main_query_1(sK11,sK12) | main_query_0(sK11,sK12)) [skolemisation 16,17]
19. tab(X0,sK0(X0,X1),sK1(X0,X1)) | ~main_query_0(X0,X1) [cnf transformation 11]
20. tab(X1,sK2(X0,X1),sK3(X0,X1)) | ~main_query_0(X0,X1) [cnf transformation 11]
21. a_fst = b_snd | ~main_query_0(X0,X1) [cnf transformation 11]
22. b_fst = a_snd | ~main_query_0(X0,X1) [cnf transformation 11]
23. lessthan(sK0(X0,X1),sK2(X0,X1)) | ~main_query_0(X0,X1) [cnf transformation 11]
24. main_query_0(X0,X1) | ~lessthan(X2,X4) | b_fst != a_snd | a_fst != b_snd | ~tab(X1,X4,X5) | ~tab(X0,X2,X3) [cnf transformation 11]
25. tab(X0,sK4(X0,X1),sK5(X0,X1)) | ~main_query_1(X0,X1) [cnf transformation 15]
26. tab(X1,sK6(X0,X1),sK7(X0,X1)) | ~main_query_1(X0,X1) [cnf transformation 15]
28. a_fst = b_snd | ~main_query_1(X0,X1) [cnf transformation 15]
29. b_fst = a_snd | ~main_query_1(X0,X1) [cnf transformation 15]
31. lessthan(sK4(X0,X1),sK6(X0,X1)) | ~main_query_1(X0,X1) [cnf transformation 15]
32. main_query_1(X0,X1) | ~lessthan(X2,X4) | c_fst != c_fst | b_fst != a_snd | a_fst != b_snd | ~tab(X6,X7,X8) | ~tab(X1,X4,X5) | ~tab(X0,X2,X3) [cnf transformation 15]
33. main_query_1(sK11,sK12) | main_query_0(sK11,sK12) [cnf transformation 18]
34. ~main_query_1(sK11,sK12) | ~main_query_0(sK11,sK12) [cnf transformation 18]
35. main_query_1(X0,X1) | ~lessthan(X2,X4) | b_fst != a_snd | a_fst != b_snd | ~tab(X6,X7,X8) | ~tab(X1,X4,X5) | ~tab(X0,X2,X3) [trivial inequality removal 32]
37. 0 <=> a_fst = b_snd [avatar definition]
40. 1 <=> a_fst != b_snd [avatar definition]
43. 2 <=> b_fst = a_snd [avatar definition]
46. 3 <=> b_fst != a_snd [avatar definition]
49. 4 <=> ! [X1,X3,X5,X0,X2,X4] : (main_query_0(X0,X1) | ~tab(X0,X2,X3) | ~tab(X1,X4,X5) | ~lessthan(X2,X4)) [avatar definition]
50. ~tab(X1,X4,X5) | ~tab(X0,X2,X3) | main_query_0(X0,X1) | ~lessthan(X2,X4) <- {4} [avatar component clause 49]
51. ~1 | ~3 | 4 [avatar split clause 24,49,46,40]
53. 6 <=> ! [X1,X0] : ~main_query_0(X0,X1) [avatar definition]
54. ~main_query_0(X0,X1) <- {6} [avatar component clause 53]
55. 6 | 2 [avatar split clause 22,43,53]
56. 6 | 0 [avatar split clause 21,37,53]
58. 8 <=> ! [X7,X8,X6] : ~tab(X6,X7,X8) [avatar definition]
59. ~tab(X6,X7,X8) <- {8} [avatar component clause 58]
61. 10 <=> ! [X1,X3,X5,X0,X2,X4] : (main_query_1(X0,X1) | ~tab(X0,X2,X3) | ~tab(X1,X4,X5) | ~lessthan(X2,X4)) [avatar definition]
62. ~tab(X1,X4,X5) | ~tab(X0,X2,X3) | main_query_1(X0,X1) | ~lessthan(X2,X4) <- {10} [avatar component clause 61]
63. 8 | ~1 | ~3 | 10 [avatar split clause 35,61,46,40,58]
65. 12 <=> ! [X1,X0] : ~main_query_1(X0,X1) [avatar definition]
66. ~main_query_1(X0,X1) <- {12} [avatar component clause 65]
67. 12 | 2 [avatar split clause 29,43,65]
68. 12 | 0 [avatar split clause 28,37,65]
70. 14 <=> main_query_0(sK11,sK12) [avatar definition]
71. main_query_0(sK11,sK12) <- {14} [avatar component clause 70]
73. 15 <=> ~main_query_0(sK11,sK12) [avatar definition]
74. ~main_query_0(sK11,sK12) <- {15} [avatar component clause 73]
76. 16 <=> main_query_1(sK11,sK12) [avatar definition]
77. main_query_1(sK11,sK12) <- {16} [avatar component clause 76]
79. 17 <=> ~main_query_1(sK11,sK12) [avatar definition]
80. ~main_query_1(sK11,sK12) <- {17} [avatar component clause 79]
81. ~15 | ~17 [avatar split clause 34,79,73]
82. 14 | 16 [avatar split clause 33,76,70]
83. $false <- {6, 14} [resolution 54,71]
84. ~6 | ~14 [avatar contradiction clause 83]
85. $false <- {12, 16} [subsumption resolution 77,66]
86. ~12 | ~16 [avatar contradiction clause 85]
87. ~main_query_0(X0,X1) <- {8} [resolution 19,59]
88. 6 | ~8 [avatar split clause 87,58,53]
100. ~lessthan(X16,sK6(X19,X18)) | main_query_0(X15,X18) | ~tab(X15,X16,X17) | ~main_query_1(X19,X18) <- {4} [resolution 50,26]
106. ~lessthan(X6,sK2(X9,X8)) | main_query_1(X5,X8) | ~tab(X5,X6,X7) | ~main_query_0(X9,X8) <- {10} [resolution 62,20]
114. main_query_0(X0,X1) | ~tab(X0,sK4(X2,X1),X3) | ~main_query_1(X2,X1) | ~main_query_1(X2,X1) <- {4} [resolution 100,31]
115. ~tab(X0,sK4(X2,X1),X3) | main_query_0(X0,X1) | ~main_query_1(X2,X1) <- {4} [duplicate literal removal 114]
116. main_query_0(X0,X1) | ~main_query_1(X0,X1) | ~main_query_1(X0,X1) <- {4} [resolution 115,25]
117. ~main_query_1(X0,X1) | main_query_0(X0,X1) <- {4} [duplicate literal removal 116]
118. main_query_1(X0,X1) | ~tab(X0,sK0(X2,X1),X3) | ~main_query_0(X2,X1) | ~main_query_0(X2,X1) <- {10} [resolution 106,23]
119. ~tab(X0,sK0(X2,X1),X3) | main_query_1(X0,X1) | ~main_query_0(X2,X1) <- {10} [duplicate literal removal 118]
120. main_query_1(X0,X1) | ~main_query_0(X0,X1) | ~main_query_0(X0,X1) <- {10} [resolution 119,19]
121. ~main_query_0(X0,X1) | main_query_1(X0,X1) <- {10} [duplicate literal removal 120]
122. main_query_1(sK11,sK12) <- {10, 14} [resolution 121,71]
123. $false <- {10, 14, 17} [subsumption resolution 122,80]
124. ~10 | ~14 | 17 [avatar contradiction clause 123]
125. main_query_0(sK11,sK12) <- {4, 16} [resolution 77,117]
126. $false <- {4, 15, 16} [subsumption resolution 125,74]
127. ~4 | 15 | ~16 [avatar contradiction clause 126]
128. $false [avatar sat refutation 51,55,56,63,67,68,81,82,84,86,88,124,127]
% SZS output end Proof for test1_symm
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.060 s
% ------------------------------
% ------------------------------
