% Refutation found. Thanks to Tanya!
% SZS status Theorem for output_query06_query07
% SZS output start Proof for output_query06_query07
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X5,X6] : (where_3(X5,X6) <=> eq(X5,X6)) [input]
3. ! [X5] : (select_1(X5) <=> ? [X0,X1,X2,X3,X4] : (where_3(X5,X6) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X5,X6] : (where_6(X5,X6) <=> eq(X5,X6)) [input]
7. ! [X5] : (select_4(X5) <=> ? [X0,X1,X2,X3,X4] : (where_6(X5,X6) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X7] : (main_query_2(X7) <=> select_4(X7)) [input]
50. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
51. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 50]
52. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [rectify 2]
53. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : (where_3(X0,X6) & from_2(X1,X2,X3,X4,X5))) [rectify 3]
54. ! [X6] : ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : (where_3(X0,X6) & from_2(X1,X2,X3,X4,X5))) [closure 53]
55. ! [X6,X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : (where_3(X0,X6) & from_2(X1,X2,X3,X4,X5))) [flattening 54]
56. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
57. ! [X0,X1] : (where_6(X0,X1) <=> eq(X0,X1)) [rectify 6]
58. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4,X5] : (where_6(X0,X6) & from_5(X1,X2,X3,X4,X5))) [rectify 7]
59. ! [X6] : ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4,X5] : (where_6(X0,X6) & from_5(X1,X2,X3,X4,X5))) [closure 58]
60. ! [X6,X0] : (select_4(X0) <=> ? [X1,X2,X3,X4,X5] : (where_6(X0,X6) & from_5(X1,X2,X3,X4,X5))) [flattening 59]
61. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
103. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 51]
181. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 103]
182. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
183. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 52]
184. ! [X6,X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : (~where_3(X0,X6) | ~from_2(X1,X2,X3,X4,X5))) & (? [X1,X2,X3,X4,X5] : (where_3(X0,X6) & from_2(X1,X2,X3,X4,X5)) | ~select_1(X0))) [nnf transformation 55]
185. ! [X0,X1] : ((select_1(X1) | ! [X2,X3,X4,X5,X6] : (~where_3(X1,X0) | ~from_2(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_3(X1,X0) & from_2(X7,X8,X9,X10,X11)) | ~select_1(X1))) [rectify 184]
186. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_3(X1,X0) & from_2(X7,X8,X9,X10,X11)) => (where_3(X1,X0) & from_2(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)))) [choice axiom]
187. ! [X0,X1] : ((select_1(X1) | ! [X2,X3,X4,X5,X6] : (~where_3(X1,X0) | ~from_2(X2,X3,X4,X5,X6))) & ((where_3(X1,X0) & from_2(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1))) | ~select_1(X1))) [skolemisation 185,186]
188. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 56]
189. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
190. ! [X0,X1] : ((where_6(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_6(X0,X1))) [nnf transformation 57]
191. ! [X6,X0] : ((select_4(X0) | ! [X1,X2,X3,X4,X5] : (~where_6(X0,X6) | ~from_5(X1,X2,X3,X4,X5))) & (? [X1,X2,X3,X4,X5] : (where_6(X0,X6) & from_5(X1,X2,X3,X4,X5)) | ~select_4(X0))) [nnf transformation 60]
192. ! [X0,X1] : ((select_4(X1) | ! [X2,X3,X4,X5,X6] : (~where_6(X1,X0) | ~from_5(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_6(X1,X0) & from_5(X7,X8,X9,X10,X11)) | ~select_4(X1))) [rectify 191]
193. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_6(X1,X0) & from_5(X7,X8,X9,X10,X11)) => (where_6(X1,X0) & from_5(sK5(X0,X1),sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
194. ! [X0,X1] : ((select_4(X1) | ! [X2,X3,X4,X5,X6] : (~where_6(X1,X0) | ~from_5(X2,X3,X4,X5,X6))) & ((where_6(X1,X0) & from_5(sK5(X0,X1),sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X1))) [skolemisation 192,193]
195. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 61]
196. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 181]
197. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK10) | ~main_query_1(sK10)) & (main_query_2(sK10) | main_query_1(sK10))) [choice axiom]
198. (~main_query_2(sK10) | ~main_query_1(sK10)) & (main_query_2(sK10) | main_query_1(sK10)) [skolemisation 196,197]
199. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 182]
200. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 182]
201. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 183]
202. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 183]
203. from_2(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_1(X1) [cnf transformation 187]
204. where_3(X1,X0) | ~select_1(X1) [cnf transformation 187]
205. select_1(X1) | ~where_3(X1,X0) | ~from_2(X2,X3,X4,X5,X6) [cnf transformation 187]
206. ~main_query_1(X0) | select_1(X0) [cnf transformation 188]
207. ~select_1(X0) | main_query_1(X0) [cnf transformation 188]
208. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 189]
209. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 189]
210. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 190]
211. ~eq(X0,X1) | where_6(X0,X1) [cnf transformation 190]
212. from_5(sK5(X0,X1),sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X1) [cnf transformation 194]
213. where_6(X1,X0) | ~select_4(X1) [cnf transformation 194]
214. select_4(X1) | ~where_6(X1,X0) | ~from_5(X2,X3,X4,X5,X6) [cnf transformation 194]
215. ~main_query_2(X0) | select_4(X0) [cnf transformation 195]
216. ~select_4(X0) | main_query_2(X0) [cnf transformation 195]
258. main_query_2(sK10) | main_query_1(sK10) [cnf transformation 198]
259. ~main_query_2(sK10) | ~main_query_1(sK10) [cnf transformation 198]
299. 0 <=> ! [X3,X5,X2,X4,X6] : ~from_2(X2,X3,X4,X5,X6) [avatar definition]
300. ~from_2(X2,X3,X4,X5,X6) <- {0} [avatar component clause 299]
302. 2 <=> ! [X1,X0] : (select_1(X1) | ~where_3(X1,X0)) [avatar definition]
303. ~where_3(X1,X0) | select_1(X1) <- {2} [avatar component clause 302]
304. 0 | 2 [avatar split clause 205,302,299]
306. 4 <=> ! [X3,X5,X2,X4,X6] : ~from_5(X2,X3,X4,X5,X6) [avatar definition]
307. ~from_5(X2,X3,X4,X5,X6) <- {4} [avatar component clause 306]
309. 6 <=> ! [X1,X0] : (select_4(X1) | ~where_6(X1,X0)) [avatar definition]
310. ~where_6(X1,X0) | select_4(X1) <- {6} [avatar component clause 309]
311. 4 | 6 [avatar split clause 214,309,306]
313. 8 <=> main_query_1(sK10) [avatar definition]
314. main_query_1(sK10) <- {8} [avatar component clause 313]
316. 9 <=> ~main_query_1(sK10) [avatar definition]
317. ~main_query_1(sK10) <- {9} [avatar component clause 316]
319. 10 <=> main_query_2(sK10) [avatar definition]
320. main_query_2(sK10) <- {10} [avatar component clause 319]
322. 11 <=> ~main_query_2(sK10) [avatar definition]
323. ~main_query_2(sK10) <- {11} [avatar component clause 322]
324. ~9 | ~11 [avatar split clause 259,322,316]
325. 8 | 10 [avatar split clause 258,319,313]
326. select_1(sK10) <- {8} [resolution 206,314]
329. eq(X0,X1) | ~select_1(X0) [resolution 201,204]
331. eq(X0,X1) | ~select_4(X0) [resolution 210,213]
332. where_6(X0,X1) | ~select_1(X0) [resolution 211,329]
334. where_3(X2,X3) | ~select_4(X2) [resolution 331,202]
338. ~select_4(X2) | select_1(X2) <- {2} [resolution 334,303]
340. tab1(sK0(X1,X0),sK1(X1,X0),sK2(X1,X0),sK3(X1,X0),sK4(X1,X0)) | ~select_1(X0) [resolution 203,199]
341. tab1(sK5(X1,X0),sK6(X1,X0),sK7(X1,X0),sK8(X1,X0),sK9(X1,X0)) | ~select_4(X0) [resolution 212,208]
343. from_5(sK0(X1,X0),sK1(X1,X0),sK2(X1,X0),sK3(X1,X0),sK4(X1,X0)) | ~select_1(X0) [resolution 340,209]
345. ~select_1(X0) <- {4} [subsumption resolution 343,307]
347. from_2(sK5(X3,X2),sK6(X3,X2),sK7(X3,X2),sK8(X3,X2),sK9(X3,X2)) | ~select_4(X2) [resolution 341,200]
348. $false <- {4, 8} [resolution 345,326]
349. ~4 | ~8 [avatar contradiction clause 348]
353. select_1(sK10) <- {8} [resolution 314,206]
354. ~select_1(X0) | select_4(X0) <- {6} [resolution 310,332]
358. select_4(sK10) <- {6, 8} [resolution 354,353]
361. main_query_2(sK10) <- {6, 8} [resolution 358,216]
362. $false <- {6, 8, 11} [subsumption resolution 361,323]
363. ~6 | ~8 | 11 [avatar contradiction clause 362]
364. select_4(sK10) <- {10} [resolution 320,215]
365. select_1(sK10) <- {2, 10} [resolution 364,338]
368. main_query_1(sK10) <- {2, 10} [resolution 365,207]
369. $false <- {2, 9, 10} [subsumption resolution 368,317]
370. ~2 | 9 | ~10 [avatar contradiction clause 369]
372. ~select_4(X1) <- {0} [resolution 300,347]
373. $false <- {0, 10} [resolution 372,364]
374. ~0 | ~10 [avatar contradiction clause 373]
375. $false [avatar sat refutation 304,311,324,325,349,363,370,374]
% SZS output end Proof for output_query06_query07
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.029 s
% ------------------------------
% ------------------------------
