% Refutation found. Thanks to Tanya!
% SZS status Theorem for output_query06_query10
% SZS output start Proof for output_query06_query10
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X5,X6] : (where_3(X5,X6) <=> eq(X5,X6)) [input]
3. ! [X5] : (select_1(X5) <=> ? [X0,X1,X2,X3,X4] : (where_3(X5,X6) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X5] : (main_query_1(X5) <=> select_1(X5)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X6,X0] : (where_6(X6,X0) <=> eq(X0,X6)) [input]
7. ! [X0] : (select_4(X0) <=> ? [X7,X8,X9,X10,X11] : (where_6(X6,X0) & from_5(X7,X8,X9,X10,X11))) [input]
8. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
50. ! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [input]
51. ~! [X5] : (main_query_1(X5) <=> main_query_2(X5)) [negated conjecture 50]
52. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [rectify 2]
53. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : (where_3(X0,X6) & from_2(X1,X2,X3,X4,X5))) [rectify 3]
54. ! [X6] : ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : (where_3(X0,X6) & from_2(X1,X2,X3,X4,X5))) [closure 53]
55. ! [X6,X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : (where_3(X0,X6) & from_2(X1,X2,X3,X4,X5))) [flattening 54]
56. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
57. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
58. ! [X0,X1] : (where_6(X0,X1) <=> eq(X1,X0)) [rectify 6]
59. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4,X5] : (where_6(X6,X0) & from_5(X1,X2,X3,X4,X5))) [rectify 7]
60. ! [X6] : ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4,X5] : (where_6(X6,X0) & from_5(X1,X2,X3,X4,X5))) [closure 59]
61. ! [X6,X0] : (select_4(X0) <=> ? [X1,X2,X3,X4,X5] : (where_6(X6,X0) & from_5(X1,X2,X3,X4,X5))) [flattening 60]
62. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
104. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 51]
182. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 104]
183. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
184. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 52]
185. ! [X6,X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : (~where_3(X0,X6) | ~from_2(X1,X2,X3,X4,X5))) & (? [X1,X2,X3,X4,X5] : (where_3(X0,X6) & from_2(X1,X2,X3,X4,X5)) | ~select_1(X0))) [nnf transformation 55]
186. ! [X0,X1] : ((select_1(X1) | ! [X2,X3,X4,X5,X6] : (~where_3(X1,X0) | ~from_2(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_3(X1,X0) & from_2(X7,X8,X9,X10,X11)) | ~select_1(X1))) [rectify 185]
187. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_3(X1,X0) & from_2(X7,X8,X9,X10,X11)) => (where_3(X1,X0) & from_2(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)))) [choice axiom]
188. ! [X0,X1] : ((select_1(X1) | ! [X2,X3,X4,X5,X6] : (~where_3(X1,X0) | ~from_2(X2,X3,X4,X5,X6))) & ((where_3(X1,X0) & from_2(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1))) | ~select_1(X1))) [skolemisation 186,187]
189. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 56]
190. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 57]
191. ! [X0,X1] : ((where_6(X0,X1) | ~eq(X1,X0)) & (eq(X1,X0) | ~where_6(X0,X1))) [nnf transformation 58]
192. ! [X6,X0] : ((select_4(X0) | ! [X1,X2,X3,X4,X5] : (~where_6(X6,X0) | ~from_5(X1,X2,X3,X4,X5))) & (? [X1,X2,X3,X4,X5] : (where_6(X6,X0) & from_5(X1,X2,X3,X4,X5)) | ~select_4(X0))) [nnf transformation 61]
193. ! [X0,X1] : ((select_4(X1) | ! [X2,X3,X4,X5,X6] : (~where_6(X0,X1) | ~from_5(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_6(X0,X1) & from_5(X7,X8,X9,X10,X11)) | ~select_4(X1))) [rectify 192]
194. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_6(X0,X1) & from_5(X7,X8,X9,X10,X11)) => (where_6(X0,X1) & from_5(sK5(X0,X1),sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
195. ! [X0,X1] : ((select_4(X1) | ! [X2,X3,X4,X5,X6] : (~where_6(X0,X1) | ~from_5(X2,X3,X4,X5,X6))) & ((where_6(X0,X1) & from_5(sK5(X0,X1),sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X1))) [skolemisation 193,194]
196. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 62]
197. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 182]
198. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK10) | ~main_query_1(sK10)) & (main_query_2(sK10) | main_query_1(sK10))) [choice axiom]
199. (~main_query_2(sK10) | ~main_query_1(sK10)) & (main_query_2(sK10) | main_query_1(sK10)) [skolemisation 197,198]
200. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 183]
201. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 183]
202. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 184]
203. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 184]
204. from_2(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_1(X1) [cnf transformation 188]
205. where_3(X1,X0) | ~select_1(X1) [cnf transformation 188]
206. select_1(X1) | ~where_3(X1,X0) | ~from_2(X2,X3,X4,X5,X6) [cnf transformation 188]
207. ~main_query_1(X0) | select_1(X0) [cnf transformation 189]
208. ~select_1(X0) | main_query_1(X0) [cnf transformation 189]
209. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 190]
210. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 190]
211. ~where_6(X0,X1) | eq(X1,X0) [cnf transformation 191]
212. ~eq(X1,X0) | where_6(X0,X1) [cnf transformation 191]
213. from_5(sK5(X0,X1),sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X1) [cnf transformation 195]
214. where_6(X0,X1) | ~select_4(X1) [cnf transformation 195]
215. select_4(X1) | ~where_6(X0,X1) | ~from_5(X2,X3,X4,X5,X6) [cnf transformation 195]
216. ~main_query_2(X0) | select_4(X0) [cnf transformation 196]
217. ~select_4(X0) | main_query_2(X0) [cnf transformation 196]
259. main_query_2(sK10) | main_query_1(sK10) [cnf transformation 199]
260. ~main_query_2(sK10) | ~main_query_1(sK10) [cnf transformation 199]
300. 0 <=> ! [X3,X5,X2,X4,X6] : ~from_2(X2,X3,X4,X5,X6) [avatar definition]
301. ~from_2(X2,X3,X4,X5,X6) <- {0} [avatar component clause 300]
303. 2 <=> ! [X1,X0] : (select_1(X1) | ~where_3(X1,X0)) [avatar definition]
304. ~where_3(X1,X0) | select_1(X1) <- {2} [avatar component clause 303]
305. 0 | 2 [avatar split clause 206,303,300]
307. 4 <=> ! [X3,X5,X2,X4,X6] : ~from_5(X2,X3,X4,X5,X6) [avatar definition]
308. ~from_5(X2,X3,X4,X5,X6) <- {4} [avatar component clause 307]
310. 6 <=> ! [X1,X0] : (select_4(X1) | ~where_6(X0,X1)) [avatar definition]
311. ~where_6(X0,X1) | select_4(X1) <- {6} [avatar component clause 310]
312. 4 | 6 [avatar split clause 215,310,307]
314. 8 <=> main_query_1(sK10) [avatar definition]
315. main_query_1(sK10) <- {8} [avatar component clause 314]
317. 9 <=> ~main_query_1(sK10) [avatar definition]
318. ~main_query_1(sK10) <- {9} [avatar component clause 317]
320. 10 <=> main_query_2(sK10) [avatar definition]
321. main_query_2(sK10) <- {10} [avatar component clause 320]
323. 11 <=> ~main_query_2(sK10) [avatar definition]
324. ~main_query_2(sK10) <- {11} [avatar component clause 323]
325. ~9 | ~11 [avatar split clause 260,323,317]
326. 8 | 10 [avatar split clause 259,320,314]
327. select_1(sK10) <- {8} [resolution 207,315]
330. eq(X0,X1) | ~select_1(X0) [resolution 202,205]
332. eq(X0,X1) | ~select_4(X0) [resolution 211,214]
333. where_6(X0,X1) | ~select_1(X1) [resolution 212,330]
335. where_3(X2,X3) | ~select_4(X2) [resolution 332,203]
339. ~select_4(X2) | select_1(X2) <- {2} [resolution 335,304]
341. tab1(sK0(X1,X0),sK1(X1,X0),sK2(X1,X0),sK3(X1,X0),sK4(X1,X0)) | ~select_1(X0) [resolution 204,200]
342. tab1(sK5(X1,X0),sK6(X1,X0),sK7(X1,X0),sK8(X1,X0),sK9(X1,X0)) | ~select_4(X0) [resolution 213,209]
344. from_5(sK0(X1,X0),sK1(X1,X0),sK2(X1,X0),sK3(X1,X0),sK4(X1,X0)) | ~select_1(X0) [resolution 341,210]
346. ~select_1(X0) <- {4} [subsumption resolution 344,308]
348. from_2(sK5(X3,X2),sK6(X3,X2),sK7(X3,X2),sK8(X3,X2),sK9(X3,X2)) | ~select_4(X2) [resolution 342,201]
349. $false <- {4, 8} [resolution 346,327]
350. ~4 | ~8 [avatar contradiction clause 349]
354. select_1(sK10) <- {8} [resolution 315,207]
355. ~select_1(X0) | select_4(X0) <- {6} [resolution 311,333]
359. select_4(sK10) <- {6, 8} [resolution 355,354]
362. main_query_2(sK10) <- {6, 8} [resolution 359,217]
363. $false <- {6, 8, 11} [subsumption resolution 362,324]
364. ~6 | ~8 | 11 [avatar contradiction clause 363]
365. select_4(sK10) <- {10} [resolution 321,216]
366. select_1(sK10) <- {2, 10} [resolution 365,339]
369. main_query_1(sK10) <- {2, 10} [resolution 366,208]
370. $false <- {2, 9, 10} [subsumption resolution 369,318]
371. ~2 | 9 | ~10 [avatar contradiction clause 370]
373. ~select_4(X1) <- {0} [resolution 301,348]
374. $false <- {0, 10} [resolution 373,365]
375. ~0 | ~10 [avatar contradiction clause 374]
376. $false [avatar sat refutation 305,312,325,326,350,364,371,375]
% SZS output end Proof for output_query06_query10
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.028 s
% ------------------------------
% ------------------------------
