% Refutation found. Thanks to Tanya!
% SZS status Theorem for output_query09_query10
% SZS output start Proof for output_query09_query10
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X5] : (where_3(X5) <=> eq(X5,X5)) [input]
3. ! [X5] : (select_1(X5) <=> ? [X0,X1,X2,X3,X4] : (where_3(X5) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X6] : (main_query_1(X6) <=> select_1(X6)) [input]
5. ! [X7,X8,X9,X10,X11] : (from_5(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
6. ! [X12,X0] : (where_6(X12,X0) <=> eq(X0,X12)) [input]
7. ! [X0] : (select_4(X0) <=> ? [X7,X8,X9,X10,X11] : (where_6(X12,X0) & from_5(X7,X8,X9,X10,X11))) [input]
8. ! [X5] : (main_query_2(X5) <=> select_4(X5)) [input]
48. ! [X6] : (main_query_1(X6) <=> main_query_2(X6)) [input]
49. ~! [X6] : (main_query_1(X6) <=> main_query_2(X6)) [negated conjecture 48]
50. ! [X0] : (where_3(X0) <=> eq(X0,X0)) [rectify 2]
51. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : (where_3(X0) & from_2(X1,X2,X3,X4,X5))) [rectify 3]
52. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
53. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
54. ! [X0,X1] : (where_6(X0,X1) <=> eq(X1,X0)) [rectify 6]
55. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4,X5] : (where_6(X6,X0) & from_5(X1,X2,X3,X4,X5))) [rectify 7]
56. ! [X6] : ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4,X5] : (where_6(X6,X0) & from_5(X1,X2,X3,X4,X5))) [closure 55]
57. ! [X6,X0] : (select_4(X0) <=> ? [X1,X2,X3,X4,X5] : (where_6(X6,X0) & from_5(X1,X2,X3,X4,X5))) [flattening 56]
58. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
98. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 49]
172. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 98]
173. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
174. ! [X0] : ((where_3(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_3(X0))) [nnf transformation 50]
175. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : (~where_3(X0) | ~from_2(X1,X2,X3,X4,X5))) & (? [X1,X2,X3,X4,X5] : (where_3(X0) & from_2(X1,X2,X3,X4,X5)) | ~select_1(X0))) [nnf transformation 51]
176. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : (~where_3(X0) | ~from_2(X1,X2,X3,X4,X5))) & (? [X6,X7,X8,X9,X10] : (where_3(X0) & from_2(X6,X7,X8,X9,X10)) | ~select_1(X0))) [rectify 175]
177. ! [X0] : (? [X6,X7,X8,X9,X10] : (where_3(X0) & from_2(X6,X7,X8,X9,X10)) => (where_3(X0) & from_2(sK0(X0),sK1(X0),sK2(X0),sK3(X0),sK4(X0)))) [choice axiom]
178. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : (~where_3(X0) | ~from_2(X1,X2,X3,X4,X5))) & ((where_3(X0) & from_2(sK0(X0),sK1(X0),sK2(X0),sK3(X0),sK4(X0))) | ~select_1(X0))) [skolemisation 176,177]
179. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 52]
180. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 53]
181. ! [X0,X1] : ((where_6(X0,X1) | ~eq(X1,X0)) & (eq(X1,X0) | ~where_6(X0,X1))) [nnf transformation 54]
182. ! [X6,X0] : ((select_4(X0) | ! [X1,X2,X3,X4,X5] : (~where_6(X6,X0) | ~from_5(X1,X2,X3,X4,X5))) & (? [X1,X2,X3,X4,X5] : (where_6(X6,X0) & from_5(X1,X2,X3,X4,X5)) | ~select_4(X0))) [nnf transformation 57]
183. ! [X0,X1] : ((select_4(X1) | ! [X2,X3,X4,X5,X6] : (~where_6(X0,X1) | ~from_5(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_6(X0,X1) & from_5(X7,X8,X9,X10,X11)) | ~select_4(X1))) [rectify 182]
184. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_6(X0,X1) & from_5(X7,X8,X9,X10,X11)) => (where_6(X0,X1) & from_5(sK5(X0,X1),sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)))) [choice axiom]
185. ! [X0,X1] : ((select_4(X1) | ! [X2,X3,X4,X5,X6] : (~where_6(X0,X1) | ~from_5(X2,X3,X4,X5,X6))) & ((where_6(X0,X1) & from_5(sK5(X0,X1),sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1))) | ~select_4(X1))) [skolemisation 183,184]
186. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 58]
187. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 172]
188. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK10) | ~main_query_1(sK10)) & (main_query_2(sK10) | main_query_1(sK10))) [choice axiom]
189. (~main_query_2(sK10) | ~main_query_1(sK10)) & (main_query_2(sK10) | main_query_1(sK10)) [skolemisation 187,188]
190. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 173]
191. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 173]
192. eq(X0,X0) | ~where_3(X0) [cnf transformation 174]
193. ~eq(X0,X0) | where_3(X0) [cnf transformation 174]
194. from_2(sK0(X0),sK1(X0),sK2(X0),sK3(X0),sK4(X0)) | ~select_1(X0) [cnf transformation 178]
195. ~select_1(X0) | where_3(X0) [cnf transformation 178]
196. select_1(X0) | ~where_3(X0) | ~from_2(X1,X2,X3,X4,X5) [cnf transformation 178]
197. ~main_query_1(X0) | select_1(X0) [cnf transformation 179]
198. ~select_1(X0) | main_query_1(X0) [cnf transformation 179]
199. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 180]
200. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 180]
201. ~where_6(X0,X1) | eq(X1,X0) [cnf transformation 181]
202. ~eq(X1,X0) | where_6(X0,X1) [cnf transformation 181]
203. from_5(sK5(X0,X1),sK6(X0,X1),sK7(X0,X1),sK8(X0,X1),sK9(X0,X1)) | ~select_4(X1) [cnf transformation 185]
204. where_6(X0,X1) | ~select_4(X1) [cnf transformation 185]
205. select_4(X1) | ~where_6(X0,X1) | ~from_5(X2,X3,X4,X5,X6) [cnf transformation 185]
206. ~main_query_2(X0) | select_4(X0) [cnf transformation 186]
207. ~select_4(X0) | main_query_2(X0) [cnf transformation 186]
247. main_query_2(sK10) | main_query_1(sK10) [cnf transformation 189]
248. ~main_query_2(sK10) | ~main_query_1(sK10) [cnf transformation 189]
286. 0 <=> ! [X1,X3,X5,X2,X4] : ~from_2(X1,X2,X3,X4,X5) [avatar definition]
287. ~from_2(X1,X2,X3,X4,X5) <- {0} [avatar component clause 286]
289. 2 <=> ! [X0] : (select_1(X0) | ~where_3(X0)) [avatar definition]
290. ~where_3(X0) | select_1(X0) <- {2} [avatar component clause 289]
291. 0 | 2 [avatar split clause 196,289,286]
293. 4 <=> ! [X3,X5,X2,X4,X6] : ~from_5(X2,X3,X4,X5,X6) [avatar definition]
294. ~from_5(X2,X3,X4,X5,X6) <- {4} [avatar component clause 293]
296. 6 <=> ! [X1,X0] : (select_4(X1) | ~where_6(X0,X1)) [avatar definition]
297. ~where_6(X0,X1) | select_4(X1) <- {6} [avatar component clause 296]
298. 4 | 6 [avatar split clause 205,296,293]
300. 8 <=> main_query_1(sK10) [avatar definition]
301. main_query_1(sK10) <- {8} [avatar component clause 300]
303. 9 <=> ~main_query_1(sK10) [avatar definition]
304. ~main_query_1(sK10) <- {9} [avatar component clause 303]
306. 10 <=> main_query_2(sK10) [avatar definition]
307. main_query_2(sK10) <- {10} [avatar component clause 306]
309. 11 <=> ~main_query_2(sK10) [avatar definition]
310. ~main_query_2(sK10) <- {11} [avatar component clause 309]
311. ~9 | ~11 [avatar split clause 248,309,303]
312. 8 | 10 [avatar split clause 247,306,300]
313. select_1(sK10) <- {8} [resolution 197,301]
318. eq(X0,X1) | ~select_4(X0) [resolution 201,204]
319. where_6(X0,X0) | ~where_3(X0) [resolution 202,192]
320. ~select_4(X0) | where_3(X0) [resolution 318,193]
325. tab1(sK0(X0),sK1(X0),sK2(X0),sK3(X0),sK4(X0)) | ~select_1(X0) [resolution 194,190]
326. tab1(sK5(X1,X0),sK6(X1,X0),sK7(X1,X0),sK8(X1,X0),sK9(X1,X0)) | ~select_4(X0) [resolution 203,199]
328. from_5(sK0(X0),sK1(X0),sK2(X0),sK3(X0),sK4(X0)) | ~select_1(X0) [resolution 325,200]
330. ~select_1(X0) <- {4} [subsumption resolution 328,294]
331. $false <- {4, 8} [resolution 330,313]
332. ~4 | ~8 [avatar contradiction clause 331]
336. select_1(sK10) <- {8} [resolution 301,197]
338. ~where_3(X1) | select_4(X1) <- {6} [resolution 297,319]
340. where_3(sK10) <- {8} [resolution 336,195]
342. from_2(sK5(X3,X2),sK6(X3,X2),sK7(X3,X2),sK8(X3,X2),sK9(X3,X2)) | ~select_4(X2) [resolution 326,191]
344. select_4(sK10) <- {6, 8} [resolution 338,340]
346. main_query_2(sK10) <- {6, 8} [resolution 344,207]
347. $false <- {6, 8, 11} [subsumption resolution 346,310]
348. ~6 | ~8 | 11 [avatar contradiction clause 347]
349. select_4(sK10) <- {10} [resolution 307,206]
350. where_3(sK10) <- {10} [resolution 349,320]
353. select_1(sK10) <- {2, 10} [resolution 350,290]
355. main_query_1(sK10) <- {2, 10} [resolution 353,198]
357. $false <- {2, 9, 10} [subsumption resolution 355,304]
358. ~2 | 9 | ~10 [avatar contradiction clause 357]
360. ~select_4(X0) <- {0} [resolution 342,287]
362. $false <- {0, 10} [resolution 360,349]
363. ~0 | ~10 [avatar contradiction clause 362]
364. $false [avatar sat refutation 291,298,311,312,332,348,358,363]
% SZS output end Proof for output_query09_query10
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.030 s
% ------------------------------
% ------------------------------
