% Refutation found. Thanks to Tanya!
% SZS status Theorem for output_query08_query11
% SZS output start Proof for output_query08_query11
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X5,X6] : (where_3(X5,X6) <=> eq(X6,X5)) [input]
3. ! [X6] : (select_1(X6) <=> ? [X0,X1,X2,X3,X4] : (where_3(X5,X6) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X7] : (main_query_1(X7) <=> select_1(X7)) [input]
5. ! [X8,X9,X10,X11,X12] : (from_5(X8,X9,X10,X11,X12) <=> tab1(X8,X9,X10,X11,X12)) [input]
6. ! [X0,X9] : (where_6(X0,X9) <=> eq(X0,X9)) [input]
7. ! [X0] : (select_4(X0) <=> ? [X8,X9,X10,X11,X12] : (where_6(X0,X9) & from_5(X8,X9,X10,X11,X12))) [input]
8. ! [X6] : (main_query_2(X6) <=> select_4(X6)) [input]
49. ! [X7] : (main_query_1(X7) <=> main_query_2(X7)) [input]
50. ~! [X7] : (main_query_1(X7) <=> main_query_2(X7)) [negated conjecture 49]
51. ! [X0,X1] : (where_3(X0,X1) <=> eq(X1,X0)) [rectify 2]
52. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : (where_3(X6,X0) & from_2(X1,X2,X3,X4,X5))) [rectify 3]
53. ! [X6] : ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : (where_3(X6,X0) & from_2(X1,X2,X3,X4,X5))) [closure 52]
54. ! [X6,X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : (where_3(X6,X0) & from_2(X1,X2,X3,X4,X5))) [flattening 53]
55. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
56. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 5]
57. ! [X0,X1] : (where_6(X0,X1) <=> eq(X0,X1)) [rectify 6]
58. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4,X5] : (where_6(X0,X2) & from_5(X1,X2,X3,X4,X5))) [rectify 7]
59. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
100. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 50]
176. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 100]
177. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
178. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X1,X0)) & (eq(X1,X0) | ~where_3(X0,X1))) [nnf transformation 51]
179. ! [X6,X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : (~where_3(X6,X0) | ~from_2(X1,X2,X3,X4,X5))) & (? [X1,X2,X3,X4,X5] : (where_3(X6,X0) & from_2(X1,X2,X3,X4,X5)) | ~select_1(X0))) [nnf transformation 54]
180. ! [X0,X1] : ((select_1(X1) | ! [X2,X3,X4,X5,X6] : (~where_3(X0,X1) | ~from_2(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_3(X0,X1) & from_2(X7,X8,X9,X10,X11)) | ~select_1(X1))) [rectify 179]
181. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_3(X0,X1) & from_2(X7,X8,X9,X10,X11)) => (where_3(X0,X1) & from_2(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)))) [choice axiom]
182. ! [X0,X1] : ((select_1(X1) | ! [X2,X3,X4,X5,X6] : (~where_3(X0,X1) | ~from_2(X2,X3,X4,X5,X6))) & ((where_3(X0,X1) & from_2(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1))) | ~select_1(X1))) [skolemisation 180,181]
183. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 55]
184. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 56]
185. ! [X0,X1] : ((where_6(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_6(X0,X1))) [nnf transformation 57]
186. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4,X5] : (~where_6(X0,X2) | ~from_5(X1,X2,X3,X4,X5))) & (? [X1,X2,X3,X4,X5] : (where_6(X0,X2) & from_5(X1,X2,X3,X4,X5)) | ~select_4(X0))) [nnf transformation 58]
187. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4,X5] : (~where_6(X0,X2) | ~from_5(X1,X2,X3,X4,X5))) & (? [X6,X7,X8,X9,X10] : (where_6(X0,X7) & from_5(X6,X7,X8,X9,X10)) | ~select_4(X0))) [rectify 186]
188. ! [X0] : (? [X6,X7,X8,X9,X10] : (where_6(X0,X7) & from_5(X6,X7,X8,X9,X10)) => (where_6(X0,sK6(X0)) & from_5(sK5(X0),sK6(X0),sK7(X0),sK8(X0),sK9(X0)))) [choice axiom]
189. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4,X5] : (~where_6(X0,X2) | ~from_5(X1,X2,X3,X4,X5))) & ((where_6(X0,sK6(X0)) & from_5(sK5(X0),sK6(X0),sK7(X0),sK8(X0),sK9(X0))) | ~select_4(X0))) [skolemisation 187,188]
190. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 59]
191. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 176]
192. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK10) | ~main_query_1(sK10)) & (main_query_2(sK10) | main_query_1(sK10))) [choice axiom]
193. (~main_query_2(sK10) | ~main_query_1(sK10)) & (main_query_2(sK10) | main_query_1(sK10)) [skolemisation 191,192]
194. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 177]
195. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 177]
196. ~where_3(X0,X1) | eq(X1,X0) [cnf transformation 178]
197. ~eq(X1,X0) | where_3(X0,X1) [cnf transformation 178]
198. from_2(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_1(X1) [cnf transformation 182]
199. where_3(X0,X1) | ~select_1(X1) [cnf transformation 182]
200. select_1(X1) | ~where_3(X0,X1) | ~from_2(X2,X3,X4,X5,X6) [cnf transformation 182]
201. ~main_query_1(X0) | select_1(X0) [cnf transformation 183]
202. ~select_1(X0) | main_query_1(X0) [cnf transformation 183]
203. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 184]
204. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 184]
205. ~where_6(X0,X1) | eq(X0,X1) [cnf transformation 185]
206. ~eq(X0,X1) | where_6(X0,X1) [cnf transformation 185]
207. from_5(sK5(X0),sK6(X0),sK7(X0),sK8(X0),sK9(X0)) | ~select_4(X0) [cnf transformation 189]
208. where_6(X0,sK6(X0)) | ~select_4(X0) [cnf transformation 189]
209. ~from_5(X1,X2,X3,X4,X5) | ~where_6(X0,X2) | select_4(X0) [cnf transformation 189]
210. ~main_query_2(X0) | select_4(X0) [cnf transformation 190]
211. ~select_4(X0) | main_query_2(X0) [cnf transformation 190]
252. main_query_2(sK10) | main_query_1(sK10) [cnf transformation 193]
253. ~main_query_2(sK10) | ~main_query_1(sK10) [cnf transformation 193]
292. 0 <=> ! [X3,X5,X2,X4,X6] : ~from_2(X2,X3,X4,X5,X6) [avatar definition]
293. ~from_2(X2,X3,X4,X5,X6) <- {0} [avatar component clause 292]
295. 2 <=> ! [X1,X0] : (select_1(X1) | ~where_3(X0,X1)) [avatar definition]
296. ~where_3(X0,X1) | select_1(X1) <- {2} [avatar component clause 295]
297. 0 | 2 [avatar split clause 200,295,292]
299. 4 <=> main_query_1(sK10) [avatar definition]
300. main_query_1(sK10) <- {4} [avatar component clause 299]
302. 5 <=> ~main_query_1(sK10) [avatar definition]
303. ~main_query_1(sK10) <- {5} [avatar component clause 302]
305. 6 <=> main_query_2(sK10) [avatar definition]
306. main_query_2(sK10) <- {6} [avatar component clause 305]
308. 7 <=> ~main_query_2(sK10) [avatar definition]
309. ~main_query_2(sK10) <- {7} [avatar component clause 308]
310. ~5 | ~7 [avatar split clause 253,308,302]
311. 4 | 6 [avatar split clause 252,305,299]
312. select_1(sK10) <- {4} [resolution 201,300]
315. eq(X0,X1) | ~select_1(X0) [resolution 196,199]
317. where_6(X0,X1) | ~select_1(X0) [resolution 206,315]
318. eq(X0,sK6(X0)) | ~select_4(X0) [resolution 208,205]
322. where_3(sK6(X1),X1) | ~select_4(X1) [resolution 318,197]
325. ~select_4(X1) | select_1(X1) <- {2} [resolution 322,296]
326. tab1(sK5(X0),sK6(X0),sK7(X0),sK8(X0),sK9(X0)) | ~select_4(X0) [resolution 207,203]
327. ~where_6(X2,sK6(X1)) | ~select_4(X1) | select_4(X2) [resolution 207,209]
329. ~select_4(X1) | select_4(X2) | ~select_1(X2) [resolution 327,317]
332. 8 <=> ! [X2] : (select_4(X2) | ~select_1(X2)) [avatar definition]
333. ~select_1(X2) | select_4(X2) <- {8} [avatar component clause 332]
335. 10 <=> ! [X1] : ~select_4(X1) [avatar definition]
336. ~select_4(X1) <- {10} [avatar component clause 335]
337. 8 | 10 [avatar split clause 329,335,332]
338. tab1(sK0(X1,X0),sK1(X1,X0),sK2(X1,X0),sK3(X1,X0),sK4(X1,X0)) | ~select_1(X0) [resolution 198,194]
339. select_4(sK10) <- {4, 8} [resolution 333,312]
341. from_2(sK5(X1),sK6(X1),sK7(X1),sK8(X1),sK9(X1)) | ~select_4(X1) [resolution 326,195]
343. main_query_2(sK10) <- {4, 8} [resolution 339,211]
344. $false <- {4, 7, 8} [subsumption resolution 343,309]
345. ~4 | 7 | ~8 [avatar contradiction clause 344]
346. from_5(sK0(X1,X0),sK1(X1,X0),sK2(X1,X0),sK3(X1,X0),sK4(X1,X0)) | ~select_1(X0) [resolution 338,204]
350. ~select_1(X2) | ~where_6(X3,sK1(X4,X2)) | select_4(X3) [resolution 346,209]
351. ~where_6(X3,sK1(X4,X2)) | ~select_1(X2) <- {10} [subsumption resolution 350,336]
352. ~select_1(X0) | ~select_1(X1) <- {10} [resolution 351,317]
354. 12 <=> ! [X1] : ~select_1(X1) [avatar definition]
355. ~select_1(X1) <- {12} [avatar component clause 354]
356. 12 | 12 | ~10 [avatar split clause 352,335,354,354]
357. select_4(sK10) <- {6} [resolution 306,210]
358. select_1(sK10) <- {2, 6} [resolution 357,325]
361. main_query_1(sK10) <- {2, 6} [resolution 358,202]
362. $false <- {2, 5, 6} [subsumption resolution 361,303]
363. ~2 | 5 | ~6 [avatar contradiction clause 362]
364. ~select_1(X0) <- {0} [resolution 293,198]
365. ~select_4(X1) <- {0} [resolution 293,341]
366. 12 | ~0 [avatar split clause 364,292,354]
367. 10 | ~0 [avatar split clause 365,292,335]
368. $false <- {6, 10} [resolution 336,357]
369. ~6 | ~10 [avatar contradiction clause 368]
371. select_1(sK10) <- {4} [resolution 300,201]
372. $false <- {4, 12} [subsumption resolution 371,355]
373. ~4 | ~12 [avatar contradiction clause 372]
374. $false [avatar sat refutation 297,310,311,337,345,356,363,366,367,369,373]
% SZS output end Proof for output_query08_query11
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.024 s
% ------------------------------
% ------------------------------
