% Refutation found. Thanks to Tanya!
% SZS status Theorem for output_query01_query04
% SZS output start Proof for output_query01_query04
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X5] : (select_1(X5) <=> ? [X0,X1,X2,X3,X4] : ($true & from_2(X0,X1,X2,X3,X4))) [input]
3. ! [X6] : (main_query_1(X6) <=> select_1(X6)) [input]
4. ! [X7,X8,X9,X10,X11] : (from_4(X7,X8,X9,X10,X11) <=> tab1(X7,X8,X9,X10,X11)) [input]
5. ! [X0] : (select_3(X0) <=> ? [X7,X8,X9,X10,X11] : ($true & from_4(X7,X8,X9,X10,X11))) [input]
6. ! [X6] : (main_query_2(X6) <=> select_3(X6)) [input]
40. ! [X6] : (main_query_1(X6) <=> main_query_2(X6)) [input]
41. ~! [X6] : (main_query_1(X6) <=> main_query_2(X6)) [negated conjecture 40]
42. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : ($true & from_2(X1,X2,X3,X4,X5))) [rectify 2]
43. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : from_2(X1,X2,X3,X4,X5)) [true and false elimination 42]
44. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 3]
45. ! [X0,X1,X2,X3,X4] : (from_4(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [rectify 4]
46. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4,X5] : ($true & from_4(X1,X2,X3,X4,X5))) [rectify 5]
47. ! [X0] : (select_3(X0) <=> ? [X1,X2,X3,X4,X5] : from_4(X1,X2,X3,X4,X5)) [true and false elimination 46]
48. ! [X0] : (main_query_2(X0) <=> select_3(X0)) [rectify 6]
82. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 41]
144. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 82]
145. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
146. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : ~from_2(X1,X2,X3,X4,X5)) & (? [X1,X2,X3,X4,X5] : from_2(X1,X2,X3,X4,X5) | ~select_1(X0))) [nnf transformation 43]
147. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : ~from_2(X1,X2,X3,X4,X5)) & (? [X6,X7,X8,X9,X10] : from_2(X6,X7,X8,X9,X10) | ~select_1(X0))) [rectify 146]
148. ? [X6,X7,X8,X9,X10] : from_2(X6,X7,X8,X9,X10) => from_2(sK0,sK1,sK2,sK3,sK4) [choice axiom]
149. ! [X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : ~from_2(X1,X2,X3,X4,X5)) & (from_2(sK0,sK1,sK2,sK3,sK4) | ~select_1(X0))) [skolemisation 147,148]
150. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 44]
151. ! [X0,X1,X2,X3,X4] : ((from_4(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_4(X0,X1,X2,X3,X4))) [nnf transformation 45]
152. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4,X5] : ~from_4(X1,X2,X3,X4,X5)) & (? [X1,X2,X3,X4,X5] : from_4(X1,X2,X3,X4,X5) | ~select_3(X0))) [nnf transformation 47]
153. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4,X5] : ~from_4(X1,X2,X3,X4,X5)) & (? [X6,X7,X8,X9,X10] : from_4(X6,X7,X8,X9,X10) | ~select_3(X0))) [rectify 152]
154. ? [X6,X7,X8,X9,X10] : from_4(X6,X7,X8,X9,X10) => from_4(sK5,sK6,sK7,sK8,sK9) [choice axiom]
155. ! [X0] : ((select_3(X0) | ! [X1,X2,X3,X4,X5] : ~from_4(X1,X2,X3,X4,X5)) & (from_4(sK5,sK6,sK7,sK8,sK9) | ~select_3(X0))) [skolemisation 153,154]
156. ! [X0] : ((main_query_2(X0) | ~select_3(X0)) & (select_3(X0) | ~main_query_2(X0))) [nnf transformation 48]
157. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 144]
158. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK10) | ~main_query_1(sK10)) & (main_query_2(sK10) | main_query_1(sK10))) [choice axiom]
159. (~main_query_2(sK10) | ~main_query_1(sK10)) & (main_query_2(sK10) | main_query_1(sK10)) [skolemisation 157,158]
160. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 145]
161. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 145]
162. from_2(sK0,sK1,sK2,sK3,sK4) | ~select_1(X0) [cnf transformation 149]
163. select_1(X0) | ~from_2(X1,X2,X3,X4,X5) [cnf transformation 149]
164. ~main_query_1(X0) | select_1(X0) [cnf transformation 150]
165. ~select_1(X0) | main_query_1(X0) [cnf transformation 150]
166. ~from_4(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 151]
167. ~tab1(X0,X1,X2,X3,X4) | from_4(X0,X1,X2,X3,X4) [cnf transformation 151]
168. from_4(sK5,sK6,sK7,sK8,sK9) | ~select_3(X0) [cnf transformation 155]
169. select_3(X0) | ~from_4(X1,X2,X3,X4,X5) [cnf transformation 155]
170. ~main_query_2(X0) | select_3(X0) [cnf transformation 156]
171. ~select_3(X0) | main_query_2(X0) [cnf transformation 156]
205. main_query_2(sK10) | main_query_1(sK10) [cnf transformation 159]
206. ~main_query_2(sK10) | ~main_query_1(sK10) [cnf transformation 159]
238. 0 <=> ! [X1,X3,X5,X2,X4] : ~from_2(X1,X2,X3,X4,X5) [avatar definition]
239. ~from_2(X1,X2,X3,X4,X5) <- {0} [avatar component clause 238]
241. 2 <=> ! [X0] : select_1(X0) [avatar definition]
242. select_1(X0) <- {2} [avatar component clause 241]
243. 0 | 2 [avatar split clause 163,241,238]
245. 4 <=> ! [X0] : ~select_1(X0) [avatar definition]
246. ~select_1(X0) <- {4} [avatar component clause 245]
251. 6 <=> from_2(sK0,sK1,sK2,sK3,sK4) [avatar definition]
252. from_2(sK0,sK1,sK2,sK3,sK4) <- {6} [avatar component clause 251]
253. 4 | 6 [avatar split clause 162,251,245]
255. 8 <=> ! [X1,X3,X5,X2,X4] : ~from_4(X1,X2,X3,X4,X5) [avatar definition]
256. ~from_4(X1,X2,X3,X4,X5) <- {8} [avatar component clause 255]
258. 10 <=> ! [X0] : select_3(X0) [avatar definition]
259. select_3(X0) <- {10} [avatar component clause 258]
260. 8 | 10 [avatar split clause 169,258,255]
262. 12 <=> ! [X0] : ~select_3(X0) [avatar definition]
263. ~select_3(X0) <- {12} [avatar component clause 262]
268. 14 <=> from_4(sK5,sK6,sK7,sK8,sK9) [avatar definition]
269. from_4(sK5,sK6,sK7,sK8,sK9) <- {14} [avatar component clause 268]
270. 12 | 14 [avatar split clause 168,268,262]
272. 16 <=> main_query_1(sK10) [avatar definition]
273. main_query_1(sK10) <- {16} [avatar component clause 272]
275. 17 <=> ~main_query_1(sK10) [avatar definition]
276. ~main_query_1(sK10) <- {17} [avatar component clause 275]
278. 18 <=> main_query_2(sK10) [avatar definition]
279. main_query_2(sK10) <- {18} [avatar component clause 278]
281. 19 <=> ~main_query_2(sK10) [avatar definition]
282. ~main_query_2(sK10) <- {19} [avatar component clause 281]
283. ~17 | ~19 [avatar split clause 206,281,275]
284. 16 | 18 [avatar split clause 205,278,272]
287. select_1(sK10) <- {16} [resolution 164,273]
288. $false <- {4, 16} [subsumption resolution 287,246]
289. ~4 | ~16 [avatar contradiction clause 288]
290. select_3(sK10) <- {18} [resolution 170,279]
291. $false <- {12, 18} [subsumption resolution 290,263]
292. ~12 | ~18 [avatar contradiction clause 291]
298. main_query_2(X0) <- {10} [resolution 171,259]
302. tab1(sK5,sK6,sK7,sK8,sK9) <- {14} [resolution 166,269]
304. from_2(sK5,sK6,sK7,sK8,sK9) <- {14} [resolution 302,161]
305. $false <- {0, 14} [subsumption resolution 304,239]
306. ~0 | ~14 [avatar contradiction clause 305]
307. $false <- {10, 19} [subsumption resolution 282,298]
308. ~10 | 19 [avatar contradiction clause 307]
312. main_query_1(X0) <- {2} [resolution 242,165]
313. tab1(sK0,sK1,sK2,sK3,sK4) <- {6} [resolution 252,160]
315. from_4(sK0,sK1,sK2,sK3,sK4) <- {6} [resolution 313,167]
317. $false <- {6, 8} [subsumption resolution 315,256]
318. ~6 | ~8 [avatar contradiction clause 317]
319. $false <- {2, 17} [subsumption resolution 276,312]
320. ~2 | 17 [avatar contradiction clause 319]
321. $false [avatar sat refutation 243,253,260,270,283,284,289,292,306,308,318,320]
% SZS output end Proof for output_query01_query04
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.025 s
% ------------------------------
% ------------------------------
