% Refutation found. Thanks to Tanya!
% SZS status Theorem for output_query08_query09
% SZS output start Proof for output_query08_query09
1. ! [X0,X1,X2,X3,X4] : (from_2(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
2. ! [X5,X6] : (where_3(X5,X6) <=> eq(X5,X6)) [input]
3. ! [X5] : (select_1(X5) <=> ? [X0,X1,X2,X3,X4] : (where_3(X5,X6) & from_2(X0,X1,X2,X3,X4))) [input]
4. ! [X7] : (main_query_1(X7) <=> select_1(X7)) [input]
5. ! [X0,X1,X2,X3,X4] : (from_5(X0,X1,X2,X3,X4) <=> tab1(X0,X1,X2,X3,X4)) [input]
6. ! [X5] : (where_6(X5) <=> eq(X5,X5)) [input]
7. ! [X5] : (select_4(X5) <=> ? [X0,X1,X2,X3,X4] : (where_6(X5) & from_5(X0,X1,X2,X3,X4))) [input]
8. ! [X7] : (main_query_2(X7) <=> select_4(X7)) [input]
48. ! [X7] : (main_query_1(X7) <=> main_query_2(X7)) [input]
49. ~! [X7] : (main_query_1(X7) <=> main_query_2(X7)) [negated conjecture 48]
50. ! [X0,X1] : (where_3(X0,X1) <=> eq(X0,X1)) [rectify 2]
51. ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : (where_3(X0,X6) & from_2(X1,X2,X3,X4,X5))) [rectify 3]
52. ! [X6] : ! [X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : (where_3(X0,X6) & from_2(X1,X2,X3,X4,X5))) [closure 51]
53. ! [X6,X0] : (select_1(X0) <=> ? [X1,X2,X3,X4,X5] : (where_3(X0,X6) & from_2(X1,X2,X3,X4,X5))) [flattening 52]
54. ! [X0] : (main_query_1(X0) <=> select_1(X0)) [rectify 4]
55. ! [X0] : (where_6(X0) <=> eq(X0,X0)) [rectify 6]
56. ! [X0] : (select_4(X0) <=> ? [X1,X2,X3,X4,X5] : (where_6(X0) & from_5(X1,X2,X3,X4,X5))) [rectify 7]
57. ! [X0] : (main_query_2(X0) <=> select_4(X0)) [rectify 8]
97. ~! [X0] : (main_query_1(X0) <=> main_query_2(X0)) [rectify 49]
171. ? [X0] : (main_query_1(X0) <~> main_query_2(X0)) [ennf transformation 97]
172. ! [X0,X1,X2,X3,X4] : ((from_2(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_2(X0,X1,X2,X3,X4))) [nnf transformation 1]
173. ! [X0,X1] : ((where_3(X0,X1) | ~eq(X0,X1)) & (eq(X0,X1) | ~where_3(X0,X1))) [nnf transformation 50]
174. ! [X6,X0] : ((select_1(X0) | ! [X1,X2,X3,X4,X5] : (~where_3(X0,X6) | ~from_2(X1,X2,X3,X4,X5))) & (? [X1,X2,X3,X4,X5] : (where_3(X0,X6) & from_2(X1,X2,X3,X4,X5)) | ~select_1(X0))) [nnf transformation 53]
175. ! [X0,X1] : ((select_1(X1) | ! [X2,X3,X4,X5,X6] : (~where_3(X1,X0) | ~from_2(X2,X3,X4,X5,X6))) & (? [X7,X8,X9,X10,X11] : (where_3(X1,X0) & from_2(X7,X8,X9,X10,X11)) | ~select_1(X1))) [rectify 174]
176. ! [X1,X0] : (? [X7,X8,X9,X10,X11] : (where_3(X1,X0) & from_2(X7,X8,X9,X10,X11)) => (where_3(X1,X0) & from_2(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)))) [choice axiom]
177. ! [X0,X1] : ((select_1(X1) | ! [X2,X3,X4,X5,X6] : (~where_3(X1,X0) | ~from_2(X2,X3,X4,X5,X6))) & ((where_3(X1,X0) & from_2(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1))) | ~select_1(X1))) [skolemisation 175,176]
178. ! [X0] : ((main_query_1(X0) | ~select_1(X0)) & (select_1(X0) | ~main_query_1(X0))) [nnf transformation 54]
179. ! [X0,X1,X2,X3,X4] : ((from_5(X0,X1,X2,X3,X4) | ~tab1(X0,X1,X2,X3,X4)) & (tab1(X0,X1,X2,X3,X4) | ~from_5(X0,X1,X2,X3,X4))) [nnf transformation 5]
180. ! [X0] : ((where_6(X0) | ~eq(X0,X0)) & (eq(X0,X0) | ~where_6(X0))) [nnf transformation 55]
181. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4,X5] : (~where_6(X0) | ~from_5(X1,X2,X3,X4,X5))) & (? [X1,X2,X3,X4,X5] : (where_6(X0) & from_5(X1,X2,X3,X4,X5)) | ~select_4(X0))) [nnf transformation 56]
182. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4,X5] : (~where_6(X0) | ~from_5(X1,X2,X3,X4,X5))) & (? [X6,X7,X8,X9,X10] : (where_6(X0) & from_5(X6,X7,X8,X9,X10)) | ~select_4(X0))) [rectify 181]
183. ! [X0] : (? [X6,X7,X8,X9,X10] : (where_6(X0) & from_5(X6,X7,X8,X9,X10)) => (where_6(X0) & from_5(sK5(X0),sK6(X0),sK7(X0),sK8(X0),sK9(X0)))) [choice axiom]
184. ! [X0] : ((select_4(X0) | ! [X1,X2,X3,X4,X5] : (~where_6(X0) | ~from_5(X1,X2,X3,X4,X5))) & ((where_6(X0) & from_5(sK5(X0),sK6(X0),sK7(X0),sK8(X0),sK9(X0))) | ~select_4(X0))) [skolemisation 182,183]
185. ! [X0] : ((main_query_2(X0) | ~select_4(X0)) & (select_4(X0) | ~main_query_2(X0))) [nnf transformation 57]
186. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) [nnf transformation 171]
187. ? [X0] : ((~main_query_2(X0) | ~main_query_1(X0)) & (main_query_2(X0) | main_query_1(X0))) => ((~main_query_2(sK10) | ~main_query_1(sK10)) & (main_query_2(sK10) | main_query_1(sK10))) [choice axiom]
188. (~main_query_2(sK10) | ~main_query_1(sK10)) & (main_query_2(sK10) | main_query_1(sK10)) [skolemisation 186,187]
189. ~from_2(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 172]
190. ~tab1(X0,X1,X2,X3,X4) | from_2(X0,X1,X2,X3,X4) [cnf transformation 172]
191. ~where_3(X0,X1) | eq(X0,X1) [cnf transformation 173]
192. ~eq(X0,X1) | where_3(X0,X1) [cnf transformation 173]
193. from_2(sK0(X0,X1),sK1(X0,X1),sK2(X0,X1),sK3(X0,X1),sK4(X0,X1)) | ~select_1(X1) [cnf transformation 177]
194. where_3(X1,X0) | ~select_1(X1) [cnf transformation 177]
195. select_1(X1) | ~where_3(X1,X0) | ~from_2(X2,X3,X4,X5,X6) [cnf transformation 177]
196. ~main_query_1(X0) | select_1(X0) [cnf transformation 178]
197. ~select_1(X0) | main_query_1(X0) [cnf transformation 178]
198. ~from_5(X0,X1,X2,X3,X4) | tab1(X0,X1,X2,X3,X4) [cnf transformation 179]
199. ~tab1(X0,X1,X2,X3,X4) | from_5(X0,X1,X2,X3,X4) [cnf transformation 179]
200. eq(X0,X0) | ~where_6(X0) [cnf transformation 180]
201. ~eq(X0,X0) | where_6(X0) [cnf transformation 180]
202. from_5(sK5(X0),sK6(X0),sK7(X0),sK8(X0),sK9(X0)) | ~select_4(X0) [cnf transformation 184]
203. ~select_4(X0) | where_6(X0) [cnf transformation 184]
204. select_4(X0) | ~where_6(X0) | ~from_5(X1,X2,X3,X4,X5) [cnf transformation 184]
205. ~main_query_2(X0) | select_4(X0) [cnf transformation 185]
206. ~select_4(X0) | main_query_2(X0) [cnf transformation 185]
246. main_query_2(sK10) | main_query_1(sK10) [cnf transformation 188]
247. ~main_query_2(sK10) | ~main_query_1(sK10) [cnf transformation 188]
285. 0 <=> ! [X3,X5,X2,X4,X6] : ~from_2(X2,X3,X4,X5,X6) [avatar definition]
286. ~from_2(X2,X3,X4,X5,X6) <- {0} [avatar component clause 285]
288. 2 <=> ! [X1,X0] : (select_1(X1) | ~where_3(X1,X0)) [avatar definition]
289. ~where_3(X1,X0) | select_1(X1) <- {2} [avatar component clause 288]
290. 0 | 2 [avatar split clause 195,288,285]
292. 4 <=> ! [X1,X3,X5,X2,X4] : ~from_5(X1,X2,X3,X4,X5) [avatar definition]
293. ~from_5(X1,X2,X3,X4,X5) <- {4} [avatar component clause 292]
295. 6 <=> ! [X0] : (select_4(X0) | ~where_6(X0)) [avatar definition]
296. ~where_6(X0) | select_4(X0) <- {6} [avatar component clause 295]
297. 4 | 6 [avatar split clause 204,295,292]
299. 8 <=> main_query_1(sK10) [avatar definition]
300. main_query_1(sK10) <- {8} [avatar component clause 299]
302. 9 <=> ~main_query_1(sK10) [avatar definition]
303. ~main_query_1(sK10) <- {9} [avatar component clause 302]
305. 10 <=> main_query_2(sK10) [avatar definition]
306. main_query_2(sK10) <- {10} [avatar component clause 305]
308. 11 <=> ~main_query_2(sK10) [avatar definition]
309. ~main_query_2(sK10) <- {11} [avatar component clause 308]
310. ~9 | ~11 [avatar split clause 247,308,302]
311. 8 | 10 [avatar split clause 246,305,299]
312. select_1(sK10) <- {8} [resolution 196,300]
316. eq(X0,X1) | ~select_1(X0) [resolution 191,194]
317. ~select_1(X0) | where_6(X0) [resolution 316,201]
318. where_3(X0,X0) | ~where_6(X0) [resolution 192,200]
323. ~where_6(X1) | select_1(X1) <- {2} [resolution 318,289]
326. tab1(sK5(X0),sK6(X0),sK7(X0),sK8(X0),sK9(X0)) | ~select_4(X0) [resolution 202,198]
328. tab1(sK0(X1,X0),sK1(X1,X0),sK2(X1,X0),sK3(X1,X0),sK4(X1,X0)) | ~select_1(X0) [resolution 193,189]
330. from_2(sK5(X1),sK6(X1),sK7(X1),sK8(X1),sK9(X1)) | ~select_4(X1) [resolution 326,190]
331. from_5(sK0(X1,X0),sK1(X1,X0),sK2(X1,X0),sK3(X1,X0),sK4(X1,X0)) | ~select_1(X0) [resolution 328,199]
333. ~select_1(X0) <- {4} [subsumption resolution 331,293]
334. $false <- {4, 8} [resolution 333,312]
335. ~4 | ~8 [avatar contradiction clause 334]
339. select_1(sK10) <- {8} [resolution 300,196]
340. where_6(sK10) <- {8} [resolution 339,317]
343. select_4(sK10) <- {6, 8} [resolution 340,296]
346. main_query_2(sK10) <- {6, 8} [resolution 343,206]
348. $false <- {6, 8, 11} [subsumption resolution 346,309]
349. ~6 | ~8 | 11 [avatar contradiction clause 348]
350. select_4(sK10) <- {10} [resolution 306,205]
352. where_6(sK10) <- {10} [resolution 350,203]
354. select_1(sK10) <- {2, 10} [resolution 352,323]
356. main_query_1(sK10) <- {2, 10} [resolution 354,197]
357. $false <- {2, 9, 10} [subsumption resolution 356,303]
358. ~2 | 9 | ~10 [avatar contradiction clause 357]
360. ~select_4(X1) <- {0} [resolution 286,330]
361. $false <- {0, 10} [resolution 360,350]
362. ~0 | ~10 [avatar contradiction clause 361]
363. $false [avatar sat refutation 290,297,310,311,335,349,358,362]
% SZS output end Proof for output_query08_query09
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.029 s
% ------------------------------
% ------------------------------
