% Refutation found. Thanks to Tanya!
% SZS status Theorem for test9
% SZS output start Proof for test9
14. ! [X7] : (main_query_0(X7) <=> main_query_1(X7)) [input]
15. ~! [X7] : (main_query_0(X7) <=> main_query_1(X7)) [negated conjecture 14]
16. ! [X7] : (main_query_0(X7) <=> ? [X8,X9,X10,X11,X12] : (X7 = X11 & X7 = X9 & $true & tab(X7,X11,X12) & tab(X8,X9,X10))) [input]
17. ! [X7] : (main_query_1(X7) <=> ? [X11,X12] : (X7 = X11 & $true & tab(X7,X11,X12))) [input]
25. ~! [X0] : (main_query_0(X0) <=> main_query_1(X0)) [rectify 15]
26. ! [X0] : (main_query_0(X0) <=> ? [X1,X2,X3,X4,X5] : (X0 = X4 & X0 = X2 & $true & tab(X0,X4,X5) & tab(X1,X2,X3))) [rectify 16]
27. ! [X0] : (main_query_0(X0) <=> ? [X1,X2,X3,X4,X5] : (X0 = X4 & X0 = X2 & tab(X0,X4,X5) & tab(X1,X2,X3))) [true and false elimination 26]
28. ! [X0] : (main_query_1(X0) <=> ? [X1,X2] : (X0 = X1 & $true & tab(X0,X1,X2))) [rectify 17]
29. ! [X0] : (main_query_1(X0) <=> ? [X1,X2] : (X0 = X1 & tab(X0,X1,X2))) [true and false elimination 28]
51. ? [X0] : (main_query_0(X0) <~> main_query_1(X0)) [ennf transformation 25]
52. ? [X0] : ((~main_query_1(X0) | ~main_query_0(X0)) & (main_query_1(X0) | main_query_0(X0))) [nnf transformation 51]
53. ? [X0] : ((~main_query_1(X0) | ~main_query_0(X0)) & (main_query_1(X0) | main_query_0(X0))) => ((~main_query_1(sK0) | ~main_query_0(sK0)) & (main_query_1(sK0) | main_query_0(sK0))) [choice axiom]
54. (~main_query_1(sK0) | ~main_query_0(sK0)) & (main_query_1(sK0) | main_query_0(sK0)) [skolemisation 52,53]
55. ! [X0] : ((main_query_0(X0) | ! [X1,X2,X3,X4,X5] : (X0 != X4 | X0 != X2 | ~tab(X0,X4,X5) | ~tab(X1,X2,X3))) & (? [X1,X2,X3,X4,X5] : (X0 = X4 & X0 = X2 & tab(X0,X4,X5) & tab(X1,X2,X3)) | ~main_query_0(X0))) [nnf transformation 27]
56. ! [X0] : ((main_query_0(X0) | ! [X1,X2,X3,X4,X5] : (X0 != X4 | X0 != X2 | ~tab(X0,X4,X5) | ~tab(X1,X2,X3))) & (? [X6,X7,X8,X9,X10] : (X0 = X9 & X0 = X7 & tab(X0,X9,X10) & tab(X6,X7,X8)) | ~main_query_0(X0))) [rectify 55]
57. ! [X0] : (? [X6,X7,X8,X9,X10] : (X0 = X9 & X0 = X7 & tab(X0,X9,X10) & tab(X6,X7,X8)) => (sK4(X0) = X0 & sK2(X0) = X0 & tab(X0,sK4(X0),sK5(X0)) & tab(sK1(X0),sK2(X0),sK3(X0)))) [choice axiom]
58. ! [X0] : ((main_query_0(X0) | ! [X1,X2,X3,X4,X5] : (X0 != X4 | X0 != X2 | ~tab(X0,X4,X5) | ~tab(X1,X2,X3))) & ((sK4(X0) = X0 & sK2(X0) = X0 & tab(X0,sK4(X0),sK5(X0)) & tab(sK1(X0),sK2(X0),sK3(X0))) | ~main_query_0(X0))) [skolemisation 56,57]
59. ! [X0] : ((main_query_1(X0) | ! [X1,X2] : (X0 != X1 | ~tab(X0,X1,X2))) & (? [X1,X2] : (X0 = X1 & tab(X0,X1,X2)) | ~main_query_1(X0))) [nnf transformation 29]
60. ! [X0] : ((main_query_1(X0) | ! [X1,X2] : (X0 != X1 | ~tab(X0,X1,X2))) & (? [X3,X4] : (X0 = X3 & tab(X0,X3,X4)) | ~main_query_1(X0))) [rectify 59]
61. ! [X0] : (? [X3,X4] : (X0 = X3 & tab(X0,X3,X4)) => (sK6(X0) = X0 & tab(X0,sK6(X0),sK7(X0)))) [choice axiom]
62. ! [X0] : ((main_query_1(X0) | ! [X1,X2] : (X0 != X1 | ~tab(X0,X1,X2))) & ((sK6(X0) = X0 & tab(X0,sK6(X0),sK7(X0))) | ~main_query_1(X0))) [skolemisation 60,61]
76. main_query_1(sK0) | main_query_0(sK0) [cnf transformation 54]
77. ~main_query_1(sK0) | ~main_query_0(sK0) [cnf transformation 54]
79. tab(X0,sK4(X0),sK5(X0)) | ~main_query_0(X0) [cnf transformation 58]
81. ~main_query_0(X0) | sK4(X0) = X0 [cnf transformation 58]
82. main_query_0(X0) | X0 != X4 | X0 != X2 | ~tab(X0,X4,X5) | ~tab(X1,X2,X3) [cnf transformation 58]
83. tab(X0,sK6(X0),sK7(X0)) | ~main_query_1(X0) [cnf transformation 62]
84. ~main_query_1(X0) | sK6(X0) = X0 [cnf transformation 62]
85. main_query_1(X0) | X0 != X1 | ~tab(X0,X1,X2) [cnf transformation 62]
98. main_query_0(X4) | X2 != X4 | ~tab(X4,X4,X5) | ~tab(X1,X2,X3) [equality resolution 82]
99. ~tab(X4,X4,X5) | main_query_0(X4) | ~tab(X1,X4,X3) [equality resolution 98]
100. ~tab(X1,X1,X2) | main_query_1(X1) [equality resolution 85]
102. 0 <=> main_query_0(sK0) [avatar definition]
103. main_query_0(sK0) <- {0} [avatar component clause 102]
105. 1 <=> ~main_query_0(sK0) [avatar definition]
106. ~main_query_0(sK0) <- {1} [avatar component clause 105]
108. 2 <=> main_query_1(sK0) [avatar definition]
109. main_query_1(sK0) <- {2} [avatar component clause 108]
111. 3 <=> ~main_query_1(sK0) [avatar definition]
112. ~main_query_1(sK0) <- {3} [avatar component clause 111]
113. ~1 | ~3 [avatar split clause 77,111,105]
114. 0 | 2 [avatar split clause 76,108,102]
115. sK0 = sK6(sK0) <- {2} [resolution 84,109]
116. tab(sK0,sK0,sK7(sK0)) | ~main_query_1(sK0) <- {2} [superposition 83,115]
117. tab(sK0,sK0,sK7(sK0)) <- {2} [subsumption resolution 116,109]
121. main_query_0(sK0) | ~tab(X0,sK0,X1) <- {2} [resolution 99,117]
122. ~tab(X0,sK0,X1) <- {1, 2} [subsumption resolution 121,106]
123. $false <- {1, 2} [resolution 122,117]
124. 1 | ~2 [avatar contradiction clause 123]
125. sK0 = sK4(sK0) <- {0} [resolution 103,81]
127. tab(sK0,sK0,sK5(sK0)) | ~main_query_0(sK0) <- {0} [superposition 79,125]
128. tab(sK0,sK0,sK5(sK0)) <- {0} [subsumption resolution 127,103]
132. main_query_1(sK0) <- {0} [resolution 128,100]
133. $false <- {0, 3} [subsumption resolution 132,112]
134. ~0 | 3 [avatar contradiction clause 133]
135. $false [avatar sat refutation 113,114,124,134]
% SZS output end Proof for test9
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.057 s
% ------------------------------
% ------------------------------
